<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating Actual Length of an object Using a Linear Scale</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A scale drawing of a bridge shows a length of \({@drawing_length@}\) cm. </p>
<p>If 1 cm on the drawing represents \({@scale@}\) meters, what is the actual length of the bridge?</p>
<p>[[input:ans1]] [[validation:ans1]]m[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To find the actual length of the bridge we use the formula:</p>
<p><em>Actual Length = Scale factor × Drawing length</em></p>
<p>Given;</p>
<ul>
     <li>Scale factor =1&nbsp;<strong>:\({@scale@}\)</strong></li>
<p>This means that every 1 cm on the drawing equals \({@scale@}\) meters in real life.</p>
     <li>Measured drawing length of the bridge  = <strong>\({@drawing_length@}\) cm</strong> </li>
</ul>
<p>Substitute the values into the formula:</p>
<p>Actual Length = \({@scale@}\) × \({@drawing_length@}\) = \({@actual_length@}\)</p>
<p>Therefore, the actual length of the bridge is \({@actual_length@}\) meters.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>scale: rand(91) + 10;/*Scale factor of the bridge*/
drawing_length: rand(20) + 1;/*scale or drawing length of the bridge*/
actual_length: scale * drawing_length;/*Actual length of the bridge*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@scale@}\)</p>
<p>\({@drawing_length@}\)</p>
<p>\({@actual_length@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Purity Ekadeli</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the learner to use the given scale to calculate the actual length of the bridge given the drawing length. </p]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>actual_length</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*mistake1*/
actual_length_scale_mistake1: scale/drawing_length;/*Actual length of the bridge*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>actual_length</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the wrong answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>actual_length_scale_mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You mistakenly divided the scale by the drawing length instead of multiplying the drawing length by the scale to find the actual length of the bridge.</p>
<p>The actual length is calculated by multiplying the scale by the drawing length.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
