<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Conversion of ratio scales to statement form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The ratio scale of a map is given as \({@scale@}: {@scale:ratio@}\). Convert this ratio scale to statement form</p>
<p> [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>1 cm on the map represents \({@scale_ratio@}\)in reality.</p>
<p>Since 100 cm = 1 m, and 1,000 m = 1 km</p>
<p>\({@ scale_ratio@}\) = \({@ratio_scale@} \div 100000\) </p>
<p>= \({@statement_scale@}\)</p>
<p>Thus, the statement form is:
\({@scale@}\) cm represents \({@statement_scale@}\)km.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>scale:1;/*Ratio 1 for the scales*/
scale_ratio:rand_with_step(25000,100000,1000);/*ration of the range ranges from 25000 to 100000*/
statement_scale:float(scale_ratio/100000);/*statement ratio scale*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>1 cm on the map represents \({@scale_ratio@}\)in reality.</p>
<p>Since 100 cm = 1 m, and 1,000 m = 1 km</p>
<p>\({@ scale_ratio@}\) = \({@ratio_scale@} \div 100000\) </p>
<p>= \({@statement_scale@}\)</p>
<p>Thus, the statement form is:
\({@scale@}\) cm represents \({@statement_scale@}\)km.</p>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Conversion of ratio scale to statement form.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
