<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Interactive: Sketching Quadratic Graphs y=ax²+bx+c</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Move the three points in the grid below to sketch the parabola with equation \(y = {@a@}x^2 {@S1(b)@}x {@S1(c)@}\).</p>

<p><b>Note:</b> Make sure you move all three points, otherwise your answer won't be assessed.</p>

<center>
[[jsxgraph input-ref-ans1="sR1" input-ref-ans2="sR2" input-ref-ans3="sR3"]]
JXG.Options.text.useMathJax = true;
var board = JXG.JSXGraph.initBoard(divid, {
  axis: true,
  showCopyright: false,
  boundingbox: [-5,10,5,-2],
  keepaspectratio: true,
  grid:true
});

var p1 = board.create('point',[{#x1#},{#y1#}],{
  name:'',
  snapToGrid:true,
  size:3,
  snapSizeX:0.5,
  snapSizeY:0.5,
  showInfobox:false
});

var p2 = board.create('point',[{#x2#},{#y2#}],{
  name:'',
  snapToGrid:true,
  size:3,
  snapSizeX:0.5,
  snapSizeY:0.5,
  showInfobox:false
});

var p3 = board.create('point',[{#x3#},{#y3#}],{
  name:'',
  snapToGrid:true,
  size:3,
  snapSizeX:0.5,
  snapSizeY:0.5,
  showInfobox:false
});

var curve = board.create('curve', [function(t){return t;}, function(t){
  var x1=p1.X(), y1=p1.Y(), x2=p2.X(), y2=p2.Y(), x3=p3.X(), y3=p3.Y();
  var denom = (x1-x2)*(x1-x3)*(x2-x3);
  var a = (x3*(y2-y1)+x2*(y1-y3)+x1*(y3-y2))/denom;
  var b = (x3*x3*(y1-y2)+x2*x2*(y3-y1)+x1*x1*(y2-y3))/denom;
  var c = (x2*x3*(x2-x3)*y1+x3*x1*(x3-x1)*y2+x1*x2*(x1-x2)*y3)/denom;
  return a*t*t + b*t + c;
}, -5, 5], {strokeWidth:2});

stack_jxg.bind_point(sR1,p1);
stack_jxg.bind_point(sR2,p2);
stack_jxg.bind_point(sR3,p3);

board.unsuspendUpdate();
[[/jsxgraph]]
</center>

<p style="visibility:hidden">[[input:ans1]] [[validation:ans1]]</p>
<p style="visibility:hidden">[[input:ans2]] [[validation:ans2]]</p>
<p style="visibility:hidden">[[input:ans3]] [[validation:ans3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Generate random quadratic coefficients */
a: rand(3)+1;
b: (rand(5)-2)*2;
c: rand(5)+1;

/* Correct points - vertex and two symmetric points */
h: -b/(2*a);
k: a*h^2 + b*h + c;
x1: h;
y1: k;
x2: h+1;
y2: a*x2^2 + b*x2 + c;
x3: h-1;
y3: a*x3^2 + b*x3 + c;

/* Starting points (not correct) */
sx1: rand_with_prohib(-3,3,[x1,x2,x3]);
sy1: rand_with_prohib(-1,9,[y1,y2,y3]);
sx2: rand_with_prohib(-3,3,[sx1,x1,x2,x3]);
sy2: rand_with_prohib(-1,9,[sy1,y1,y2,y3]);
sx3: rand_with_prohib(-3,3,[sx1,sx2,x1,x2,x3]);
sy3: rand_with_prohib(-1,9,[sy1,sy2,y1,y2,y3]);

/* String formatting functions */
S(n):= if (n=1) then "+" elseif (n=-1) then "-" elseif (n>=0) then sconcat("+", string(n)) else sconcat("-", string(abs(n)));
S1(n):= if (n>=0) then sconcat("+", string(n)) else sconcat("-", string(abs(n)));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>y = {@a@}x^2 {@S1(b)@}x {@S1(c)@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>[x1,y1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>[x2,y2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>[x3,y3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Calculate the quadratic equation from student points */
x1: ans1[1]; y1: ans1[2];
x2: ans2[1]; y2: ans2[2];
x3: ans3[1]; y3: ans3[2];

denom: (x1-x2)*(x1-x3)*(x2-x3);
sans_a: (x3*(y2-y1)+x2*(y1-y3)+x1*(y3-y2))/denom;
sans_b: (x3*x3*(y1-y2)+x2*x2*(y3-y1)+x1*x1*(y2-y3))/denom;
sans_c: (x2*x3*(x2-x3)*y1+x3*x1*(x3-x1)*y2+x1*x2*(x1-x2)*y3)/denom;

sans_equation: y = sans_a*x^2 + sans_b*x + sans_c;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sans_equation</sans>
        <tans>y=a*x^2+b*x+c</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>sans_a</sans>
        <tans>a</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The coefficient of x² in your equation is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The coefficient of x² in your equation is incorrect.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>sans_b</sans>
        <tans>b</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The coefficient of x in your equation is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The coefficient of x in your equation is incorrect.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>sans_c</sans>
        <tans>c</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The constant term in your equation is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The constant term in your equation is incorrect.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>[x1,y1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>[x2,y2]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>[x3,y3]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
