<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Question on JSXGraph</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the perimeter of the polygon below. [[input:ans1]] [[validation:ans1]]</p>
[[jsxgraph]]
 var board = JXG.JSXGraph.initBoard('jxgbox', {boundingbox: [-5, 10, 10, -5],
                keepAspectRatio:true, showCopyright:false, showNavigation:false ,axis:false});

        const A = board.create('point', [0,6], { name: 'A', visible: false });
        const B = board.create('point', [4,6], { name: 'B', visible: false });
        const C = board.create('point', [4,3], { name: 'C', visible: false });
        const D = board.create('point', [8,3], { name: 'D', visible: false });
        const E = board.create('point', [8,0], { name: 'E', visible: false });
        const F = board.create('point', [0,0], { name: 'F', visible: false });

        board.create('polygon', [A, B, C, D,E,F], {borders: {
                                               strokeWidth: 1,
                                               strokeColor: 'blue',},
                                              fillColor: 'none',});

        board.create('line',[[0,5.7],[0.3,5.7]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        board.create('line',[[0.3,5.7],[0.3,6]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        
        board.create('line',[[3.7,5.7],[4,5.7]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        board.create('line',[[3.7,5.7],[3.7,6]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        
        board.create('line',[[4,3.3],[4.3,3.3]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        board.create('line',[[4.3,3.3],[4.3,3]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});

       board.create('line',[[7.7,2.7],[8,2.7]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        board.create('line',[[7.7,2.7],[7.7,3]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});

        board.create('line',[[7.7,0.3],[8,0.3]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        board.create('line',[[7.7,0.3],[7.7,0]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});

        board.create('line',[[0,0.3],[0.3,0.3]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});
        board.create('line',[[0.3,0.3],[0.3,0]],{straightFirst: false,
                                    straightLast: false,
                                    strokeWidth: 1,
                                    strokeColor:  'blue'});

        board.create('text', [1.5,6.5, '4 cm'], { fontSize: 8 });
        board.create('text', [4.2, 4.5, '3 cm'], { fontSize: 8});
        board.create('text', [6,3.5, '4 cm'], { fontSize: 8 });
        board.create('text', [8.2, 1.5, '3 cm'], { fontSize: 8});
        board.create('text', [4,-0.45, '8 cm'], { fontSize: 8 });
        board.create('text', [-1.5, 3, '6 cm'], { fontSize: 8});
[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>1</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text>1</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
