<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the discount amount on a purchase</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The marked price of a \(\text{@item_bought@}\) is sh \({@marked_price@}\). \(\text{@pple_names@}\) bought the \(\text{@item_bought@}\) for sh  \({@selling_price@}\).
How much was the discount? sh [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Marked price \(= \text{sh } {@marked_price@}\)</p>
<p>Selling price \(= \text{sh } {@selling_price@}\)</p>
<p>Discount \(=\) Marked price \(-\) Selling price</p>
<p>\(\quad = \text{sh } {@marked_price@} - \text{sh } {@selling_price@}\)</p>
<p>\( \qquad \qquad = \text{sh } {@discount_amount@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[selling_price:rand_with_step(5000,8000,10);
discount_amount:rand_with_step(100,500,20);
marked_price:selling_price+discount_amount;
pple_names:rand(["Murunga","Rodgers","Mike","Shillah","Emmaculate","Mary","John","Erick","Joseph","Purity","Juliet","Evans","Caren","Dennis"]);
item_bought:rand(["bicycle","Unicycle","Skateboard"]);
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Discount \({@discount_amount@}\)</p>
<p>Marked price \({@marked_price@}\)</p>
<p>Selling price \({@selling_price@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">https://docs.google.com/document/d/1qgLFImMM_WpGEo5jHdWW-u0UAHcFByvqPi3jkBf6DG8/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">This requires the student to first identify the marked price and the selling price of the item from the question. Then, to calculate the discount amount, the student needs to subtract the selling price from the marked price.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>discount_amount</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correctly calculated discount amount for the purchased item.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>discount_amount</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
