<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Construction of a triangle given two sides and an angle.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Construct a triangle \(PQR\) with sides \(PQ = {@PQ@} cm,\, PR = {@PR@} cm\) and \(\angle QPR = {@angle_QPR@}^\circ.\) What is the measure of the side \(QR\)? (Write your answer correct to 1 d.p)</p>
<p>\(QR\) = [[input:ans1]] \(cm\)[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox',{
                boundingbox:[-5,10,10,-1], 
                axis: false,
                showCopyright:false,keepAspectRatio:true
    
            });

            var p1 = board.create('point', [0,0] , {name:'P',fillColor:'red',strokeColor:'red', point: {visible: false}});
            var p2 = board.create('point', {#coordinate_pointQ#} , {name:'Q',fillColor:'red',strokeColor:'red', visible: true});
            var p3 = board.create('point', {#coordinate_pointR#} , {name:'R',fillColor:'red',strokeColor:'red', visible: true});

            var b1 = board.create('segment',['P','Q'],{name:''});
            var b2 = board.create('segment',['P','R'],{name:''});
            var b3 = board.create('segment',['R','Q'],{name:''});

            var angle1 = board.create('angle', [p2,p1,p3], {type:'sector', orthoType:'square', orthoSensitivity:2, radius:0.6, name:'{#angle_QPR#}&deg'});
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*randomize the lengths of the triangle*/
PQ:rand_with_step(4,9,.1);
PR:rand_with_step(5,9,.1);

/*angle QPR*/
angle_QPR:rand_with_step(30,125,1);

/*coordinates of the triangle */
dtr(d) := d*%pi/180; /* Convert degrees to radians */
angle_to_point(r, d) := [float(r*cos(dtr(d))), float(r*sin(dtr(d)))];
coordinate_pointQ: angle_to_point(PQ, 0);
coordinate_pointR: angle_to_point(PR, angle_QPR);

angle_QPR_radians:angle_QPR*%pi/180;

/*compute the side lengths of the triangle*/
side_PQ:PQ;
side_PR:PR;
side_QR:sqrt(PQ^2+PR^2-2*PQ*PR*cos(angle_QPR_radians));
ta:decimalplaces(side_QR,1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@PQ@}</p>
<p>{@PR@}</p>
<p>{@side_QR@}</p>
<p>{@ta@}</p>
<p>{@coordinate_pointQ@}</p>
<p>{@coordinate_pointR@}</p>
<p>{@angle_QPR@}</p>
<p>{@angle_QPR_radians@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct value of the length QR </description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>823962729</deployedseed>
    <deployedseed>1514731692</deployedseed>
    <deployedseed>315176210</deployedseed>
    <deployedseed>1359872323</deployedseed>
    <deployedseed>323267992</deployedseed>
    <deployedseed>1985382323</deployedseed>
    <deployedseed>150479622</deployedseed>
    <deployedseed>2002080693</deployedseed>
    <deployedseed>43036402</deployedseed>
    <deployedseed>101177452</deployedseed>
    <deployedseed>613067939</deployedseed>
    <deployedseed>601815172</deployedseed>
    <deployedseed>1228592688</deployedseed>
    <deployedseed>551061638</deployedseed>
    <deployedseed>304902060</deployedseed>
    <deployedseed>414332601</deployedseed>
    <deployedseed>1398623890</deployedseed>
    <deployedseed>867214982</deployedseed>
    <deployedseed>1385323532</deployedseed>
    <deployedseed>2131043910</deployedseed>
  </question>
</quiz>
