<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the LCM of the given numbers</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p> What is the least common number that can be divided by \({@least_numb_1@},{@least_numb_2@}\) and \({@least_numb_3@}\) without a remainder. </p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3> Worked solution</h3>
<p>To find the LCM of \({@least_numb_1@}, {@least_numb_2@}\), and \({@least_numb_3@}\), let's use prime factorization.</p>
<p>The prime factorization of the numbers are:</p>
<p>
<ul>
  <li>\({@least_numb_1@} = {@f1@} \)</li>
  <li>\({@least_numb_2@} = {@f2@} \)</li>
  <li>\({@least_numb_3@} = {@f3@} \)</li>
</ul>
</p>
<p>Next, take the highest power of each prime to compute the LCM:</p>
<p> Hence, the \(\text{LCM is} = {@least_multiple@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>least_numb_1: rand(30)+2;
least_numb_2: rand_with_prohib(4,10,[least_numb_1]);
least_numb_3: rand_with_prohib(10,20,[least_numb_1,least_numb_2]);
least_multiple: lcm(least_numb_1,least_numb_2,least_numb_3);
least_numb_1 : rand(30) + 2;
least_numb_2 : rand_with_prohib(4, 10, [least_numb_1]);
least_numb_3 : rand_with_prohib(10, 20, [least_numb_1, least_numb_2]);

least_multiple : lcm(least_numb_1, least_numb_2, least_numb_3);

f1 : string(factor(least_numb_1));
f2 : string(factor(least_numb_2));
f3 : string(factor(least_numb_3));


</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p> \({@least_multiple@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text>The question requires the learner to:
1) Write down the multiples of each number.
2) Identify the common multiples shared by both numbers.
3) The smallest common multiple is the LCM.</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>least_multiple</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>least_multiple</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
