<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying numbers divisible by 11 from a given list.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p> Which one of the following numbers are divisible by 11</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>A number is divisible by \(11\) if the difference between the sums of alternating digits in the number is zero or a multiple of \(11\).</p>
<p>
<ul>
<li>
<p>\( {@num1_11@} \) is divisible by \(11\) since the sum of alternating digits equals to </p>
<p>\( {@digit1_num1_11@} + {@digit3_num1_11@} = {@addition_of_digits_num1_11@}\)</p>
<p>\( {@digit2_num1_11@} + {@digit4_num1_11@} ={@addition_of_4_digits_num1_11@}\)</p>
<p> And the difference between the sums of alternating digits is :</p>
<p>\({@addition_of_digits_num1_11@} - {@addition_of_4_digits_num1_11@} = {@subtraction_num1_11@} \)</p>
</li>
</ul>
</p>

<p>
<ul>
<li>
<p>\( {@num2_11@} \) is divisible by \(11\)  since the sum of alternating digits equals to </p>
<p>\({@digit1_num2_11@}+{@digit3_num2_11@}+{@digit5_num2_11@} = {@addition_of_digits_num2_11@}\)</p>
<p>\({@digit2_num2_11@}+{@digit4_num2_11@}= {@addition_of_2_digits_num2_11@}\)</p>
<p> And the difference between the sums of alternating digits is :</p>
<p>\({@addition_of_digits_num2_11@} - {@addition_of_2_digits_num2_11@}= {@subtraction_of_digits_num2_11@}\)</p>
</li>
</ul>
</p>

<p>
<ul>
<li>
<p>\( {@num3_11@} \) is not divisible by \(11\) since the sum of alternating digits equals to </p>
<p>\({@digit1_num3_11@}+{@digit3_num3_11@} = {@addition_of_digits_num3_11 @}\) </p>
<p>\({@digit2_num3_11@}+{@digit4_num3_11@} = {@addition_of_2_digits_num3_11@}\)</p>
<p> And the difference between the sums of alternating digits is :</p>
<p>\({@addition_of_digits_num3_11@} - {@addition_of_2_digits_num3_11@} = {@subtraction_of_digits_num3_11@}\)</p>
</li>
</ul>
</p>

<p>
<ul>
<li>
<p>\( {@num4_11@} \) is not divisible by \(11\) since the sum of alternating digits equals to </p>
<p> \({@digit1_num4_11@}+{@digit3_num4_11@} = {@addition_of_digits_num4_11@} \)</p>
<p>\({@digit2_num4_11@}+{@digit4_num4_11@} ={@addition_of_2_digits_num4_11@} \)</p>
<p> And the difference between the sums of alternating digits is :</p>
<p>\({@addition_of_digits_num4_11@}-{@addition_of_2_digits_num4_11@} = {@subtraction_of_digits_num4_11@}\)</p>
</li>
</ul>
</p>

<p>
<ul>
<li>
<p>\( {@num5_11@} \) is not divisible by \(11\) since the sum of alternating digits equals to </p>
<p>\({@digit1_num5_11@}+{@digit3_num5_11@} = {@addition_of_digits_num5_11@}\)</p>
<p>\({@digit2_num5_11@}+{@digit4_num5_11@} = {@addition_of_2_digits_num4_11@}\) </p>
<p> And the difference between the sums of alternating digits is :</p>
<p>\({@addition_of_digits_num5_11@}-{@addition_of_2_digits_num4_11@} ={@ subtraction_of_digits_num4_11@}\)</p>
</li>
</ul>
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*Divisibility test for 11*/

/*numbers divisible by 11*/
num1_11:rand_with_step(3311,8000,11);
num2_11:rand_with_step(21780,99990,11);

/* Define a function to check if a number is NOT divisible by 11 */
is_not_divisible_by_11(n) := mod(n,11) # 0;
/* Generate a list of 20 random numbers between 100 and 1000 */
random_list : makelist(rand(4000) + 4000, i, 1, 40);
/* Filter the numbers that are NOT divisible by 11 */
filtered_list : sublist(random_list, lambda([n], is_not_divisible_by_11(n)));

/*numbers not divisible by nine*/
num3_11:filtered_list[1];
num4_11:filtered_list[2];
num5_11:filtered_list[3];


/* Extracting digits for num1_11*/
digit1_num1_11:floor(num1_11/1000);
digit2_num1_11:floor(mod(num1_11 / 100,10));
digit3_num1_11:floor(mod(num1_11 / 10,10));
digit4_num1_11:floor(mod(num1_11,10));
addition_of_digits_num1_11:digit1_num1_11+digit3_num1_11;
addition_of_4_digits_num1_11:digit2_num1_11+digit4_num1_11;
subtraction_num1_11:addition_of_digits_num1_11 -addition_of_4_digits_num1_11;


/* Extracting digits for num2_11*/
digit1_num2_11:floor(num2_11/10000);
digit2_num2_11:floor(mod(num2_11/1000,10));
digit3_num2_11:floor(mod(num2_11/ 100,10));
digit4_num2_11:floor(mod(num2_11/ 10,10));
digit5_num2_11:floor(mod(num2_11,10));
addition_of_digits_num2_11:digit1_num2_11+digit3_num2_11+digit5_num2_11;
addition_of_2_digits_num2_11:digit2_num2_11+digit4_num2_11;
subtraction_of_digits_num2_11:addition_of_digits_num2_11-addition_of_2_digits_num2_11;

/* Extracting digits for num3_11*/
digit1_num3_11:floor(num3_11/1000);
digit2_num3_11:floor(mod(num3_11/ 100,10));
digit3_num3_11:floor(mod(num3_11 / 10,10));
digit4_num3_11:floor(mod(num3_11,10));
addition_of_digits_num3_11:digit1_num3_11+digit3_num3_11;
addition_of_2_digits_num3_11:digit2_num3_11+digit4_num3_11;
subtraction_of_digits_num3_11:addition_of_digits_num3_11-addition_of_2_digits_num3_11;

/* Extracting digits for num4_11*/
digit1_num4_11:floor(num4_11/1000);
digit2_num4_11:floor(mod(num4_11 / 100,10));
digit3_num4_11:floor(mod(num4_11 / 10,10));
digit4_num4_11:floor(mod(num4_11,10));
addition_of_digits_num4_11:digit1_num4_11+digit3_num4_11;
addition_of_2_digits_num4_11:digit2_num4_11+digit4_num4_11;
subtraction_of_digits_num4_11:addition_of_digits_num4_11-addition_of_2_digits_num4_11;


/* Extracting digits for num5_11*/
digit1_num5_11:floor(num5_11/1000);
digit2_num5_11:floor(mod(num5_11 / 100,10));
digit3_num5_11:floor(mod(num5_11/ 10,10));
digit4_num5_11:floor(mod(num5_11,10));
addition_of_digits_num5_11:digit1_num5_11+digit3_num5_11;
addition_of_2_digits_num4_11:digit2_num5_11+digit4_num5_11;
subtraction_of_digits_num4_11:addition_of_digits_num5_11-addition_of_2_digits_num4_11;


/*correct answer*/
correct_answers:rand_selection([
[1, true, num1_11],[2,true,num2_11]], 2);

/*incorrect answers*/
incorrect_answers:rand_selection([
[3, false, num3_11],
[4, false, num4_11],
[5, false, num5_11]
], 3);

ta: random_permutation(append(correct_answers,incorrect_answers));
ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index); 



</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>number divisible by 11-\({@num1_11@}\)</p>
<p>number divisible by 11- \({@num2_11@}\)</p>
<p>number not divisible by 11-\({@num3_11@}\)</p>
<p>number not divisible by 11-\({@num4_11@}\)</p>
<p>number not divisible by 11-\({@num5_11@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Sheila Cherotich </p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the student to test for divisibility of 11 by first summing the alternate digits in the number. Then, find the difference between the sums of alternating digits the number will be divisible by 11 if the result is zero or multiple of 11.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1871844735</deployedseed>
    <deployedseed>1898496946</deployedseed>
    <deployedseed>223472099</deployedseed>
    <deployedseed>1547277794</deployedseed>
    <deployedseed>1064801468</deployedseed>
    <deployedseed>2072252822</deployedseed>
    <deployedseed>61651630</deployedseed>
    <deployedseed>326462583</deployedseed>
    <deployedseed>942785212</deployedseed>
    <deployedseed>627402186</deployedseed>
    <deployedseed>1808211043</deployedseed>
    <deployedseed>1990065320</deployedseed>
    <deployedseed>1594770343</deployedseed>
    <deployedseed>1345414518</deployedseed>
    <deployedseed>1106137563</deployedseed>
    <deployedseed>51102760</deployedseed>
    <deployedseed>595152882</deployedseed>
    <deployedseed>368628000</deployedseed>
    <deployedseed>352427423</deployedseed>
    <deployedseed>1862796982</deployedseed>
    <deployedseed>647825782</deployedseed>
    <deployedseed>1280532956</deployedseed>
    <deployedseed>402611191</deployedseed>
    <deployedseed>938911415</deployedseed>
    <deployedseed>345507179</deployedseed>
    <deployedseed>803461651</deployedseed>
    <deployedseed>2121496040</deployedseed>
    <deployedseed>1288354763</deployedseed>
    <deployedseed>1787243045</deployedseed>
    <deployedseed>2103052098</deployedseed>
    <deployedseed>566240976</deployedseed>
    <deployedseed>1388693199</deployedseed>
    <deployedseed>979066826</deployedseed>
    <deployedseed>1222520640</deployedseed>
    <deployedseed>1308068826</deployedseed>
    <deployedseed>1139403541</deployedseed>
    <deployedseed>2039303917</deployedseed>
    <deployedseed>498072750</deployedseed>
    <deployedseed>1957584603</deployedseed>
    <deployedseed>1629083178</deployedseed>
    <deployedseed>1770981147</deployedseed>
    <deployedseed>539948860</deployedseed>
    <deployedseed>279804692</deployedseed>
    <deployedseed>1639329403</deployedseed>
    <deployedseed>366597869</deployedseed>
    <deployedseed>207724718</deployedseed>
    <deployedseed>1526487440</deployedseed>
    <deployedseed>634927852</deployedseed>
    <deployedseed>444510618</deployedseed>
    <deployedseed>974282108</deployedseed>
    <deployedseed>1014266855</deployedseed>
    <deployedseed>638502751</deployedseed>
    <deployedseed>339689648</deployedseed>
    <deployedseed>1505088319</deployedseed>
    <deployedseed>1051389058</deployedseed>
    <deployedseed>714397953</deployedseed>
    <deployedseed>1826560917</deployedseed>
    <deployedseed>1969901348</deployedseed>
    <deployedseed>961833412</deployedseed>
    <deployedseed>1552617419</deployedseed>
    <deployedseed>32066161</deployedseed>
    <deployedseed>1765995546</deployedseed>
    <deployedseed>487600354</deployedseed>
    <deployedseed>1098027450</deployedseed>
    <deployedseed>423232155</deployedseed>
    <deployedseed>873377942</deployedseed>
    <deployedseed>1869327421</deployedseed>
    <deployedseed>1227299370</deployedseed>
    <deployedseed>1169221955</deployedseed>
    <deployedseed>2091456405</deployedseed>
    <deployedseed>276846904</deployedseed>
    <deployedseed>1086501035</deployedseed>
    <deployedseed>1937121999</deployedseed>
    <deployedseed>514729945</deployedseed>
    <deployedseed>492521327</deployedseed>
    <deployedseed>505902045</deployedseed>
    <deployedseed>1817384196</deployedseed>
    <deployedseed>19906642</deployedseed>
    <deployedseed>933094581</deployedseed>
    <deployedseed>405301175</deployedseed>
    <deployedseed>1963206021</deployedseed>
    <deployedseed>704409136</deployedseed>
    <deployedseed>1153017614</deployedseed>
    <deployedseed>1632913574</deployedseed>
    <deployedseed>511648350</deployedseed>
    <deployedseed>844649408</deployedseed>
    <deployedseed>705970959</deployedseed>
    <deployedseed>1648583934</deployedseed>
    <deployedseed>854065984</deployedseed>
    <deployedseed>488956988</deployedseed>
    <deployedseed>231325475</deployedseed>
    <deployedseed>667856126</deployedseed>
    <deployedseed>946493034</deployedseed>
    <deployedseed>751726288</deployedseed>
    <deployedseed>1683326474</deployedseed>
    <deployedseed>1874560243</deployedseed>
    <deployedseed>1375832356</deployedseed>
    <deployedseed>954243069</deployedseed>
    <deployedseed>920449723</deployedseed>
    <deployedseed>752864446</deployedseed>
    <deployedseed>1599875673</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
