<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying numbers which are divisible by both 5 and 9 from a given list.</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Which of the following numbers are divisible by both 5 and 9?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[[[ if test="member(1, ta_index)" ]]

<p>A number is divisible by 5 if the last digit digit is 0.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]

<p>A number is divisible by 5 if the last digit is 5.</p>

[[/ if ]]

[[ if test="member(3, ta_index)" ]]

<p>A number is divisible by 9 if the sum of the digits is divisible by 9.</p>

[[/ if ]]

[[ if test="member(1, ta_index)" ]]

<p>A number is divisible by both 5 and 9 if it meets the rule of both 5 and 9</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]

<p>The number must be a multiple of their least common multiple to be divisible by both 5 and 9</p>

[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>num_div_by_5_9_1:rand_with_step(4500, 6000, 45);/*numbers divisible by both 5 and 9*/
num_div_by_5_9_2:rand_with_step(45000, 95000, 45);/*numbers divisible by both 5 and 9*/
num_div_by_5_9_3:rand_with_step(135000, 230000, 45);/*numbers divisible by both 5 and 9*/
num_div_by_5_9_4:rand_with_step(225000, 350000, 45);/*numbers divisible by both 5 and 9*/
num_div_by_5_9_5:rand_with_step(855000, 1000000, 45);/*numbers divisible by both 5 and 9*/
num_not_div_by_5_9_1:rand_with_step(2800, 40000,14);/*numbers not divisible by both 5 and 9*/
num_not_div_by_5_9_2:rand_with_step(42000, 50000, 14);/*numbers not divisible by both 5 and 9*/
num_not_div_by_5_9_3:rand_with_step(56000, 90000, 14);/*numbers not divisible by both 5 and 9*/
num_not_div_by_5_9_4:rand_with_step(98000, 150000, 14);/*numbers not divisible by both 5 and 9*/
num_not_div_by_5_9_5:rand_with_step(210000, 1000000, 14);/*numbers not divisible by both 5 and 9*/
correct_answer:rand_selection([
[1, true, num_div_by_5_9_1],[2, true, num_div_by_5_9_2],
[3, true, num_div_by_5_9_3],[4, true, num_div_by_5_9_4],
[5, true, num_div_by_5_9_5]],1);

incorrect_answers:rand_selection([
[6, false, num_not_div_by_5_9_1],[7, false, num_not_div_by_5_9_2],
[8, false, num_not_div_by_5_9_3],[9, false, num_not_div_by_5_9_4],
[10, false, num_not_div_by_5_9_5]], 3);

ta: random_permutation(append(correct_answer,incorrect_answers));
ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@num_div_by_5_9_1@}\)</p>
<p>\({@num_div_by_5_9_3@}\)</p>
<p>\({@num_div_by_5_9_4@}\)</p>
<p>\({@num_not_div_by_5_9_2@}\)</p>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Purity Ekadeli</p>
<p id="concept">https://docs.google.com/document/d/1R9ZLYIdNJpq0EwxRIkBibNw1-OtkDNMk_3_NZTezbN8/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires one to identify which numbers are  both divisible by 5 and 9. The following rules has to be applied, a number is divisible by 5 if the last digit is 0 or 5. A number is divisible by 9 if the sum of its digits is divisible by 9 </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
