<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying numbers which are divisible by 4</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Which of the following numbers are divisible by \(4\)?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>A number is divisible by \(4\) if its digit in the ones and tens place value are both zero or the number formed by the digits in the ones and tens place value is divisible by \(4\).</p>
<p>
<ul>
<li>
<p>\({@num1_div_by_4@}\) is divisible by \(4\) since {@message1@}</p>
<p>\({@addition1@}\) {@divide1@} \({@num1@}\) {@equal@}  \({@division_of_digits_num1_div_by_4@}\)</p>
</li>
<li>
<p>\({@num2_div_by_4@}\) is divisible by \(4\) since {@message2@}</p>
<p>\({@addition2@}\) {@divide2@}  \({@num2@}\) {@equal2@} \({@division_of_digits_num2_div_by_4@}\)</p>
</li>
<li>
<p>\({@num1_not_div_by_4@}\) is not divisible by \(4\) since \({@addition_of_digits_num1_not_div_by_4@}\) is not divisible by \(4\)</p>
<p>\(\frac{{@addition_of_digits_num1_not_div_by_4@}}{4} = {@floor(addition_of_digits_num1_not_div_by_4/4)@}  \text{  remainder   } {@mod(addition_of_digits_num1_not_div_by_4,4)@} \)</p>
</li>
<li>
<p>\({@num2_not_div_by_4@}\) is not divisible by \(4\) since \({@addition_of_digits_num2_not_div_by_4@}\) is not divisible by \(4\).
</p>
<p>\(\frac{{@addition_of_digits_num2_not_div_by_4@}}{4} = {@floor(addition_of_digits_num2_not_div_by_4/4)@}  \text{  remainder   } {@mod(addition_of_digits_num2_not_div_by_4,4)@} \)</p>
</li>
<li>
<p>\({@num3_not_div_by_4@}\) is not divisible by \(4\) since \({@addition_of_digits_num3_not_div_by_4@}\) is not divisible by \(4\).
</p>
<p>\(\frac{{@addition_of_digits_num3_not_div_by_4@}}{4} = {@floor(addition_of_digits_num3_not_div_by_4/4)@}  \text{  remainder   } {@mod(addition_of_digits_num3_not_div_by_4,4)@}\)</p>
</li>
</ul>
</p>
<p></p>
<p></p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Divisibility test for 4*/

/*numbers divisible by four*/
num1_div_by_4:rand_with_step(4004,8000,4);
num2_div_by_4:rand_with_step(8004,9999,4);

/* Define a function to check if a number is NOT divisible by 4 */
is_not_divisible_by_4(n) := mod(n,4) # 0;
/* Generate a list of 20 random numbers between 4000 and 8000 */
random_list : makelist(rand(4000) + 4000, i, 1, 40);
/* Filter the numbers that are NOT divisible by 4 */
filtered_list : sublist(random_list, lambda([n], is_not_divisible_by_4(n)));

/*numbers not divisible by four*/
num1_not_div_by_4:filtered_list[1];
num2_not_div_by_4:filtered_list[2];
num3_not_div_by_4:filtered_list[3];


/* Extracting digits for num1_div_by_4*/
digit3_num1_div_by_4:floor(mod(num1_div_by_4 / 10,10));
digit4_num1_div_by_4:floor(mod(num1_div_by_4,10));
addition1:((digit3_num1_div_by_4*10)+digit4_num1_div_by_4);
/*simp:false;
f1:addition1/4;
simp:true;*/
division_of_digits_num1_div_by_4:((digit3_num1_div_by_4*10)+digit4_num1_div_by_4)/4;
answer_div:division_of_digits_num1_div_by_4;

if division_of_digits_num1_div_by_4=0 then (message1:"the digits in the ones and tens place value are both zero")
else(message1:"the number formed by the digits in the ones and tens place value is divisible by 4.");
if division_of_digits_num1_div_by_4=0 then (display_calculation:"",addition1:"",equal:"",f1:"",division_of_digits_num1_div_by_4:"",divide1:"",num1:"")else(display_calculation :"",addition1:((digit3_num1_div_by_4*10)+digit4_num1_div_by_4),equal:"\\(= \\)",divide1:"\\(\\div\\)",num1:4,simp:false,f1:addition1/4,division_of_digits_num1_div_by_4:answer_div);

/* Extracting digits for num2_div_by_4*/
digit3_num2_div_by_4:floor(mod(num2_div_by_4 / 10,10));
digit4_num2_div_by_4:floor(mod(num2_div_by_4,10));
addition2:((digit3_num2_div_by_4*10)+digit4_num2_div_by_4);
/*simp:false;
f2:addition2/4;
simp:true;*/
division_of_digits_num2_div_by_4:((digit3_num2_div_by_4*10)+digit4_num2_div_by_4)/4;
answer_div2:division_of_digits_num2_div_by_4;

if division_of_digits_num2_div_by_4=0 then (message2:"the digits in the ones and tens place value are both zero")
else(message2:"the number formed by the digits in the ones and tens place value is divisible by \\(4\\).");
if division_of_digits_num2_div_by_4=0 then (display_calculation:"",addition2:"",equal2:"",f2:"",division_of_digits_num2_div_by_4:"",divide2:"",num2:"")else(display_calculation :"",simp:true,addition2:((digit3_num2_div_by_4*10)+digit4_num2_div_by_4),equal2:"\\( =\\)",divide2:"\\(\\div\\)",num2:4,simp:false,f2:addition2/4,division_of_digits_num2_div_by_4:answer_div2);

simp:true;
/* Extracting digits for num1_not_div_by_4*/
digit3_num1_not_div_by_4:floor(mod(num1_not_div_by_4 / 10,10));
digit4_num1_not_div_by_4:floor(mod(num1_not_div_by_4,10));
addition_of_digits_num1_not_div_by_4:((digit3_num1_not_div_by_4*10)+digit4_num1_not_div_by_4);

/* Extracting digits for num2_not_div_by_4*/
digit3_num2_not_div_by_4:floor(mod(num2_not_div_by_4 / 10,10));
digit4_num2_not_div_by_4:floor(mod(num2_not_div_by_4,10));
addition_of_digits_num2_not_div_by_4:((digit3_num2_not_div_by_4*10)+digit4_num2_not_div_by_4);

/* Extracting digits for num3_not_div_by_4*/
digit3_num3_not_div_by_4:floor(mod(num3_not_div_by_4 / 10,10));
digit4_num3_not_div_by_4:floor(mod(num3_not_div_by_4,10));
addition_of_digits_num3_not_div_by_4:((digit3_num3_not_div_by_4*10)+digit4_num3_not_div_by_4);


/*correct answer*/
correct_answers:rand_selection([
[1, true, num1_div_by_4],[2,true,num2_div_by_4]], 2);

/*incorrect answers*/
incorrect_answers:rand_selection([
[3, false, num1_not_div_by_4],
[4, false, num2_not_div_by_4],
[5, false, num3_not_div_by_4]
], 3);

ta: random_permutation(append(correct_answers,incorrect_answers));
ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index); 


]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Number divisible by 4: \({@num1_div_by_4@}\)</p>
<p>Number divisible by 4: \({@num2_div_by_4@}\)</p>
<p>Number not divisible by 4:  \({@num1_not_div_by_4@}\)</p>
<p>Number not divisible by 4:  \({@num2_not_div_by_4@}\)</p>
<p>Number not divisible by 4:   \({@num3_not_div_by_4@}\)</p>
<p></p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the student to test if the digits in the ones and tens place value are both zero or the number formed by the digits in the ones and tens place value is divisible 4 then the given number is also divisible  by 4.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct numbers which are divisible by 4</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
