<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Travel graph  showing the journey made by two motorist A and B</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The graph below shows the journey made by two motorist, A and B </p>

[[jsxgraph width="900px" height="500px"]]

var board = JXG.JSXGraph.initBoard("jxgbox", {
    boundingbox: [-1, 900, 12, -100], // Ensuring full coverage
    axis: false,
    showCopyright: false,
    showNavigation: false,
    grid: false
});

// X-axis (Time) - Bold Black Line
board.create('line', [[0, 0], [10, 0]], { strokeColor: 'black', strokeWidth: 3, fixed: true });

// Y-axis (Distance) - Bold Black Line
board.create('line', [[0, 0], [0, 800]], { strokeColor: 'black', strokeWidth: 3, fixed: true });

// X-axis labels and grid lines (Time)
var timeLabels = ['7 AM', '8 AM', '9 AM', '10 AM', '11 AM', '12 PM', '1 PM', '2 PM', '3 PM', '4 PM', '5 PM'];
for (var i = 0; i < timeLabels.length; i++) {
    board.create('text', [i, -50, timeLabels[i]], { anchorX: 'middle', anchorY: 'top', fixed: true });
    board.create('line', [[i, 0], [i, 800]], { strokeColor: 'gray', strokeWidth: 1, fixed: true }); // Vertical grid lines
}

// Y-axis labels and grid lines (Distance in km)
for (var j = 0; j <= 800; j += 200) {
    board.create('text', [-0.5, j, j.toString() + " km"], { anchorX: 'right', anchorY: 'middle', fixed: true });
    board.create('line', [[0, j], [9, j]], { strokeColor: 'gray', strokeWidth: 1, fixed: true }); // Horizontal grid lines
}

// Motorist A (7 AM - 5 PM, 800 km)
var A_start = board.create('point', [0, 0], { name: "A Start", color: "red", size: 4, fixed: true });
var A_meet = board.create('point', [('{#random_time#}'), ('{#random_distance #}')], { name: "Meet", color: "green", size: 5, fixed: true });
var A_end = board.create('point', [10, 800], { name: "A End", color: "red", size: 4, fixed: true });

board.create('segment', [A_start, A_meet], { strokeWidth: 2, strokeColor: 'red' });
board.create('segment', [A_meet, A_end], { strokeWidth: 2, strokeColor: 'red' });

// Motorist B (8 AM - 4 PM, 800 km)
var B_start = board.create('point', [1, 0], { name: "B Start", color: "blue", size: 4, fixed: true });
var B_meet = board.create('point', [('{#random_time#}'), ('{#random_distance #}')], { name: "Meet", color: "green", size: 5, fixed: true }); // Intersection at 12 PM, 400 km
var B_end = board.create('point', [9, 800], { name: "B End", color: "blue", size: 4, fixed: true });

board.create('segment', [B_start, B_meet], { strokeWidth: 2, strokeColor: 'blue' });
board.create('segment', [B_meet, B_end], { strokeWidth: 2, strokeColor: 'blue' });

// Labels for points
board.create('text', [() => A_start.X() + 0.2, () => A_start.Y() + 30, "7 AM, 0 km"], { anchorX: 'left', anchorY: 'bottom' });
board.create('text', [() => A_meet.X() + 0.2, () => A_meet.Y() + 30, " ('{#random_distance #}')km"], { anchorX: 'left', anchorY: 'bottom' });
board.create('text', [() => A_end.X() + 0.2, () => A_end.Y() + 30, "5 PM, 800 km"], { anchorX: 'left', anchorY: 'bottom' });

board.create('text', [() => B_start.X() + 0.2, () => B_start.Y() + 30, "8 AM, 0 km"], { anchorX: 'left', anchorY: 'bottom' });
board.create('text', [() => B_meet.X() + 0.2, () => B_meet.Y() + 30, " ('{#random_distance #}')km"], { anchorX: 'left', anchorY: 'bottom' });
board.create('text', [() => B_end.X() + 0.2, () => B_end.Y() + 30, "4 PM, 800 km"], { anchorX: 'left', anchorY: 'bottom' });

[[/jsxgraph]]

<p> <br> </p>
<p> Use the graph to answer the following<br> </p>

<ol>
<li><p>How many hours had motorist B taken when motorist A started his journey? [[input:ans1]]Hours [[validation:ans1]][[feedback:prt1]]</p></li>
<li><p>what was the Speed of A after meeting B?[[input:ans2]]Kilometres/hour [[validation:ans2]][[feedback:prt2]]</p></li>
<li><p>what was the Speed of B after meeting A?[[input:ans6]]Kilometres/hour [[validation:ans6]][[feedback:prt6]]</p></li>
<li><p>What distance did A cover before meeting B? [[input:ans3]]Kilometres [[validation:ans3]][[feedback:prt3]]</p></li>
<li><p>What was the average speed of motorist A for the whole journey? [[input:ans4]] [[validation:ans4]]Kilometres/hour[[feedback:prt4]]</p></li>
<li><p>What was the average speed of motorist B for the whole journey? [[input:ans5]] [[validation:ans5]]Kilometres/hour[[feedback:prt5]]</p></li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<ol>
<li><p>The hours  that motorist B took when motorist A started his journey was \({@tans1@}\) hour</p></li>
<li><p>The Speed of A after meeting B was \({@tans2@}\) Kms/hr </p></li>
<li><p>The distance  A had covered before meeting B was \({@tans3@}\)Kms</p></li>
<li><p>The average speed of motorist A for the whole journey was  \({@tans4@}\)Kms/hr</p></li>
<li><p>The average speed of motorist B for the whole journey was  \({@tans5@}\)Kms/hr</p></li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Define fixed values */
start_A : 1;
end_A : 10;
start_B : 2;
end_B : 9;
total_distance : 800;

/* Randomize meeting point */
random_time : rand(2)+6; /* Meeting time between 10 AM and 2 PM */
random_distance : rand(200)+400; /* Meeting distance between 200 km and 600 km */


/* Calculate speeds */
speed_A : decimalplaces(total_distance / (end_A - start_A),2); /* Speed of A */
speed_B : decimalplaces(total_distance / (end_B - start_B),2); /* Speed of B */


/* Calculate speeds for A and B */
speed_A1 : random_distance / (random_time - start_A); /* Speed before meeting */
speed_A2 : (total_distance - random_distance) / (end_A - random_time); /* Speed after meeting */

speed_B1 : random_distance / (random_time - start_B); /* Speed before meeting */
speed_B2 : (total_distance - random_distance) / (end_B - random_time); /* Speed after meeting */

/* Distance covered by B before A started */
distance_B_at_A_start : speed_B * (start_A - start_B); /* Should be 0 */

/* Distance covered by A when B started */
distance_A_at_B_start : speed_A * (start_B - start_A);

/* Distance covered by A before meeting */
distance_A_before_meeting : random_distance;

/* "12 PM, 400 km", ('{#random_time#}') PM*/



tans1:1;
tans2:decimalplaces(speed_A2,2);
tans6: decimalplaces(speed_B2, 2) ;
tans3:distance_A_before_meeting;
tans4: speed_A;
tans5: speed_B ;





]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\({@tans1@}\)
\({@tans2@}\)
\({@tans3@}\)
\({@tans4@}\)
\({@tans5@}\)
\({@tans6@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Joseph Baya</p>
<p id="reviewer"></p>
<p id="description">This question evaluates the ability to interpret a distance-time graph representing the journeys of two motorists, A and B. It involves determining travel durations, calculating speeds before and after meeting, and computing average speeds for the entire journey. The focus is on understanding motion, speed, and time relationships using graphical data.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>tans1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>tans2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>tans3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>tans4</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>tans5</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>tans6</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>tans1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>tans2</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answers</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>tans3</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answers</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>tans4</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answers</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans5</sans>
        <tans>tans5</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answers</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans6</sans>
        <tans>tans6</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1413619640</deployedseed>
    <deployedseed>714323462</deployedseed>
    <deployedseed>1098306112</deployedseed>
    <deployedseed>432862846</deployedseed>
    <deployedseed>1784569649</deployedseed>
    <deployedseed>1086244354</deployedseed>
    <deployedseed>322028338</deployedseed>
    <deployedseed>165931580</deployedseed>
    <deployedseed>102914514</deployedseed>
    <deployedseed>669768440</deployedseed>
    <deployedseed>213009021</deployedseed>
    <deployedseed>273851200</deployedseed>
    <deployedseed>985492735</deployedseed>
    <deployedseed>1517435405</deployedseed>
    <deployedseed>135877505</deployedseed>
    <deployedseed>865260109</deployedseed>
    <deployedseed>75755760</deployedseed>
    <deployedseed>768957005</deployedseed>
    <deployedseed>542384695</deployedseed>
    <deployedseed>1606488494</deployedseed>
    <deployedseed>1821021728</deployedseed>
    <deployedseed>564017324</deployedseed>
    <deployedseed>1659361751</deployedseed>
    <deployedseed>1956242645</deployedseed>
    <deployedseed>933610234</deployedseed>
    <deployedseed>182881710</deployedseed>
    <deployedseed>67268406</deployedseed>
    <deployedseed>1241535050</deployedseed>
    <deployedseed>1219377957</deployedseed>
    <deployedseed>1563111956</deployedseed>
    <deployedseed>1201896312</deployedseed>
    <deployedseed>1132664154</deployedseed>
    <deployedseed>1616930276</deployedseed>
    <deployedseed>1546497995</deployedseed>
    <deployedseed>1251789796</deployedseed>
    <deployedseed>527726424</deployedseed>
    <deployedseed>136132765</deployedseed>
    <deployedseed>1113179741</deployedseed>
    <deployedseed>222568639</deployedseed>
    <deployedseed>1220385160</deployedseed>
    <deployedseed>746108113</deployedseed>
    <deployedseed>1122766351</deployedseed>
    <deployedseed>1677825747</deployedseed>
    <deployedseed>2033907658</deployedseed>
    <deployedseed>430234014</deployedseed>
    <deployedseed>101421310</deployedseed>
    <deployedseed>1414247286</deployedseed>
    <deployedseed>1828901262</deployedseed>
    <deployedseed>1868921596</deployedseed>
    <deployedseed>536990116</deployedseed>
    <deployedseed>667044628</deployedseed>
    <deployedseed>921719989</deployedseed>
    <deployedseed>1988286247</deployedseed>
    <deployedseed>1916412063</deployedseed>
    <deployedseed>246275437</deployedseed>
    <deployedseed>334650983</deployedseed>
    <deployedseed>514261745</deployedseed>
    <deployedseed>278203230</deployedseed>
    <deployedseed>1656440899</deployedseed>
    <deployedseed>115047165</deployedseed>
    <deployedseed>1079595452</deployedseed>
    <deployedseed>894373501</deployedseed>
    <deployedseed>109453914</deployedseed>
    <deployedseed>1744719391</deployedseed>
    <deployedseed>418376427</deployedseed>
    <deployedseed>317389634</deployedseed>
    <deployedseed>295035234</deployedseed>
    <deployedseed>1202648849</deployedseed>
    <deployedseed>1296807781</deployedseed>
    <deployedseed>328254352</deployedseed>
    <deployedseed>38149922</deployedseed>
    <deployedseed>1801454027</deployedseed>
    <deployedseed>702113826</deployedseed>
    <deployedseed>1642815031</deployedseed>
    <deployedseed>1249807502</deployedseed>
    <deployedseed>760339081</deployedseed>
    <deployedseed>1236345268</deployedseed>
    <deployedseed>2086678354</deployedseed>
    <deployedseed>1800450788</deployedseed>
    <deployedseed>635467888</deployedseed>
    <deployedseed>1331402687</deployedseed>
    <deployedseed>2053071193</deployedseed>
    <deployedseed>2046189115</deployedseed>
    <deployedseed>178198736</deployedseed>
    <deployedseed>1945761651</deployedseed>
    <deployedseed>1651642910</deployedseed>
    <deployedseed>1580349608</deployedseed>
    <deployedseed>2067138506</deployedseed>
    <deployedseed>532368462</deployedseed>
    <deployedseed>499058139</deployedseed>
    <deployedseed>1826809457</deployedseed>
    <deployedseed>259786335</deployedseed>
    <deployedseed>2006731909</deployedseed>
    <deployedseed>2014815258</deployedseed>
    <deployedseed>1640691799</deployedseed>
    <deployedseed>1694411114</deployedseed>
    <deployedseed>1777608684</deployedseed>
    <deployedseed>1632466850</deployedseed>
    <deployedseed>344387196</deployedseed>
    <deployedseed>1782193933</deployedseed>
    <deployedseed>1658746317</deployedseed>
    <deployedseed>1729729963</deployedseed>
    <deployedseed>180540812</deployedseed>
    <deployedseed>143082541</deployedseed>
    <deployedseed>1526056842</deployedseed>
    <deployedseed>1167178592</deployedseed>
    <deployedseed>1606253207</deployedseed>
    <deployedseed>1811235136</deployedseed>
  </question>
</quiz>
