<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the square of a given whole number</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Work out: \({@number@}^2\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To find the square of a whole number, we multiply the number by itself.</p>
<p>In this case, we need to find: \({@number@}^2\)</p>
<p>\( \begin{align} {@number@}^2  & = {@number@} \times {@number@}\\ &= {@squared_number@} \end{align}\)</p>
<p>Thus, the square of \({@number@}\) is \({@squared_number@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>number:rand(50)+4; /* random number */
squared_number:number^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Number: \({@number@}\)</p>
<p>Squared number: \({@squared_number@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">To find the square of a number, we multiply the given number by itself.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>squared_number</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct squared number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>squared_number</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1153556128</deployedseed>
    <deployedseed>555288386</deployedseed>
    <deployedseed>842855219</deployedseed>
    <deployedseed>1989826029</deployedseed>
    <deployedseed>190114826</deployedseed>
    <deployedseed>526329064</deployedseed>
    <deployedseed>1049509596</deployedseed>
    <deployedseed>1528745182</deployedseed>
    <deployedseed>844301301</deployedseed>
    <deployedseed>1216087708</deployedseed>
    <deployedseed>1291797524</deployedseed>
    <deployedseed>847703015</deployedseed>
    <deployedseed>634711848</deployedseed>
    <deployedseed>1264932965</deployedseed>
    <deployedseed>1896208643</deployedseed>
    <deployedseed>1322155264</deployedseed>
    <deployedseed>1719250132</deployedseed>
    <deployedseed>1715474517</deployedseed>
    <deployedseed>752351503</deployedseed>
    <deployedseed>674263620</deployedseed>
  </question>
</quiz>
