<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the product of two given decimal numbers</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Work out:  {@first_number_decimal1@} \(  \times \) {@second_number_decimal1@}   [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>

<p>Leave your answer to four decimal places</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To solve this, first convert each decimal number into a fraction.</p>
<p>\({@first_number_decimal1@} = \frac{{@first_number@}}{100}\) and \({@second_number_decimal1@} = \frac{{@second_number@}}{100}\)</p>
<p>Multiply the fractions by multiplying the numerators together and the denominators together</p>
<p>\( \frac{{@first_number@}}{100} \times \frac{{@second_number@}}{100} = \frac{{@first_number@} \times {@second_number@}}{100 \times 100} = \frac{{@product_of_first_second@}}{10000}\)</p>
<p>Finally, simplify the result and convert it back into a decimal.</p>
<p>\(\frac{{@product_of_first_second@}}{10000} = {@ta1@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>first_number:rand(1000)+1;
first_number_decimal:decimalplaces(float(first_number/100),2);
first_number_decimal1:dispdp(first_number_decimal,2);

second_number:rand_with_prohib(11,1000,[first_number]);
second_number_decimal:decimalplaces(float(second_number/100),2);
second_number_decimal1:dispdp(second_number_decimal,2);

product_of_first_second:first_number*second_number;
ta:decimalplaces(first_number_decimal*second_number_decimal,4);
ta1:dispdp(ta,4);


</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>First decimal number: \({@first_number_decimal@}\)</p>
<p>Second decimal number: \({@second_number_decimal@}\)</p>
<p></p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p  id="author">Daniel Murunga</p>
<p  id="concept"></p>
<p  id="reviewer"></p>
<p  id="reviewer"></p>
<p  id="description">To calculate the product of two decimal numbers, first convert them into fractions. Multiply the numerators together and the denominators together. Finally, simplify the resulting fraction and convert it back into a decimal to obtain the final answer.r </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct division of the two decimal numbers</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
