<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the volume of a cylinder.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the figure below. Note.\( Leave \,your \,answer \,in \,2 \,decimal \,places.\)</p>

[[jsxgraph]]

 var board = JXG.JSXGraph.initBoard('jxgbox', { // STACK: when moving to STACK change 'jxgboard' for divid
            boundingbox: [-5, 5, 5, -5],
            keepaspectratio: true,
            showCopyright: false,
            axis:false,
            showNavigation: true
        });

        // Defining the points of the elips

        var b1 = board.create('point', [-1, -1], {'visible':false, 'fixed':true});
        var b2 = board.create('point', [1, -1], {'visible':false, 'fixed':true});
        var b3 = board.create('point', [1.2, -1], {'visible':false, 'fixed':true});
        var b3o = board.create('point', [-1.2, -1], {'visible':false, 'fixed':true});
        var t1 = board.create('point', [-1, 1], {'visible':false, 'fixed':true});
        var t2 = board.create('point', [1, 1], {'visible':false, 'fixed':true});
        var t3 = board.create('point', [1.2, 1], {'visible':false, 'fixed':true});
        var t3o = board.create('point', [-1.2, 1], {'visible':false, 'fixed':true});

        // Drawing the elipse

        var ell_bottom = board.create('ellipse', [b1, b2, b3] , {fillColor: 'pink', borders: {strokeColor: 'green'}});
        var ell_top = board.create('ellipse', [t1, t2, t3] , {fillColor: 'white', borders: {strokeColor: 'black'}});

        // Drawing the heiht of the cylynder.
        var l1 = board.create('segment', [b3, t3]);
        var l2 = board.create('segment', [b3o, t3o]);

        // Defining the points of the the radius
        var A = board.create('point', [-1.2, 1], {name: 'A',visible:false});
        var B = board.create('point', [0, 1], {name: 'B',visible:false});

        // Drawing the radius of the cylinder.

        var line1 = board.create('segment', [A, B], { name: 'Line AB', strokeWidth: 2, 'fixed':true });

        // Labelling the height and the radius of the cylinder.

        var labelA = board.create('text', [-0.5, 1.2, '{#radius#}cm'], {fontSize: 12 ,'fixed':true});
        var labelC = board.create('text', [1.3,0, '{#height#}cm'], {fontSize: 12 ,'fixed':true});

 










[[/jsxgraph]]
<p>Volume \(=\) [[input:ans1]] [[validation:ans1]]\(cm^2\) [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>radius:rand(10)+4;
height:rand(15)+radius;
pi_1:3.142;
volume:float(pi_1*radius^2*height);
ta:decimalplaces(volume,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The height  \({@height@}\)</p>
<p>The volume   \({@volume@}\)</p>
<p> The volume in 2 decimal places \({@ta@}\)</p>
<p></p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct volume of the cylinder.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
