<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Average Monthly Temperature in Kisii</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p></p><p></p>

<p>The table below shows the average temperature (in degrees) for Kisii over six months.</p>

<table border="1" cellpadding="5">
<tr><th>Month</th><th>Rainfall (mm)</th></tr>
<tr><td>January</td><td>\({@Jan@}\)</td></tr>
<tr><td>February</td><td>\({@Feb@}\)</td></tr>
<tr><td>March</td><td>\({@Mar@}\)</td></tr>
<tr><td>April</td><td>\({@Apr@}\)</td></tr>
<tr><td>May</td><td>\({@May@}\)</td></tr>
<tr><td>June</td><td>\({@Jun@}\)</td></tr>
</table>

<p>Complete the bar graph by  By dragging the red points, complete the bar graph below to represent the average temperature (degrees) for the corresponding month.</p>

[[jsxgraph width="750px" height="450px" input-ref-ans1="sR1" input-ref-ans2="sR2" input-ref-ans3="sR3" input-ref-ans4="sR4" input-ref-ans5="sR5" input-ref-ans6="sR6"]]

var months = ["Jan", "Feb", "Mar", "Apr", "May", "Jun"];
var temps = [{#temp1#}, {#temp2#}, {#temp3#}, {#temp4#}, {#temp5#}, {#temp6#}];

JXG.Options.text.useMathJax = true;

var board = JXG.JSXGraph.initBoard(divid, {
  boundingBox: [-2, 40, 7, -5],
  showCopyright: false,
  showNavigation: false,
  keepAspectRatio: false,
  grid: true
});

// X-axis labels
for (var i = 0; i < months.length; i++) {
  board.create("text", [i + 0.5, -2, months[i]], {
    fixed: true,
    fontSize: 14,
    anchorX: "middle"
  });
}

// Y-axis
board.create("axis", [[0, 0], [0, 40]], {
  highlight: false,
  name: "\\(Temp\\ (^{\\circ}\\!C)\\)",
  withLabel: true,
  label: {
    fixed: true,
    highlight: false,
    fontSize: 15,
    position: "rt",
    offset: [-30, -20],
    anchorX: "right",
    useMathJax: true,
    parse: false
  },
  ticks: {
    label: {
      highlight: false,
      anchorX: "right",
      offset: [-10, 0]
    }
  }
});

// X-axis
board.create("axis", [[0, 0], [6, 0]], {
  highlight: false,
  withLabel: false,
  ticks: {
    drawZero: false,
    insertTicks: false,
    drawLabels: false
  }
});

// Gliders and bars
var segList = [];
var pointList = [];
for (var i = 0; i < months.length; i++) {
  let x = i + 0.5;

  segList.push(board.create("segment", [[x, 0], [x, 40]], {
    fixed: true,
    visible: false
  }));

  let p = board.create("glider", [x, temps[i], segList[i]], {
    showInfoBox: false,
    size: 4,
    snapToGrid: false,
    withLabel: false
  });

  pointList.push(p);

  board.create("polygon", [
    [x - 0.3, 0],
    [x + 0.3, 0],
    function () { return [x + 0.3, p.Y()]; },
    function () { return [x - 0.3, p.Y()]; }
  ], {
    fillColor: "#FF7F50",
    fillOpacity: 0.8,
    fixed: true,
    highlight: false,
    borders: {
      strokeColor: "#CC3300",
      strokeWidth: 2
    },
    vertices: { visible: false },
    layer: 10
  });
}

// Bind to STACK inputs
stack_jxg.bind_point(sR1, pointList[0]);
stack_jxg.bind_point(sR2, pointList[1]);
stack_jxg.bind_point(sR3, pointList[2]);
stack_jxg.bind_point(sR4, pointList[3]);
stack_jxg.bind_point(sR5, pointList[4]);
stack_jxg.bind_point(sR6, pointList[5]);


board.unsuspendUpdate();



[[/jsxgraph]]

<p></p><p hidden>[[input:ans1]][[input:ans2]][[input:ans3]] [[input:ans4]][[input:ans5]][[input:ans6]]</p>
<p hidden> [[validation:ans1]][[validation:ans2]][[validation:ans3]] [[validation:ans4]][[validation:ans5]][[validation:ans6]]</p>
<p >[[feedback:prt1]] </p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The correct bar graph is given as below</p>
<p></p>
[[jsxgraph width="750px" height="450px"]]

var months = ["Jan", "Feb", "Mar", "Apr", "May", "Jun"];
var temps = [{#ta1#}, {#ta2#}, {#ta3#}, {#ta4#}, {#ta5#}, {#ta6#}];

JXG.Options.text.useMathJax = true;

var board = JXG.JSXGraph.initBoard(divid, {
  boundingBox: [-2, 40, 7, -5],
  showCopyright: false,
  showNavigation: false,
  keepAspectRatio: false,
  grid: true
});

// X-axis labels
for (var i = 0; i < months.length; i++) {
  board.create("text", [i + 0.5, -2, months[i]], {
    fixed: true,
    fontSize: 14,
    anchorX: "middle"
  });
}

// Y-axis
board.create("axis", [[0, 0], [0, 40]], {
  highlight: false,
  name: "\\(Temp\\ (^{\\circ}\\!C)\\)",
  withLabel: true,
  label: {
    fixed: true,
    highlight: false,
    fontSize: 15,
    position: "rt",
    offset: [-30, -20],
    anchorX: "right",
    useMathJax: true,
    parse: false
  },
  ticks: {
    label: {
      highlight: false,
      anchorX: "right",
      offset: [-10, 0]
    }
  }
});

// X-axis
board.create("axis", [[0, 0], [6, 0]], {
  highlight: false,
  withLabel: false,
  ticks: {
    drawZero: false,
    insertTicks: false,
    drawLabels: false
  }
});

// Gliders and bars
var segList = [];
var pointList = [];
for (var i = 0; i < months.length; i++) {
  let x = i + 0.5;

  segList.push(board.create("segment", [[x, 0], [x, 40]], {
    fixed: true,
    visible: false
  }));

  let p = board.create("glider", [x, temps[i], segList[i]], {
    showInfoBox: false,
    size: 4,
    snapToGrid: false,
    withLabel: false,
   visible: false ,
    fixed: true // Make the glider fixed
  });

  pointList.push(p);

  board.create("polygon", [
    [x - 0.3, 0],
    [x + 0.3, 0],
    function () { return [x + 0.3, p.Y()]; },
    function () { return [x - 0.3, p.Y()]; }
  ], {
    fillColor: "#FF7F50",
    fillOpacity: 0.8,
    fixed: true,
    highlight: false,
    borders: {
      strokeColor: "#CC3300",
      strokeWidth: 2
    },
    vertices: { visible: false },
    layer: 10
  });
}




[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>Jan: rand_with_step(15, 30, 5);
Feb: rand_with_step(15, 30, 5);
Mar: rand_with_step(15, 30, 5);
Apr: rand_with_step(15, 30, 5);
May: rand_with_step(15, 30, 5);
Jun: rand_with_step(15, 35, 5);

ta1: (0.5, Jan);
ta2: (1.5, Feb);
ta3: (2.5, Mar);
ta4: (3.5, Apr);
ta5: (4.5, May);
ta6: (5.5, Jun);

temp1: 25;
temp2: 27;
temp3: 28;
temp4: 26;
temp5: 24;
temp6: 23;
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\({@ta1@}\)
\({@ta2@}\)
\({@ta3@}\)
\({@ta4@}\)
\({@ta5@}\)
\({@ta6@}\)

\({@Jan@}\)
\({@Feb@}\)\
\({@Mar@}\)
\({@Apr@}\)
\({@May@}\)
\({@Jun@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Joseph BAYA</p>
<p id="reviewer"></p>
<p id="reviewer"> </p>
<p id="description">This question displays a table of the average monthly temperatures (in degrees Celsius) for Kisii over a six-month period. Students are asked to represent this data by adjusting draggable points to build a corresponding bar graph.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Correct answer */
node1_test: if ans1=ta1 and ans2=ta2 and ans3=ta3 and ans4=ta4 and ans5=ta5 and ans6=ta6 then 1 else 0;

/* Node 2: Worked solution */

/* Node 3: ans1 correct */
/* Node 4: ans1 consistent with (a) */
node4_error: [float(ta1), ans1];

/* Node 5: ans2 correct */
/* Node 6: ans2 consistent with (a) */
node6_error: [float(ta2), ans2];

/* Node 7: ans3 correct */
/* Node 8: ans3 consistent with (a) */
node8_error: [float(ta3), ans3];

/* Node 9: ans4 correct */
/* Node 10: ans4 consistent with (a) */
node10_error: [float(ta4), ans4];

/* Node 11: ans5 correct */
/* Node 12: ans5 consistent with (a) */
node12_error: [float(ta5), ans5];

/* Node 13: ans6 correct */
/* Node 14: ans6 consistent with (a) */
node14_error: [float(ta6), ans6];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>node1_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>10</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>12</truenextnode>
        <trueanswernote>prt1-11-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your bar for May  is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>11</falsenextnode>
        <falseanswernote>prt1-11-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>11</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>node12_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>12</truenextnode>
        <trueanswernote>prt1-12-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your  bar for May  is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>12</falsenextnode>
        <falseanswernote>prt1-12-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your bar for May is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>12</name>
        <description><![CDATA[<p><br>Your bar  is <b>incorrect</b>.</p>]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-13-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your bar for June  is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>13</falsenextnode>
        <falseanswernote>prt1-13-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>13</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>node14_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-14-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your bar for June  is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-14-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your  bars are incorrect.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for correct dragged bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>node4_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your  bar for January  is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your bar  for January  is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your bar for February  is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>node6_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your  bar for February  is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your bar for February  is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your bar for March  is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>node8_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your  bar for March  is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>8</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your bar for March is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>10</truenextnode>
        <trueanswernote>prt1-9-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your bar for April  is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>9</falsenextnode>
        <falseanswernote>prt1-9-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>9</name>
        <description>Checks for correct bar </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>node10_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>10</truenextnode>
        <trueanswernote>prt1-10-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your  bar for April  is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>10</falsenextnode>
        <falseanswernote>prt1-10-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your bar for April  is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1974579155</deployedseed>
    <deployedseed>508604102</deployedseed>
    <deployedseed>1983856349</deployedseed>
    <deployedseed>165980899</deployedseed>
    <deployedseed>907562205</deployedseed>
    <deployedseed>1652298137</deployedseed>
    <deployedseed>365920942</deployedseed>
    <deployedseed>180370826</deployedseed>
    <deployedseed>481001274</deployedseed>
    <deployedseed>1186210085</deployedseed>
    <deployedseed>1802517334</deployedseed>
    <deployedseed>148996867</deployedseed>
    <deployedseed>758867210</deployedseed>
    <deployedseed>1190338687</deployedseed>
    <deployedseed>549054234</deployedseed>
    <deployedseed>284085175</deployedseed>
    <deployedseed>1225417936</deployedseed>
    <deployedseed>742340904</deployedseed>
    <deployedseed>1936027509</deployedseed>
    <deployedseed>1846374376</deployedseed>
    <deployedseed>541033605</deployedseed>
    <deployedseed>1222903966</deployedseed>
    <deployedseed>988602573</deployedseed>
    <deployedseed>1206069106</deployedseed>
    <deployedseed>17728918</deployedseed>
    <deployedseed>1007990442</deployedseed>
    <deployedseed>1682923761</deployedseed>
    <deployedseed>939398362</deployedseed>
    <deployedseed>1826433237</deployedseed>
    <deployedseed>413788847</deployedseed>
    <deployedseed>1824033121</deployedseed>
    <deployedseed>1056283199</deployedseed>
    <deployedseed>1482874616</deployedseed>
    <deployedseed>1473936708</deployedseed>
    <deployedseed>883418131</deployedseed>
    <deployedseed>659801985</deployedseed>
    <deployedseed>2042754557</deployedseed>
    <deployedseed>798369056</deployedseed>
    <deployedseed>1830285917</deployedseed>
    <deployedseed>840694145</deployedseed>
    <deployedseed>1602730404</deployedseed>
    <deployedseed>831440791</deployedseed>
    <deployedseed>891207558</deployedseed>
    <deployedseed>1235701714</deployedseed>
    <deployedseed>711745055</deployedseed>
    <deployedseed>570194780</deployedseed>
    <deployedseed>988390870</deployedseed>
    <deployedseed>1660620464</deployedseed>
    <deployedseed>113491778</deployedseed>
    <deployedseed>1046913115</deployedseed>
    <deployedseed>74181516</deployedseed>
    <deployedseed>762077244</deployedseed>
    <deployedseed>1541819414</deployedseed>
    <deployedseed>1417117553</deployedseed>
    <deployedseed>485244063</deployedseed>
    <deployedseed>1031828918</deployedseed>
    <deployedseed>369583794</deployedseed>
    <deployedseed>337733793</deployedseed>
    <deployedseed>121104284</deployedseed>
    <deployedseed>1501213142</deployedseed>
    <deployedseed>1820245130</deployedseed>
    <deployedseed>649126860</deployedseed>
    <deployedseed>135262664</deployedseed>
    <deployedseed>1804520632</deployedseed>
    <deployedseed>1761084847</deployedseed>
    <deployedseed>1725056918</deployedseed>
    <deployedseed>2070256651</deployedseed>
    <deployedseed>277147635</deployedseed>
    <deployedseed>309850091</deployedseed>
    <deployedseed>1854471042</deployedseed>
    <deployedseed>1518844114</deployedseed>
    <deployedseed>1610719769</deployedseed>
    <deployedseed>904981835</deployedseed>
    <deployedseed>749930365</deployedseed>
    <deployedseed>209349377</deployedseed>
    <deployedseed>248981166</deployedseed>
    <deployedseed>1094065140</deployedseed>
    <deployedseed>1090431874</deployedseed>
    <deployedseed>1998358988</deployedseed>
    <deployedseed>955061749</deployedseed>
    <deployedseed>1895544327</deployedseed>
    <deployedseed>1587391005</deployedseed>
    <deployedseed>1357125685</deployedseed>
    <deployedseed>303248524</deployedseed>
    <deployedseed>204923146</deployedseed>
    <deployedseed>470486428</deployedseed>
    <deployedseed>528468275</deployedseed>
    <deployedseed>774862482</deployedseed>
    <deployedseed>1015707068</deployedseed>
    <deployedseed>1914346279</deployedseed>
    <deployedseed>1703923500</deployedseed>
    <deployedseed>497490239</deployedseed>
    <deployedseed>1278092366</deployedseed>
    <deployedseed>1449291559</deployedseed>
    <deployedseed>455839382</deployedseed>
    <deployedseed>1192009165</deployedseed>
    <deployedseed>178574128</deployedseed>
    <deployedseed>2759268</deployedseed>
    <deployedseed>1858783936</deployedseed>
    <deployedseed>2054624481</deployedseed>
    <deployedseed>611054241</deployedseed>
    <deployedseed>1544976439</deployedseed>
    <deployedseed>1143065499</deployedseed>
    <deployedseed>1610104181</deployedseed>
    <deployedseed>921251444</deployedseed>
    <deployedseed>1904515309</deployedseed>
    <deployedseed>197346891</deployedseed>
    <deployedseed>1834528400</deployedseed>
    <deployedseed>329570637</deployedseed>
    <deployedseed>1255348202</deployedseed>
    <deployedseed>721194360</deployedseed>
    <deployedseed>54648173</deployedseed>
    <deployedseed>302222435</deployedseed>
    <deployedseed>1766228278</deployedseed>
    <deployedseed>219822401</deployedseed>
    <deployedseed>1275242575</deployedseed>
    <deployedseed>555672608</deployedseed>
    <deployedseed>918447601</deployedseed>
  </question>
</quiz>
