<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a circle</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Calculate the area of the circle below. Use \(\pi=\frac{22}{7}\).</p>
<p>Area \(= \) [[input:ans1]] \(\text{cm}^2\) [[validation:ans1]][[feedback:prt1]]</p>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox', {
        boundingbox: [-150, 150, 150, -150],
        axis: false, grid:false,  keepAspectRatio:true, showCopyright:false });

    var center = board.create('point', [0, 0], { name: '', fixed: true, size: 2, color: '#1E90FF'});
    var pointOnCircle = board.create('point',
                                                        [{#r#}, 0], 
                                                        { name: '', fixed: true,visible: false }
                                                         );
    var circle = board.create('circle', 
                                            [center, pointOnCircle],
                                            {strokeColor: '#1E90FF'});

    var radiusLine = board.create('line', [[-{#r#},0], pointOnCircle], {
        straightFirst: false,  straightLast: false,   strokeWidth: 2,
        strokeColor:  '#1E90FF'
    });
   board.create('text', [0, {#d#}*0.1, "{#d#} cm"], {
        fontSize: 14,   anchorX: 'middle',   anchorY: 'middle',fixed:true
    });
[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution </h3>
<p>To calculate the area of a circle, first square the radius to find \(r^2\). Next, use the formula \(\text{Area }=\pi r^2 \), and multiply the squared radius by an approximation of \(\pi \), such as \(\frac{22}{7}\) to get the area of the circle​</p>
<p>Area of circle \( = \pi r^2 \) where \(r\) is the radius of the circle</p>
<p>\(\qquad \quad = \frac{22}{7} \times r \times r \)</p>
<p>\( \qquad \quad = \frac{22}{7} \times {@r@}  \times {@r@} \)</p>
<p>\(\qquad \quad = {@ta@} \text{ cm}^2\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>r:rand_with_step(7,150,7);/*radius of the circle*/
area:22/7 * r^2;/* Calculate the area of the circle */
ta:area
d:r*2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Radius of a circle {@r@}</p>
<p>Area of a circle to decimal places{@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">https://docs.google.com/document/d/1CtIfEH4PfoxyhBWYunRH2HXp4fMCO2qmmfCySjEvepE/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">To calculate the area of a circle, first square the radius to find \(r^2\). Next, use the formula Area \(=\pi r^2 \), and multiply the squared radius by an approximation of \(\pi \), such as \(\frac{22}{7}\) to get the area of the circle​</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
