<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Application of Pythagorean relationship.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The base of a ladder is \({@base_triangle_values@} m\) from the base of a vertical wall. A painter places the ladder such that it touches the top of the wall at a point \({@height_triangle_values@} m\) above the ground.Calculate the length of the ladder.</p>


<p>The length of the ladder \(=\)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]] \(m\)</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution.</h3>
<p>The figure below shows the sketch of the ladder placed on a wall</p>
[[jsxgraph ]] 
  
//  Creating the board
        var board = JXG.JSXGraph.initBoard('jxgbox', {
            boundingbox: [-4, 6, 4, -3], 
            axis: false, showCopyright: false
        });
        // Defining the points of the triangle
        var A = board.create('point', [-3, -1], {name: 'A',visible:false});
        var C = board.create('point', [-3, 5], {name: 'B',visible:false});
        var B = board.create('point', [3, -1], {name: 'C',visible:false});
        var D = board.create('point', [-2.5, 5.5], {name: 'D',visible:false});
        var E = board.create('point', [3.5, -0.5], {name: 'E',visible:false});
        

        // Difining the points of the lader
        var a = board.create('point', [-3, 5], {name: 'a',visible:false});
        var b = board.create('point', [-2.5,5.5], {name: 'b',visible:false});
        var c = board.create('point', [-2.9, 4.9], {name: 'c',visible:false});
        var d = board.create('point', [-2.4, 5.4], {name: 'd',visible:false});



        var e = board.create('point', [-2,4], {name: 'e',visible:false});
        var f = board.create('point', [-1.5,4.5], {name: 'f',visible:false});
        var g = board.create('point', [-1.9,3.9], {name: 'g',visible:false});
        var h = board.create('point', [-1.4,4.4], {name: 'h',visible:false});


        var i = board.create('point', [-1,3], {name: 'i',visible:false});
        var j = board.create('point', [-0.5,3.5], {name: 'j',visible:false});
        var k = board.create('point', [-0.9,2.9], {name: 'k',visible:false});
        var m = board.create('point', [-0.4,3.4], {name: 'm',visible:false});


        var n = board.create('point', [0,2], {name: 'n',visible:false});
        var o = board.create('point', [0.5,2.5], {name: 'o',visible:false});
        var p = board.create('point', [0.1,1.9], {name: 'p',visible:false});
        var q = board.create('point', [0.6,2.4], {name: 'q',visible:false});

       
        var r = board.create('point', [1,1], {name: 'r',visible:false});
        var s = board.create('point', [1.5,1.5], {name: 's',visible:false});
        var t = board.create('point', [1.1,0.9], {name: 't',visible:false});
        var u = board.create('point', [1.6,1.4], {name: 'u',visible:false});


        var v = board.create('point', [2,0], {name: 'v',visible:false});
        var w = board.create('point', [2.5,0.5], {name: 'w',visible:false});
        var x = board.create('point', [2.1,-0.1], {name: 'x',visible:false});
        var y = board.create('point', [2.6,0.4], {name: 'y',visible:false});

        var H = board.create('point', [2.9,-0.9], {name: 'H',visible:false});
        var   Q = board.create('point', [3.4,-0.4], {name: 'Q',visible:false});
        var z = board.create('point', [3,-1], {name: 'z',visible:false});
        var G= board.create('point', [3.5,-0.5], {name: 'G',visible:false});
        






        // Drawing the triangle using the points
        var triangle = board.create('polygon', [A, B, C], {fillColor: 'white', borders: {strokeColor: 'blue'}});
        var Line = board.create('segment', [D,E], {fillColor: 'white', borders: {strokeColor: 'blue'}});

        //Drawing the ladder
        var line = board.create('segment', [a,b], {fillColor: 'white', borders: {strokeColor: 'black'}});
        var line = board.create('segment', [c,d], {fillColor: 'white', borders: {strokeColor: 'black'}});

        var line = board.create('segment', [e,f], {fillColor: 'white', borders: {strokeColor: 'black'}});
        var line = board.create('segment', [g,h], {fillColor: 'white', borders: {strokeColor: 'black'}});

        var line = board.create('segment', [i,j], {fillColor: 'white', borders: {strokeColor: 'black'}});
        var line = board.create('segment', [k,m], {fillColor: 'white', borders: {strokeColor: 'black'}});

        var line = board.create('segment', [n,o], {fillColor: 'white', borders: {strokeColor: 'black'}});
        var line = board.create('segment', [p,q], {fillColor: 'white', borders: {strokeColor: 'black'}});

        var line = board.create('segment', [r,s], {fillColor: 'white', borders: {strokeColor: 'black'}});
        var line = board.create('segment', [t,u], {fillColor: 'white', borders: {strokeColor: 'black'}});

        var line = board.create('segment', [v,w], {fillColor: 'white', borders: {strokeColor: 'black'}});
        var line = board.create('segment', [x,y], {fillColor: 'white', borders: {strokeColor: 'black'}});

        var line = board.create('segment', [z,G], {fillColor: 'white', borders: {strokeColor: 'black'}});
        var line = board.create('segment', [H,Q], {fillColor: 'white', borders: {strokeColor: 'black'}});

        



        // Labelling the sides of the triangle
        var labelA = board.create('text', [0,-1.5,'{#base_triangle_values#}m'], {fontSize: 12});
        var labelC = board.create('text', [-3.9,2.5, '{#height_triangle_values#}m'], {fontSize: 12});





[[/jsxgraph]]
<p>For this case, the vertical wall (height), the ground (base) and the ladder (hypotenuse) forms a right-angled triangle.</p>

<p>In a right-angled triangle the square of the hypotenuse is equal to the sum of the squares on the base and height.</p>

<p>Let \({@alphab_repres_base@} = \) base, \({@alphab_repres_height@} = \) height and \({@alphab_repress_hypotenuse@} = \) hypotenuse.</p>

<p>Now to calculate the hypotenuse we use Pythagorean relationships such that:</p>
<p>\({@hypotenuse_sqr@} = {@base_and_height_squares@}\)</p>

<p>For \({@alphab_repres_base@}= {@base_triangle_values@}m\) and \({@alphab_repres_height@} = {@height_triangle_values@}\),</p>

<p>Therefore,</p>
<p>\({@hypotenuse_sqr@} = ({@base_triangle_values@}^2 + {@height_triangle_values@}^2)m^2\)</p>

<p>By squaring the numbers we have:</p>
<p>\({@hypotenuse_sqr@} = ({@squares_base@} + {@squares_height@}) m^2\)</p>

<p>Adding the numbers in the bracket we get:</p>
<p>\({@hypotenuse_sqr@} = {@squares_base + squares_height@} m^2\)</p>

<p>To get the value of \({@alphab_repress_hypotenuse@}\) we get the square root so that we have:</p>
<p>\({@alphab_repress_hypotenuse@} = \sqrt({@squares_base + squares_height@} m^2)\)</p>

<p>Therefore, the length of the ladder is {@hypotenuse_sqrt@} \(m\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*Random letters representing bases, heights and hypotenuses*/
alphab_repres_base:rand([p,q,w,r,t]);   
alphab_repres_height:rand([y,h,s,v,u]); 
alphab_repress_hypotenuse:rand([w,x,d,l,f]);

 /*random numbers for bases and heights*/
/*Base and Height values*/
base_triangle_values:rand_with_step(2,40,1);  
height_triangle_values:rand_with_step(2,60,2);

/*Computation of hypotenuse*/
/*hypotenuse*/
hypotenuse:alphab_repress_hypotenuse^2;
/*Base and Height*/
sum_base_height_squared:(alphab_repres_base+alphab_repres_height)^2;
open_sum_base_height_squared:alphab_repres_base^2+2*alphab_repres_base*alphab_repres_height+alphab_repres_height^2;
base_and_height_squares:alphab_repres_base^2+alphab_repres_height^2;
hypotenuse_sqr:alphab_repress_hypotenuse^2;

/*squares root*/
base_and_height_squares_under_square_root:sqrt(base_and_height_squares);

/*Base and Height values squared*/
squares_base:(base_triangle_values)^2;
squares_height:(height_triangle_values)^2; 
 
hypotenuse_sqrt:decimalplaces(sqrt(squares_base + squares_height),2);

/*Mistake 1*/
/*Confusing (a + b)^2 with a^2 + b^2*/
sum_hypotenuse:(base_triangle_values+height_triangle_values);
hypotenuse_square_wrong:(sum_hypotenuse)^2;
hypotenuse_wrong_1: decimalplaces(sqrt(hypotenuse_square_wrong),2);

/*Mistake 2*/
/*Forgetting to square the side lengths*/
sum_hypotenuse_base_height:(base_triangle_values+height_triangle_values);
hypotenuse_wrong_2: decimalplaces(sqrt(sum_hypotenuse_base_height),2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The length of the ladder is \({@hypotenuse_sqrt@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Henry Onyango</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">The question provided is about application for pythagorean relationship. The learner is required to identify the length of the ladder that is the hypotenuse by interpretation.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>hypotenuse_sqrt</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct length of the of the ladder.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>hypotenuse_sqrt</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Confusing \((a + b)^2 \)with \(a^2 + b^2\).</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>hypotenuse_wrong_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>When calculating the hypotenuse, you must square the lengths of the base and perpendicular height before adding them together</p>
<p>\({@alphab_repress_hypotenuse@} = ({@alphab_repres_base@}^2 + {@alphab_repres_height@}^2 )\)</p>
<p>Then, get the square root;</p>
<p>\({@alphab_repress_hypotenuse@} = \)\({@base_and_height_squares_under_square_root@}\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>You have used an incorrect method of working the sum of squares.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>hypotenuse_wrong_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have used an incorrect method of working the sum of squares.</p>
<p>\({@sum_base_height_squared@}\) is equal to \({@open_sum_base_height_squared@}\), which is greater than \({@base_and_height_squares	@}\).</p>
<p>The correct method is;  \({@base_and_height_squares	@}\).</p>
<p>Then, get the square root;</p>
<p>\({@base_and_height_squares_under_square_root@}\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>186289259</deployedseed>
    <deployedseed>260150444</deployedseed>
    <deployedseed>1423777433</deployedseed>
    <deployedseed>601769240</deployedseed>
    <deployedseed>1962978279</deployedseed>
    <deployedseed>1161106383</deployedseed>
    <deployedseed>1240216347</deployedseed>
    <deployedseed>1855288290</deployedseed>
    <deployedseed>1655079981</deployedseed>
    <deployedseed>82281419</deployedseed>
    <deployedseed>1890843774</deployedseed>
    <deployedseed>815684518</deployedseed>
    <deployedseed>2033042538</deployedseed>
    <deployedseed>2024206855</deployedseed>
    <deployedseed>267714798</deployedseed>
    <deployedseed>263133569</deployedseed>
    <deployedseed>842631188</deployedseed>
    <deployedseed>1143177633</deployedseed>
    <deployedseed>471981819</deployedseed>
    <deployedseed>1103882525</deployedseed>
    <deployedseed>752265598</deployedseed>
    <deployedseed>1221783336</deployedseed>
    <deployedseed>597976687</deployedseed>
    <deployedseed>461512355</deployedseed>
    <deployedseed>793735244</deployedseed>
    <deployedseed>438072562</deployedseed>
    <deployedseed>2129960915</deployedseed>
    <deployedseed>1011034289</deployedseed>
    <deployedseed>1739068011</deployedseed>
    <deployedseed>971742419</deployedseed>
    <deployedseed>1744389415</deployedseed>
    <deployedseed>1165319990</deployedseed>
    <deployedseed>286957574</deployedseed>
    <deployedseed>440388212</deployedseed>
    <deployedseed>54871480</deployedseed>
    <deployedseed>1875651470</deployedseed>
    <deployedseed>1424420401</deployedseed>
    <deployedseed>1523361124</deployedseed>
    <deployedseed>1743067508</deployedseed>
    <deployedseed>1056270521</deployedseed>
    <deployedseed>1848411456</deployedseed>
    <deployedseed>227512209</deployedseed>
    <deployedseed>1251332382</deployedseed>
    <deployedseed>1140957603</deployedseed>
    <deployedseed>540506320</deployedseed>
    <deployedseed>2032512904</deployedseed>
    <deployedseed>1727216407</deployedseed>
    <deployedseed>1300458403</deployedseed>
    <deployedseed>572834325</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>hypotenuse_sqrt</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Confusing \((a + b)^2 \)with \(a^2 + b^2\).</description>
      <testinput>
        <name>ans1</name>
        <value>hypotenuse_wrong_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>You have used an incorrect method of working the sum of squares.</description>
      <testinput>
        <name>ans1</name>
        <value>hypotenuse_wrong_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
