<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating probability of ducks as a decimal</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Mary keeps chickens and ducks. Out of every \({@total_birds@}\) birds, \({@ducks@}\) are ducks.</p>
<p>What is the probability of picking a duck from his birds?</p>
<p>Express your answer as a decimal.</p>
<p>Probability of picking a duck = [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To calculate the probability of picking a duck use the formula:</p>
<p>Probability of picking a duck = Number of ducks รท Total number of birds</p>
<ul>
     <li>Total number of birds = \({@total_birds@}\)</li>
     <li>Number of ducks = \({@ducks@}\)</li>
</ul>
<p>P(ducks) = {@ ducks @} รท {@ total_birds@} = {@ p_duck @} </p>
<p>Therefore, the probability of picking a duck is {@ p_duck @}. </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>total_birds : rand(31) + 20;/*total number of birds*/
ducks : rand(10) + 10;/*number of ducks*/
p_duck : decimalplaces((ducks / total_birds ) , 2);/*Probability of ducks*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@total_birds@}\)</p>
<p>\({@ducks@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Purity Ekadeli</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires a  learner to calculate the probability of picking a duck from a group of birds, given the ratio of ducks to the total number of birds, and express the answer as a decimal.</p]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>p_duck</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*mistake1*/
p_duck_mistake1: decimalplaces(( total_birds/ducks ) , 2);/*Probability of ducks*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>p_duck</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the wrong answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>p_duck_mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have divided total number of birds by number of ducks while calculating the probability</p>
<p>This is incorrect because probability  is calculated by dividing number of ducks by total number of birds</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1912432677</deployedseed>
    <deployedseed>1262831537</deployedseed>
    <deployedseed>741344151</deployedseed>
    <deployedseed>697869667</deployedseed>
    <deployedseed>2025402132</deployedseed>
    <deployedseed>2041937815</deployedseed>
    <deployedseed>1528727847</deployedseed>
    <deployedseed>1097572703</deployedseed>
    <deployedseed>1275205459</deployedseed>
    <deployedseed>651680355</deployedseed>
    <deployedseed>1952416343</deployedseed>
    <deployedseed>1769870303</deployedseed>
    <deployedseed>204492194</deployedseed>
    <deployedseed>1281131913</deployedseed>
    <deployedseed>1016619772</deployedseed>
    <deployedseed>366844646</deployedseed>
    <deployedseed>1264009586</deployedseed>
    <deployedseed>1077679782</deployedseed>
    <deployedseed>1020713988</deployedseed>
    <deployedseed>568711690</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>p_duck</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
