<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>calculating probability that the ball is being drawn as a percentage</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A box contains \({@green_balls@}\) green balls, \({@red_balls@}\) red balls, and \({@yellow_balls@}\) yellow balls. If one ball is drawn randomly from the box, find the probability that the ball drawn is: </p>
<p>Express each probability as a percentage</p>
<ol type="a">
  <li>
    <p>Green</p>
    <p>P(Green) \( = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]%</p>
  </li>
  <li>
    <p>Not red</p>
    <p>P(Not red) \( =  \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]%</p>
  </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</p>
<p>a. The probability of drawing a green ball is calculated by dividing the number of green balls by the total number of balls:</p>
<p>P(Green) = Number of green balls ÷ Total number of balls</p>
<ul>
      <li>Total number of balls = \({@total_balls@}\)</li>
      <li>Number of green balls = \({@green_balls@}\)</li>
</ul>
<p>P(Green) = \({@green_balls@}\) ÷ \({@total_balls@}\) = <strong>\({@P_green@}\)</strong></p>
<P>P(Green) = <strong>\({@P_green@}\) * 100%</strong> = <strong>\({@P_green_percent@}\)%</strong></p>
<p>Therefore, the probability of drawing a green ball is <strong>\({@P_green_percent@}\)%</strong></p>

<p>b. The probability of drawing a ball that is not red:</p>
<p>P(not red) = Number of balls that are not red ÷ Total number of balls</p>
<p>The number of balls that are not red is the sum of the green and yellow balls.</p>
<ul>
      <li>Number of balls that are not red = \({@green_balls@}\) + \({@yellow_balls@}\) = \({@total_green_yellow@}\)</li>
      <li>Total number of balls = \({@total_balls@}\)</li>
</ul>
<p>P(not red) = \({@total_green_yellow@}\) ÷ \({@total_balls@}\) = <strong>\({@P_not_red@}\)</strong></p>
<p><strong>\({@P_not_red@}\) * 100%</strong> = <strong>\({@P_not_red_percent@}\)%</strong></p>
<p>Therefore, the probability of drawing a ball that is not red is <strong>\({@P_not_red_percent@}\)%</strong></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>green_balls: rand(5) + 5;/*number of green balls*/
red_balls: rand(5) + 3;/*number of red balls*/
yellow_balls: rand(5)+1;/*number of yellow balls*/
total_green_yellow: green_balls+yellow_balls;/*sum of green and yellow balls*/
total_balls: green_balls+red_balls+yellow_balls ;/*Total number of balls*/
P_green: green_balls / total_balls;/*probability of green*/
P_green_percent : decimalplaces(P_green * 100 ,2);/*probability of green as a percentage*/
P_not_red: total_green_yellow / total_balls;/*probability of not red*/
P_not_red_percent : decimalplaces(P_not_red * 100 ,2);/*probability of not red as a percentage*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@P_green@}\)</p>
<p>\({@P_not_red@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Purity Ekadeli</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the learner to calculate and express as percentages the probabilities of drawing a green ball and a ball that is not red from a set of colored balls.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>P_green_percent</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>P_not_red_percent</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>P_green_percent</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>P_not_red_percent</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
