<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the probability of Selecting Orange Juice</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A fruit vendor sold \({@mango@}\) bottles of mango juice, \({@orange@}\) bottles of orange juice, and \({@apple@}\) bottles of apple juice in a week.</p>
<p>What is the probability (in decimal form) that the next bottle sold will be orange juice?</p>
<p>Give your answer to 2 decimal places.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Using the probability formula:</p>
<p><em>Probability = Number of favourable outcomes ÷ Total outcomes</em></p>
<p>To calculate the probability that the next bottle sold will be orange juice:</p>
<p><em>P(Orange) = Number of Orange Juice Bottles ÷ Total number of bottles</em></p>
<ul>
  <li>Mango juice bottles = <strong>{@mango@}</strong></li>
  <li>Orange juice bottles = <strong>{@orange@}</strong></li>
  <li>Apple juice bottles = <strong>{@apple@}</strong></li>
  <li>Total number of bottles = <strong>{@total@}</strong></li>
  <li>P(Orange) = {@ orange @} ÷ {@ total @} = {@ P_orange @}</li>
</ul>
<p>Therefore, Probability that the next bottle sold will be orange is  <strong>{@ P_orange @}</strong</p>
​]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>mango : rand(10) + 1;
orange : rand(10) + 1;
apple : rand(10) + 1;
total : mango + orange + apple;
P_orange : decimalplaces((orange / total) ,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@total@}\)</p>
<p>\({@P_orange@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Purity Ekadeli</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires learners to calculate and express the probability, in decimal form, of selecting an orange juice bottle from a given set of randomly distributed juice types.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>P_orange</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*mistake*/
P_orange_mistake1: decimalplaces((total / orange) ,2)</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>P_orange</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the wrong answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>P_orange_mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have divided total number of bottles by number of orange juice bottles while calculating the probability</p>
<p>This is incorrect because probability  is calculated by dividing number of orange juice bottles by total number of bottles</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>31589014</deployedseed>
    <deployedseed>625614467</deployedseed>
    <deployedseed>651220246</deployedseed>
    <deployedseed>284969007</deployedseed>
    <deployedseed>1133773487</deployedseed>
    <deployedseed>1238635684</deployedseed>
    <deployedseed>1232202898</deployedseed>
    <deployedseed>496821129</deployedseed>
    <deployedseed>888051650</deployedseed>
    <deployedseed>1768874060</deployedseed>
    <deployedseed>458343276</deployedseed>
    <deployedseed>1043609465</deployedseed>
    <deployedseed>529389373</deployedseed>
    <deployedseed>1622894136</deployedseed>
    <deployedseed>555573485</deployedseed>
    <deployedseed>1709398799</deployedseed>
    <deployedseed>504421745</deployedseed>
    <deployedseed>1368504161</deployedseed>
    <deployedseed>1075742796</deployedseed>
    <deployedseed>1549429656</deployedseed>
    <deployedseed>1988903980</deployedseed>
    <deployedseed>641881556</deployedseed>
    <deployedseed>162509225</deployedseed>
    <deployedseed>418086147</deployedseed>
    <deployedseed>1288148188</deployedseed>
    <deployedseed>878532422</deployedseed>
    <deployedseed>1281701680</deployedseed>
    <deployedseed>404812826</deployedseed>
    <deployedseed>458861184</deployedseed>
    <deployedseed>1901592300</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>P_orange</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
