<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Bar Graph of Monthly Sales</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>(a) Here is a frequency table for the time taken for students to travel to school:</p>
<center>
<table style="width:500px; border:1px solid black; text-align:left">
  <tr>
    <th>Range \([min]\)</th>
    <th>Frequency</th>
    <th>Frequency Density</th>
  </tr>
  <tr>
    <td>\(0 < t \leq {@second(first(data))@}\)</td>
    <td>{@third(first(data))@}</td>
    <td>[[input:ans1]]</td>
  </tr>
  <tr>
    <td>\({@first(second(data))@} < t \leq {@second(second(data))@}\)</td>
    <td>{@third(second(data))@}</td>
    <td>[[input:ans2]]</td>
  </tr>
  <!-- Repeat for other ranges -->
</table>
</center>
<p hidden>'[[input:ans3]][[input:ans4]][[input:ans5]][[input:ans6]]</p>
<p hidden>[[validation:ans1]] [[validation:ans2]] [[validation:ans3]] [[validation:ans4]] [[validation:ans5]] [[validation:ans6]]</p>
<p align="right"></p>
<p>[[feedback:prt1]]</p>

<p><br>(b) By dragging the red points, complete the histogram below for the time taken for students to travel to school in the morning.</p>

<p><br></p>
<center>
[[jsxgraph width="750px" height="450px" input-ref-ans7="sR1" input-ref-ans8="sR2"]]
    var ranges = [
      [0, {#second(first(data))#}],
      [{#first(second(data))#}, {#second(second(data))#}],
      ...
    ];
    JXG.Options.text.useMathJax = true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox: [-5, 15, 65, -3],
      showCopyright: false,
      showNavigation: false,
      keepAspectRatio: false,
      grid: true
    });
    // Add axes, segments, gliders, and rectangles here
    stack_jxg.bind_point(sR1, pointList[0]);
    stack_jxg.bind_point(sR2, pointList[1]);
    ...
    board.unsuspendUpdate();
[[/jsxgraph]]
</center>

<p hidden>[[input:ans7]] [[input:ans8]] [[input:ans9]] [[input:ans10]] [[input:ans11]] [[input:ans12]]</p>
<p hidden>[[validation:ans7]] [[validation:ans8]] [[validation:ans9]] [[validation:ans10]] [[validation:ans11]] [[validation:ans12]]</p>
<p align="right"><br> </p>
<p>[[feedback:prt2]]</p>

<p><br>(c) Estimate the number of students whose journey to school takes less than 20 minutes.</p>
<p align="right">[[input:ans13]] [[validation:ans13]]</p>
<p align="right"> </p>
<p>[[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<br>(a)
<p>The sales data for each month is as follows:</p>
<p><br></p>
<center>
<table style="width:450px; border:1px solid black; text-align:left">
  <tr style="height:70px; border-bottom:1px solid black">
    <th style="padding:10px">Month</th>
    <th style="padding:10px">Sales (in $1000)</th>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">{@months[1]@}</td>
    <td style="padding:10px">{@sales_data[1]@}</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">{@months[2]@}</td>
    <td style="padding:10px">{@sales_data[2]@}</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">{@months[3]@}</td>
    <td style="padding:10px">{@sales_data[3]@}</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">{@months[4]@}</td>
    <td style="padding:10px">{@sales_data[4]@}</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">{@months[5]@}</td>
    <td style="padding:10px">{@sales_data[5]@}</td>
  </tr>
  <tr style="height:70px">
    <td style="padding:10px">{@months[6]@}</td>
    <td style="padding:10px">{@sales_data[6]@}</td>
  </tr>
</table>
</center>

<p><br>(b) Here is the complete bar graph shown below for the monthly sales.</p>

<p>(c) To estimate the total sales for the first quarter (January to March), sum the sales for these months:
\[
{@sales_data[1]@} + {@sales_data[2]@} + {@sales_data[3]@} = {@sum(sales_data[i], i, 1, 3)@}.
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Define ranges and frequencies */
data: [
  [0, 5, 10],    /* Range: 0 < t <= 5, Frequency: 10 */
  [5, 10, 20],   /* Range: 5 < t <= 10, Frequency: 20 */
  [10, 15, 15],  /* Range: 10 < t <= 15, Frequency: 15 */
  [15, 20, 25],  /* Range: 15 < t <= 20, Frequency: 25 */
  [20, 25, 30],  /* Range: 20 < t <= 25, Frequency: 30 */
  [25, 60, 35]   /* Range: 25 < t <= 60, Frequency: 35 */
];

/* Compute frequency densities */
tans1: third(first(data)) / (second(first(data)) - first(first(data)));
tans2: third(second(data)) / (second(second(data)) - first(second(data)));
tans3: third(third(data)) / (second(third(data)) - first(third(data)));
tans4: third(fourth(data)) / (second(fourth(data)) - first(fourth(data)));
tans5: third(fifth(data)) / (second(fifth(data)) - first(fifth(data)));
tans6: third(sixth(data)) / (second(sixth(data)) - first(sixth(data)));

/* Coordinates for JSXGraph gliders */
tans7: [float((first(first(data)) + second(first(data))) / 2), tans1];
tans8: [float((first(second(data)) + second(second(data))) / 2), tans2];
tans9: [float((first(third(data)) + second(third(data))) / 2), tans3];
tans10: [float((first(fourth(data)) + second(fourth(data))) / 2), tans4];
tans11: [float((first(fifth(data)) + second(fifth(data))) / 2), tans5];
tans12: [float((first(sixth(data)) + second(sixth(data))) / 2), tans6];

/* Total number of students for t <= 20 */
tans13: if first(sixth(data)) < 20 then (
          first(first(data)) * tans1 +
          first(second(data)) * tans2 +
          first(third(data)) * tans3 +
          first(fourth(data)) * tans4 +
          first(fifth(data)) * tans5 +
          (20 - first(sixth(data))) * tans6
        ) else if first(fifth(data)) < 20 then (
          first(first(data)) * tans1 +
          first(second(data)) * tans2 +
          first(third(data)) * tans3 +
          first(fourth(data)) * tans4 +
          (20 - first(fifth(data))) * tans5
        ) else if first(fourth(data)) < 20 then (
          first(first(data)) * tans1 +
          first(second(data)) * tans2 +
          first(third(data)) * tans3 +
          (20 - first(fourth(data))) * tans4
        ) else if first(third(data)) < 20 then (
          first(first(data)) * tans1 +
          first(second(data)) * tans2 +
          (20 - first(third(data))) * tans3
        ) else if first(second(data)) < 20 then (
          first(first(data)) * tans1 +
          (20 - first(second(data))) * tans2
        ) else (
          20 * tans1
        );]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ans1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>ans10</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>algebraic</type>
      <tans>ans11</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>algebraic</type>
      <tans>ans12</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>algebraic</type>
      <tans>ans13</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ans2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ans3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ans4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ans5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ans6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ans7</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ans8</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>ans9</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>6.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>node1_test: if ans1=correct_data[1] and ans2=correct_data[2] and ans3=correct_data[3] and ans4=correct_data[4] and ans5=correct_data[5] and ans6=correct_data[6] then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>node1_test</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>6.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>node1_test: if ans7=correct_data[1] and ans8=correct_data[2] and ans9=correct_data[3] and ans10=correct_data[4] and ans11=correct_data[5] and ans12=correct_data[6] then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>node1_test</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>6.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>node1_test: if ans13=sum(correct_data[i], i, 1, 3) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>node1_test</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
