<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Converting speed in metres per second (m/s) to kilometres per hour (km/h)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Convert \({@speed_in_metres_var@}\) m/s to kilometres per hour. [[input:ans1]] km/h [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Convert the distance from metres to kilometres.</p>
<p>\({@speed_in_metres_var@}\) m \( = \left( \frac{@speed_in_metres_var@}{1000}\right)\) km</p>
<p>Convert the time from seconds to hours</p>
<p>\( 1 \) second \( = \left( \frac{1}{3600}\right) \) hours</p>
<p>Divide the distance in kilometres by the time in hours.</p>
<p>speed \( = \frac{\text{distance covered}}{\text{time taken}} = \frac{@speed_in_metres_var@}{1000} \div \frac{1}{3600}\)</p>
<p>\( \qquad = \frac{@speed_in_metres_var@}{1000} \times \frac{3600}{1}\)</p>
<p>\( \qquad = \frac{@speed_in_metres_var*3600@}{1000}\)</p>
<p>\( \qquad = {@speed_in_metres_var*3600/1000@}\) km/h</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>speed_in_metres_var:rand(51)+10; /* speed in metres per second */
speed_in_kilo_var:speed_in_metres_var/1000 *3600; /* speed in kilometres per hour */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>speed in m/s \({@speed_in_metres_var@}\)</p>
<p>speed in km/h \({@speed_in_kilo_var@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">https://docs.google.com/document/d/1KYBkUQ35RKv7cRLHmDgtVIH8Jk5Hsreryt5A6uv96Ig/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">Firstly, convert the given distance from meters to kilometers. Since there are 1000 meters in a kilometer, divide the given distance by 1000 to obtain the result in kilometers. Next, you need to convert the given time from seconds to hours. There are 3600 seconds in one hour, so divide the number of seconds by 3600 to convert the time to hours. Finally, calculate the speed by dividing the distance in kilometers by the time in hours. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>speed_in_kilo_var</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct converted speed in kilometres per hour.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>speed_in_kilo_var</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2061107503</deployedseed>
    <deployedseed>86282800</deployedseed>
    <deployedseed>484777803</deployedseed>
    <deployedseed>1148105629</deployedseed>
    <deployedseed>1905246958</deployedseed>
    <deployedseed>1666868064</deployedseed>
    <deployedseed>331808159</deployedseed>
    <deployedseed>168053860</deployedseed>
    <deployedseed>63948842</deployedseed>
    <deployedseed>388940539</deployedseed>
    <deployedseed>735201386</deployedseed>
    <deployedseed>281007772</deployedseed>
    <deployedseed>1016108641</deployedseed>
    <deployedseed>1102026924</deployedseed>
    <deployedseed>966200750</deployedseed>
    <deployedseed>438580978</deployedseed>
    <deployedseed>2110751813</deployedseed>
    <deployedseed>282478282</deployedseed>
    <deployedseed>700244771</deployedseed>
    <deployedseed>1991262201</deployedseed>
    <deployedseed>463008797</deployedseed>
    <deployedseed>1692017762</deployedseed>
    <deployedseed>1636952383</deployedseed>
    <deployedseed>1682878842</deployedseed>
    <deployedseed>1759113550</deployedseed>
    <deployedseed>85705016</deployedseed>
  </question>
</quiz>
