<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the surface area of a square-based pyramid from its net</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The figure below shows a pyramid. Sketch the net of the pyramid and work out the surface area of the pyramid from its net</p>
<p><b>Note:</b>Round off your final answer to <b>2</b> decimal places.</p>
<!DOCTYPE html>
<html>
<head>
  <title>3D Pyramid with Square Base</title>
</head>
<body>
  <svg width="540" height="540" viewBox="0 0 540 540" style="border:1px solid #ccc">
   !-- Base square (in perspective) now orange -->
<polygon points="144,360 324,360 324,234 144,234" fill="orange" stroke="black"/>
    <!-- All triangular faces now cyan -->
    <polygon points="252,126 324,360 324,234" fill="cyan" opacity="0.8"/>
    <polygon points="252,126 144,234 144,360" fill="cyan" opacity="0.6"/>
    <polygon points="252,126 324,360 252,360" fill="cyan" opacity="0.8"/>
    <polygon points="252,126 144,234 252,234" fill="cyan" opacity="0.8"/>

    <!-- Edges to the apex -->
    <line x1="144" y1="360" x2="252" y2="126" stroke="black"/>
    <line x1="324" y1="360" x2="252" y2="126" stroke="black"/>
    <line x1="324" y1="234" x2="252" y2="126" stroke="black"/>
    <line x1="144" y1="234" x2="252" y2="126" stroke="black"/>

    <!-- Apex -->
    <circle cx="252" cy="126" r="5.4" fill="black"/>

    <!-- Labels -->
    <text x="216" y="378" font-size="21.6" text-anchor="middle">{#width#} cm</text>
    <text x="346" y="300" font-size="21.6" text-anchor="middle">{#length#} cm</text>
    <text x="380.6" y="210" font-size="21.6" text-anchor="middle" transform="rotate(-55 291.6,252)">{#slanting_height#} cm</text>

    <!-- Vertex labels -->
    <text x="132" y="372" font-size="18" font-weight="bold">A</text>
    <text x="324" y="375" font-size="18" font-weight="bold">B</text>
    <text x="354" y="234" font-size="18" font-weight="bold">C</text>
    <text x="132" y="228" font-size="18" font-weight="bold">D</text>
    <text x="258" y="114" font-size="18" font-weight="bold">V</text>
  </svg>
</body>
</html>

<p>Surface area=[[input:ans1]]cm\(^2\) [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>(a).The figure below shows the net of the pyramid</p>

<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Geometric Net</title>
  <style>
    svg {
      background: #fff;
      border: 1px solid #ccc;
    }
    .label {
      font-size: 12px;
      fill: #000;
    }
    .dimension {
      font-size: 10px;
      fill: #000;
      text-anchor: middle;
    }
  </style>
</head>
<body>
  <svg width="720" height="720" viewBox="-180 -180 360 360">
    <!-- Square ABCD (orange) -->
    <polygon points="-36,-36 43.2,-36 43.2,36 -36,36" fill="orange" stroke="black" stroke-width="1"/>

    <!-- Top Triangle XDC (cyan) -->
    <polygon points="-36,-36 43.2,-36 0,-134.1" fill="cyan" stroke="black" stroke-width="1"/>

    <!-- Bottom Triangle ABY (cyan) -->
    <polygon points="-36,36 43.2,36 0,134.1" fill="cyan" stroke="black" stroke-width="1"/>

    <!-- Left Triangle DZA (cyan) -->
    <polygon points="-36,-36 -134.1,0 -36,36" fill="cyan" stroke="black" stroke-width="1"/>

    <!-- Right Triangle CBV (cyan) -->
    <polygon points="43.2,-36 134.1,0 43.2,36" fill="cyan" stroke="black" stroke-width="1"/>

    <!-- Perpendicular bisectors -->
    <line x1="0" y1="-134.1" x2="3.6" y2="-36" stroke="red" stroke-width="1" stroke-dasharray="4"/>
    <line x1="0" y1="134.1" x2="3.6" y2="36" stroke="red" stroke-width="1" stroke-dasharray="4"/>
    <line x1="-134.1" y1="0" x2="-36" y2="0" stroke="red" stroke-width="1" stroke-dasharray="4"/>
    <line x1="134.1" y1="0" x2="43.2" y2="0" stroke="red" stroke-width="1" stroke-dasharray="4"/>

    <!-- Labels -->
    <text x="-41" y="-41" class="label">D</text>
    <text x="48" y="-41" class="label">C</text>
    <text x="-45" y="45" class="label">A</text>
    <text x="48" y="45" class="label">B</text>
    <text x="0" y="-140" class="label" text-anchor="middle">V</text>
    <text x="0" y="148" class="label" text-anchor="middle">V</text>
    <text x="-140" y="5" class="label" text-anchor="middle">V</text>
    <text x="140" y="5" class="label" text-anchor="middle">V</text>

    <!-- Existing Dimensions -->
    <text x="-85" y="-26" class="dimension">{#slanting_height#} cm</text>
    <text x="-38" y="-80" class="dimension">{#slanting_height#} cm</text>
    <text x="-85" y="30" class="dimension">{#slanting_height#} cm</text>
    <text x="-33" y="90" class="dimension">{#slanting_height#} cm</text>
    <text x="90" y="30" class="dimension">{#slanting_height#} cm</text>
    <text x="35" y="90" class="dimension">{#slanting_height#} cm</text>
    <text x="90" y="-26" class="dimension">{#slanting_height#} cm</text>
    <text x="35" y="-80" class="dimension">{#slanting_height#} cm</text>

    <!-- New Dimensions -->
    <text x="3.6" y="32" class="dimension" text-anchor="middle">{#length#} cm</text>
    <text x="-60" y="32" class="dimension" transform="rotate(-90 -60,0)" text-anchor="middle">{#width#} cm</text>
  </svg>
</body>
</html>

<p>(b).  Calculating the surface area</p>
<p>\(\begin{align}\text{Area of the base}&=length\times width\\&={@width@}cm\times{@length@}cm\\&={@base_area@}cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Vertical height}&=\sqrt{{@slanting_height@}^2-{@half_legth@}^2}\\&=\sqrt{{@square_slanting_height@}-{@square_half_length@}}\\&={@vertical_height@} cm\end{align}\)</p>
<p>\(\begin{align}\text{Area of a triangular face}&=\frac{1}{2}\times base \times height\\&=\frac{1}{2}\times{@length@} cm\times{@vertical_height@} cm\\&={@area_of_triangle_A@}cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Since all the four triangular faces are similar,then surface area of the faces}&=4\times{@area_of_triangle_A@}\\&={@sum_area_of_triangular_faces@} cm^2\end{align}\)<p>
<p>\(\begin{align}\text{Surface area of pyramid }&=\text{Area of base}+\text{Area of triangular faces}\\&={@base_area@} cm^2+{@sum_area_of_triangular_faces@}cm^2\\&={@ta@}cm^2\end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*Surface area  of a square base pyramid*/
length:rand_with_step(5,35,1);/*randomized values for the length */
width:length;/*randomized values for the width */
slanting_height:rand_with_step(21,60,1);/*randomized values for the slanting height */

half_legth:0.5*length;/*half of the length value which will help us find the vertical height using pythagorean relationship */

base_area:length*width;/*base area of the square part*/

square_slanting_height:slanting_height^2;/*square of the slanting height*/
square_half_length:(length*0.5)^2;/*square of half the value of the length*/
vertical_height:sqrt(square_slanting_height-square_half_length);/*the value of the vertical height of the triangle which will enable us find its area*/
vertical_height_of_pyramid:sqrt(vertical_height^2-square_half_length);/*vertical height of the pyramid*/

area_of_triangle_A:0.5*length*vertical_height;/*area of the first triangle A*/
area_of_triangle_B:0.5*length*vertical_height;/*area of the second triangle B*/
area_of_triangle_C:0.5*width*vertical_height;/*area of the third triangle C*/
area_of_triangle_D:0.5*width*vertical_height;/*area of the fourth triangle D*/
sum_area_of_triangular_faces:4*area_of_triangle_A;

total_surface_area:decimalplaces((base_area)+(area_of_triangle_A*4),2);/*total surface area of the pyramid*/
ta:total_surface_area;/*teacher's answer*/
/*mistake1- confusing slanting heigth to be vertical height of the triangles*/
area_of_triangle_1:0.5*length*slanting_height;/*confusing slanting heigth to be vertical height*/
surface_area:base_area+(4*area_of_triangle_1);/*total surface area*/
ta_1:surface_area;

/*mistake 2: confusing vertical height of the pyramid to be the height of the triangles*/
area_of_triangle_2:0.5*length*vertical_height_of_pyramid;
surface_area:base_area+(4*area_of_triangle_2);/*total surface area*/
ta_2:surface_area;
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Total surface area={@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Eric Morara</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the student to find the surface area of a square-base pyramid from its net .</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct surface area</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.02</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>It looks like you confused the slanting height of the triangle with the vertical height when calculating the area.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>0.02</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you confused the slanting height of the triangle with the vertical height when calculating the area. Keep in mind that the height used in the calculation of the area of a triangle must be perpendicular to the base. You're on the right track, just be careful to identify the correct height next time</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>It seems you used the vertical height of the pyramid instead of the vertical height of the triangular face when working with the area of a triangle.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_2</tans>
        <testoptions>0.02</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It seems you used the vertical height of the pyramid instead of the vertical height of the triangular face when working with the area of a triangle. For triangle area calculations on a pyramid, use the height of the triangular face, not the vertical height of the whole pyramid. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1304574074</deployedseed>
    <deployedseed>2133722309</deployedseed>
    <deployedseed>86944550</deployedseed>
    <deployedseed>1639731779</deployedseed>
    <deployedseed>1711018908</deployedseed>
    <deployedseed>1108014335</deployedseed>
    <deployedseed>193047946</deployedseed>
    <deployedseed>1737134203</deployedseed>
    <deployedseed>581756681</deployedseed>
    <deployedseed>2133334049</deployedseed>
    <deployedseed>1009375310</deployedseed>
    <deployedseed>2126977359</deployedseed>
    <deployedseed>2008015258</deployedseed>
    <deployedseed>1654878667</deployedseed>
    <deployedseed>1438504464</deployedseed>
    <deployedseed>51318854</deployedseed>
    <deployedseed>1436231513</deployedseed>
    <deployedseed>34970773</deployedseed>
    <deployedseed>1596080662</deployedseed>
    <deployedseed>1046037262</deployedseed>
    <deployedseed>1110952750</deployedseed>
    <deployedseed>743225773</deployedseed>
    <deployedseed>940925322</deployedseed>
    <deployedseed>215368667</deployedseed>
    <deployedseed>1296817206</deployedseed>
    <deployedseed>677480307</deployedseed>
    <deployedseed>170882782</deployedseed>
    <deployedseed>2119835071</deployedseed>
    <deployedseed>1606722789</deployedseed>
    <deployedseed>2130437520</deployedseed>
    <deployedseed>1484696357</deployedseed>
    <deployedseed>1679699252</deployedseed>
    <deployedseed>1327142180</deployedseed>
    <deployedseed>1067461428</deployedseed>
    <deployedseed>158596202</deployedseed>
    <deployedseed>1439058697</deployedseed>
    <deployedseed>103975333</deployedseed>
    <deployedseed>789824666</deployedseed>
    <deployedseed>593568929</deployedseed>
    <deployedseed>1476977628</deployedseed>
    <deployedseed>328085749</deployedseed>
    <deployedseed>891882016</deployedseed>
    <deployedseed>1020906657</deployedseed>
    <deployedseed>2038948654</deployedseed>
    <deployedseed>1915114567</deployedseed>
    <deployedseed>1981355410</deployedseed>
    <deployedseed>109188079</deployedseed>
    <deployedseed>876365070</deployedseed>
    <deployedseed>775649167</deployedseed>
    <deployedseed>1677731631</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description><![CDATA[<p>It looks like you confused the slanting height of the triangle with the vertical height when calculating the area. Keep in mind that the height used in the calculation of the area of a triangle must be perpendicular to the base. You're on the right tra]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>check for correct surface area</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description><![CDATA[<p>It seems you used the vertical height of the pyramid instead of the vertical height of the triangular face when working with the area of a triangle. For triangle area calculations on a pyramid, use the height of the triangular face, not the vertical he]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
