<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the hypotenuse of a right-angled triangle</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the length of the hypotenuse \({@var_letter_hypotenuse@}\) in the right angled triangle drawn below. Write your answer correct to two decimal places.</p>
[[jsxgraph]]
  /* boundingbox:[left, top, right, bottom] */
 var board = JXG.JSXGraph.initBoard('jxgbox',{
            boundingbox:[-5,5,5,-5], 
            axis: false

        });
        var A = board.create('point',[-2,-1], {name:'A', visible:false});
        var C = board.create('point',[2,-1], {name:'C', visible:false});
        var B = board.create('point',[-2,4], {name:'B', visible:false});

        var triangle = board.create('polygon',[A,B,C],{fillcolor:'white',borders:{strokeColor:'black'}});

        var labelA = board.create('text',[0,-1.5,'{#base#} cm' ],{fontsize:12});
        var labelC = board.create('text',[-2.5,2,'{#height#} cm' ],{fontsize:12});
        var labelB = board.create('text',[0,2,'{#var_letter_hypotenuse#} cm' ],{fontsize:12});

[[/jsxgraph]]


<p>\({@var_letter_hypotenuse@} = \)[[input:ans1]] [[validation:ans1]][[feedback:prt1]] cm </p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>var_letter_hypotenuse:rand([x,y,z,r,t,y,u,p,f,d,s,q,m,n,k]);
base:rand_with_step(4,15,2);/*generates random digits for the base length of the triangle*/
height:rand_with_prohib(8,25,[base]);/*generates random digits for the height length of the triangle*/
hypotenuse:sqrt(base^2+height^2);/*calculates the length of the hypotenuse*/
ta:decimalplaces(hypotenuse,2);/*gives the length of the hypotenuse correct to 2 decimal places*/

/*mistake 1*/
hypotenuse_1:base+height;/*Adding base and height directly when finding the length of the hypotenuse*/

/*mistake 2*/
hypotenuse_2:base^2+height^2;/*Failing to find the square root after adding the squares of base and height*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@var_letter_hypotenuse@}\)
\({@ta@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Isdora Ochieng</p>
<p id="concept">https://docs.google.com/document/d/18f1snnR6wR0oRQsWiVNgv7rSnzQGHLoInV6PvbN4j9E/edit?tab=t.0#heading=h.vy8x1tbmoha9</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the learner to compute the length of the hypotenuse for a right-angled triangle giving their answer correct to two decimal places</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the hypotenuse of a triangle correct to 2 decimal places.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect answer when height and base is added to find hypotenuse</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>hypotenuse_1</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have added directly instead of finding the square root for the sum of squares of both base and height</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>0</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for the incorrect answer when squareroot is not computed</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>hypotenuse_2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have not computed for the square root of your answer.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>126911042</deployedseed>
    <deployedseed>1753559386</deployedseed>
    <deployedseed>553703926</deployedseed>
    <deployedseed>2129368695</deployedseed>
    <deployedseed>50720764</deployedseed>
    <deployedseed>673926186</deployedseed>
    <deployedseed>1114164692</deployedseed>
    <deployedseed>1878375976</deployedseed>
    <deployedseed>1657999259</deployedseed>
    <deployedseed>596108374</deployedseed>
  </question>
</quiz>
