<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the reduced echelon form of the given matrix</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the reduced echelon form of the given matrix</p>
<p>\({@mat@} = \)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>In order to find the reduced row echelon form of the matrix \({@mat@}\) we proceed as follows:</p>
<p>\(R_{2}\rightarrow   {@(1/a_21) @}R_{2}\) </p>
<p>\({@mat_A@}\)</p>
<p>\( R_{1}\rightarrow {@(1/a_11) @}R_{1}\)</p>
<p>\({@mat_B@}\)</p>
<p> \(R_{2}\rightarrow  R_{2} - R_{1}\)</p>
<p>\({@mat_C@}\)</p>
<p> \(R_{2}\rightarrow  {@a_22_2_rp@}R_{2}\)</p>
<p>\({@mat_D@}\)</p>
<p> \(R_{1}\rightarrow  R_{1} - {@a_12_1@}R_{2}\)</p>
<p>\({@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*variables a_11,a_12,a_13,a_21,a_22,a_23 are entries for our matrix */
a_11:rand(3)+1;   
a_12:rand_with_prohib(1,6,[a_11]);
a_13:rand_with_prohib(1,6,[a_11,a_12]);
a_21:rand_with_prohib(1,6,[a_11,a_12,a_13]);
a_22:rand_with_prohib(1,6,[a_11,a_12,a_13,a_21]);
a_23:rand_with_prohib(1,6,[a_11,a_12,a_13,a_21,a_22]);
mat:matrix([a_11,a_12,a_13],[a_21,a_22,a_23]); /* Matrix */
a_21_1:(1/a_21)*a_21;
a_22_1:(1/a_21)*a_22;
a_23_1:(1/a_21)*a_23;
mat_A:matrix([a_11,a_12,a_13],[a_21_1,a_22_1,a_23_1]); 
a_11_1:(1/a_11)*a_11;
a_12_1:(1/a_11)*a_12;
a_13_1:(1/a_11)*a_13;
mat_B:matrix([a_11_1,a_12_1,a_13_1],[a_21_1,a_22_1,a_23_1]);
a_21_2:a_21_1 - a_11_1;
a_22_2:a_22_1 - a_12_1;
a_23_2:a_23_1 - a_13_1;
mat_C:matrix([a_11_1,a_12_1,a_13_1],[a_21_2,a_22_2,a_23_2]);
a_22_2_rp:(1/a_22_2); /*reciprocal of a_22_2*/
a_21_3:(1/a_22_2)*a_21_2;
a_22_3:(1/a_22_2)*a_22_2;
a_23_3:(1/a_22_2)*a_23_2;
mat_D:matrix([a_11_1,a_12_1,a_13_1],[a_21_3,a_22_3,a_23_3]);
a_11_2:a_11_1 -( a_12_1*a_21_3);
a_12_2:a_12_1 - ( a_12_1*a_22_3);
a_13_2:a_13_1 -( a_12_1*a_23_3);
mat_E:matrix([a_11_2,a_12_2,a_13_2],[a_21_3,a_22_3,a_23_3]);
ta:rref(mat);  /* teachers answer*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \({@mat@}\)</p>
<p>\({@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct reduced row echelon form</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
