<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Forming linear equations involving addition</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Juma had \({@original_num_oranges@}\) oranges. He bought \({@oranges_bought@}\) more oranges. In total, he had \({@total_oranges@}\) oranges. Write an equation to show the number of oranges that he had.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Original number of oranges \( = {@original_num_oranges@}\)</p>
<p>Additional number of oranges \(= {@oranges_bought@}\) </p>
<p>Total number of oranges \( = {@total_num_oranges@}\)</p>
<p>\(\qquad\qquad\qquad\qquad{@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>original_num_oranges:rand([p,q,x,y,r]);/*initial number of oranges*/
oranges_bought:rand(5)+3;/*additional oranges bought*/
total_oranges:rand_with_step(9,15,1);/*total number of oranges*/
total_num_oranges:original_num_oranges + oranges_bought;/*sum of initial number of oranges and  additional oranges bought*/
ta:total_num_oranges=total_oranges;/*equation showing number of oranges */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@total_num_oranges@}\)</p>
<p>\({@total_oranges@}\)</p>
<p>ta</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Isdora Ochieng</p>
<p id="concept">https://docs.google.com/document/d/1Rq_9tY-VEPjNTW14GFjmt50OOfcKLaLZVYzjrg_5ojA/edit?tab=t.0#heading=h.2hldm7no7va2</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">The question requires the student to form a linear equation for the total number of oranges using initial number of oranges and additional oranges bought.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct equation of the total number of oranges</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1382653662</deployedseed>
    <deployedseed>665692789</deployedseed>
    <deployedseed>427056801</deployedseed>
    <deployedseed>1448100590</deployedseed>
    <deployedseed>529674079</deployedseed>
    <deployedseed>1564395219</deployedseed>
  </question>
</quiz>
