<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Numbers divisible by 5</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>From the numbers below, identify those that are divisible by 5?</p>


<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>A number is said to be divisible by 5 if the last digit of the number is 0 or 5. </p>
<p>According to this test, a number is divisible by 5 if its last digit is 0 or 5.</p>
<p>we need to check if the last digit of each number is zero.<p>
<p>Similarly if the last digit of each number is not zero or five then the number is not divisible by 5.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Divisibility test*/
num_div1:rand_with_step(111,10000,5);/*numbers not divisable by five*/
num_div2:rand_with_step(8000,10000,1000);/*numbers divisable by five*/
num_div3:rand_with_step(999,9999,5);/*numbers not divisable by five*/
num_div4:rand_with_step(100000,599999,5);/**numbers divisable by five*/
num_div5:rand_with_step(600000,999999,5);/*numbers not divisable by five*/
num_div6:rand_with_step(55001,90000,5);/*numbers not divisable by five*/
num_div7:rand_with_step(20000,70000,5);/*numbers divisable by five*/
num_div8:rand_with_step(9999,99999,10);/*numbers not divisable by five*/
num_div9:rand_with_step(200001,499999,5);/*numbers not divisable by five*/
num_div10:rand_with_step(10000,20000,5);/*numbers divisable by five*/
not_div_ans:(f,[num_div1,num_div3,num_div6,num_div8,num_div9]);/*numbers which are not divisable by five*/
ans:(f,[num_div2, num_div4, num_div5,num_div7,num_div10]);/*numbers divisable by five*/


correct_answer:rand_selection([
[1, true, num_div2],
[2, true, num_div4],[3, true, num_div7],[4, true, num_div5],[5, true, num_div10]
], 3);
incorrect_answers:rand_selection([
[6, false, num_div1],
[7, false, num_div3],
[8, false, num_div6],[9, false, num_div8],[10, false, num_div9]
], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);

/*ta:random_permutation([[num_div1,false],[num_div2,true],[num_div3,false],[num_div4,true],[num_div5,true],[num_div6,false],[num_div7,true],[num_div8,false],[num_div9,false],[num_div10,true]]);*/



</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>5</p>
<p>\({@ta@}\) </p>
<p>\({@num_val@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author"> Joseph Baya</p>
<p id="concept">https://docs.google.com/document/d/1NX6DgGasEp3Mz2NLIzsbb7w89EZJT2d-ApIQRT5ptxw/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">In this question, you are required to identify which numbers from a given list are divisible by 5. A number is divisible by 5 if its last digit is 0 or 5</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for all the numbers divisible by 5</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>587425987</deployedseed>
    <deployedseed>730349296</deployedseed>
    <deployedseed>929866184</deployedseed>
    <deployedseed>1100976131</deployedseed>
    <deployedseed>401904401</deployedseed>
    <deployedseed>1372619726</deployedseed>
    <deployedseed>1163608342</deployedseed>
    <deployedseed>577869025</deployedseed>
    <deployedseed>667531046</deployedseed>
    <deployedseed>325458470</deployedseed>
  </question>
</quiz>
