<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Solving for a variable in a quadratic equation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Find the possible solutions for \({@var_letter@}\) in the following quadratic equation: </p> 
<p>\({@eqn@} = 0\)</p>
<p>\({@var_letter@} = \)[[input:ans1]] [[validation:ans1]][[feedback:prt1]] and \({@var_letter@} = \)[[input:ans2]] [[validation:ans2]][[feedback:prt2]] </p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To solve for \({@var_letter@}\)  we apply the quadratic formula as follows:</p> 
<p> \({@var_letter@} = \displaystyle \frac{-b\pm \sqrt{b^{2}-4ac}}{2a}\) where \(a \) is the coefficient of \({@var_letter@}^2\) \(, b \) is the coefficient of \({@var_letter@}\) and \(c\) is the constant</p>
<p>Applying the formula above in the given equation, we get:</p>
<p> \({@var_letter@} = \displaystyle \frac{-({@coeff_variable@})\pm \sqrt{({@coeff_variable@})^{2}-(4\times{@coeff_x_sqr@}\times{@const@}})}{2\times{{{@coeff_x_sqr@}}}}\)</p>
<p>\({@var_letter@}  = \displaystyle \frac{{@coeff_variable*-1@}\pm \sqrt{{@coeff_variable^2@}-{@coeff_x_sqr*4*const@}}}{{@coeff_x_sqr*2@}}\)</p>
<p>\({@var_letter@} =\displaystyle \frac{{@coeff_variable*-1@}\pm \sqrt{{@coeff_variable^2-coeff_x_sqr*4*const@}}}{{@coeff_x_sqr*2@}}\) </p>
<p>\({@var_letter@} = \displaystyle \frac{{@coeff_variable*-1@}\pm {@sqrt(coeff_variable^2-coeff_x_sqr*4*const)@}}{{@coeff_x_sqr*2@}}\)</p>
<p>\({@var_letter@} = {@x_1@} \) and  \({@var_letter@} = {@x_2@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>var_letter:rand([x,y,z,s,t,u,v,w,t]);
x_1:rand(3)+2;/*first value of the variable*/
x_2:rand_with_prohib(1,10,[x_1]);/*second value of the variable*/
coeff_variable:-x_1-x_2;/*coefficient of x*/
const:-x_1*-x_2;/*constant*/
coeff_x_sqr:const/(x_1*x_2);
eqn_1:(var_letter-x_1)*(var_letter-(x_2));/*gives the unexpanded equation*/
eqn:ratsimp(expand(eqn_1));/*original equation*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@eqn@}\)</p>
<p>\({@x_1@}\)</p>
<p>\({@x_2@}\)</p>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author"> Sheila Cherotich</p>
<p id="concept">https://docs.google.com/document/d/1LhcSYgxtJJLV6zSSg63PL2XI3JUKXI8cpGJtGdZ6FBI/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the student to solve for a variable in a quadratic equation. Start by applying a formula to find the two roots for the stated quadratic equation which gives the possible solution of the given variable.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>x_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>x_2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct first value of x</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>x_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct second value of x</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>x_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1140593084</deployedseed>
    <deployedseed>836680140</deployedseed>
    <deployedseed>875291705</deployedseed>
    <deployedseed>1660354340</deployedseed>
    <deployedseed>1853461623</deployedseed>
    <deployedseed>1123698016</deployedseed>
    <deployedseed>1399735058</deployedseed>
    <deployedseed>872903028</deployedseed>
    <deployedseed>2013589317</deployedseed>
    <deployedseed>949496240</deployedseed>
    <deployedseed>1492064976</deployedseed>
    <deployedseed>1044206326</deployedseed>
    <deployedseed>674468288</deployedseed>
    <deployedseed>49830080</deployedseed>
    <deployedseed>1011136443</deployedseed>
    <deployedseed>757765925</deployedseed>
    <deployedseed>1574087279</deployedseed>
    <deployedseed>1508662134</deployedseed>
    <deployedseed>618768236</deployedseed>
    <deployedseed>2072176798</deployedseed>
  </question>
</quiz>
