<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Writing the Scale of Water Tank in statement form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A water tank that is \({@actual_height_m@}\) meters tall is represented in a scale drawing by a line of \({@drawing_length_cm@}\) centimeters.</p>
<p>Interpret and write the scale used in a statement form.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Given;</p>
 <ul> 
       <li>Actual height of the tank = \({@actual_height_m@}\) meters</li> 
       <li>Drawing height = \({@drawing_length_cm@}\) centimeters</li> 
</ul>
<p>To find out how much real distance 1 cm represents, divide the actual height by the drawing height:</p>
<p><em>Scale = Actual height ÷ Drawing height</em></p>
<p>We need both in the same unit so we can divide correctly.</p>
<p>Since the drawing is in centimeters, we should convert \({@actual_height_m@}\) meters to centimeters.</p>
<p>Since;</p>
<p>1 meter = 100 centimeters, so \({@actual_height_m@}\) meters = \({@actual_height_cm@}\) centimeters</p>
 <p>Now, divide the actual length (in cm) by the drawing length:</p>
 <p>Scale = \({@actual_height_cm@}\)   ÷ \({@drawing_length_cm@}\)  = 1 cm : \({@scale_cm@}\) cm</p>
 <p>Since \({@scale_cm@}\) cm = \({@scale_m@}\) meters, we write:</p>
<p>Scale: 1 cm represents \({@scale_m@}\) m</p>
 <p>Therefore, the scale used is 1 cm represents \({@scale_m@}\) meters</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>actual_height_m : rand(41) + 10; /* actual height in m */
actual_height_cm : actual_height_m * 100;/*actual height converted from m to cm*/
drawing_length_cm : rand(9) + 2; /* drawing length in cm */
scale_cm : actual_height_cm / drawing_length_cm;/*the scale in cm*/
scale_m :decimalplaces(scale_cm / 100 ,2) ;/*the scale converted to m*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<P>\({@actual_height_cm@}\)</p>
<P>\({@drawing_length_cm@}\)</p>
<P>\({@scale_m@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Purity Ekadeli</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the learner to interpret the scale used in the drawing and express it in a statement form, based on the given height of the water tank and its representation in the scale drawing. </p]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>scale_m</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*mistake1*/
scale_mistake1: decimalplaces(actual_height_m / drawing_length_cm ,2);/*the scale without converting the units to same units*/
</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>scale_m</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the wrong answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>scale_mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Dividing m by cm without converting units</p>
<p>You forgot to convert the actual height from meters to centimeters before comparing it to the drawing length.</p>
<p>This gives the wrong interpretation, because centimeters and meters are not the same units, and you can’t divide them directly without converting.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2077534004</deployedseed>
    <deployedseed>53922095</deployedseed>
    <deployedseed>1519610287</deployedseed>
    <deployedseed>1656705068</deployedseed>
    <deployedseed>1237311734</deployedseed>
    <deployedseed>689453500</deployedseed>
    <deployedseed>982104745</deployedseed>
    <deployedseed>1442908789</deployedseed>
    <deployedseed>1982711353</deployedseed>
    <deployedseed>397585994</deployedseed>
    <deployedseed>22112963</deployedseed>
    <deployedseed>1966847006</deployedseed>
    <deployedseed>1042450460</deployedseed>
    <deployedseed>1948122333</deployedseed>
    <deployedseed>305982999</deployedseed>
    <deployedseed>905257089</deployedseed>
    <deployedseed>852798186</deployedseed>
    <deployedseed>461046762</deployedseed>
    <deployedseed>749302117</deployedseed>
    <deployedseed>1207770286</deployedseed>
    <deployedseed>796671859</deployedseed>
    <deployedseed>200701977</deployedseed>
    <deployedseed>1353767988</deployedseed>
    <deployedseed>365788452</deployedseed>
    <deployedseed>1972381666</deployedseed>
    <deployedseed>1234091062</deployedseed>
    <deployedseed>1823144656</deployedseed>
    <deployedseed>801851228</deployedseed>
    <deployedseed>1331699955</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>scale_m</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
