<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Converting Temperature from Kelvin to degrees Celsius</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The temperature at the \(\text{@name_place@} \) is \({@temperature_inkelvin@}\) Kelvin. Convert this temperature to degrees Celsius.</p>
<p>Temperature at the {@name_place@} is  [[input:ans1]] \(^{\circ}C\) [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked soluition</h3>
<p>In converting units of temperature from Kelvin to degrees Celsius we should know that, the temperature of frozen water at sea level is \(0^{\circ}C\) while in Kelvin scale it is \({@units_fixedscale@}\)\(\textit{K}\). Therefore, \(0^{\circ}C\) \(=\) \({@units_fixedscale@}\)\(\textit{K}\) .</p>
<p>To convert temperature in degrees Celsius to Kelvin, \({@units_fixedscale@}\) \(\textit{K}\) is added to the temperature given in degrees Celsius \((0^{\circ}C)\).</p>
<p>A short summary to the whole statement above is that:</p>
<p>\(\textit{Degrees Celsius}\) \(=\) \(\textit{Kelvin}\) - \({@units_fixedscale@}\)</p>
<p>As for our question, we are required to convert the temperature \({@temperature_inkelvin@}\)\(\textit{K}\) to degrees Celsius.</p>
<p>We substitute the temperature given in the equation above.</p>
<p>Since temperature given is \({@temperature_inkelvin@}\)\(\textit{K}\).</p>
<p>The temperature in Degrees Celsius  \(= {@temperature_inkelvin@}-{@units_fixedscale@}\)</p>
<p>\(\qquad \qquad \qquad \qquad  \qquad  \quad \quad\quad \quad  = {@degrees_celsius@}\)</p>
<p>Therefore;</p>
<p>Temperature at the \(\text{@name_place@}\) is \({@degrees_celsius@}^{\circ}C\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[name_place:rand(["Beach","Lake,Ocean","Bar","Mountain","Hill side","Cafe","Sea side","Island","Rural Area","Country side","River side","Garden","Park"]);/*name of the places*/
units_fixedscale:273;/*Fixed unit for kelvin at 0 degrees Celsius.*/
temperature_inkelvin:rand_with_step(100,1000,5);/*Temperature given in the question in temperature.*/
degrees_celsius:temperature_inkelvin-units_fixedscale;/*The  solution to the question */




]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>{@name_place@} </p>
<p>\({@temperature_inkelvin@}\)</p>
<p>\({@degrees_celsius@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Henry Onyango</p>
<p id="concept">https://docs.google.com/document/d/1cHt74X0dvn7KfDOoqRLYclgnp2OXeGfc6Ke6uGCVkqM/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">The question given is all about units in temperature given in Kelvin . The student is  is required to convert the temperature given in Kelvin to degrees Celsius.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>degrees_celsius</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Potential mistake 1*/

degrees_celsius_mistake1:temperature_inkelvin+units_fixedscale;/*The  solution to the question */
</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct temperature in degrees Celsius.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>degrees_celsius</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Adding instead of subtracting 273 units</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>degrees_celsius_mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have failed to subtract \({@units_fixedscale@}\) units from the temperature given but instead you have added which has resulted in the wrong result.</p>
<p>Kindly when dealing with questions with units of temperature(Kelvin and degrees Celsius) check the below  step by step solution to get the correct answer in the next attempt.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>472905676</deployedseed>
    <deployedseed>996023072</deployedseed>
    <deployedseed>1061033094</deployedseed>
    <deployedseed>224894615</deployedseed>
    <deployedseed>1350164670</deployedseed>
    <deployedseed>1086162297</deployedseed>
    <deployedseed>1449794773</deployedseed>
    <deployedseed>1220853915</deployedseed>
    <deployedseed>1294265506</deployedseed>
    <deployedseed>1910043150</deployedseed>
    <deployedseed>855447522</deployedseed>
    <deployedseed>190915581</deployedseed>
    <deployedseed>1283064209</deployedseed>
    <deployedseed>1779335256</deployedseed>
    <deployedseed>1443097402</deployedseed>
    <deployedseed>328983402</deployedseed>
    <deployedseed>1318617293</deployedseed>
    <deployedseed>778905391</deployedseed>
    <deployedseed>337512490</deployedseed>
    <deployedseed>1761097376</deployedseed>
    <deployedseed>298172604</deployedseed>
    <deployedseed>2026239288</deployedseed>
    <deployedseed>1270207488</deployedseed>
    <deployedseed>985148236</deployedseed>
    <deployedseed>1843691590</deployedseed>
    <deployedseed>932978873</deployedseed>
    <deployedseed>1290215395</deployedseed>
    <deployedseed>272221833</deployedseed>
    <deployedseed>1106969732</deployedseed>
    <deployedseed>168515889</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>degrees_celsius</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
