<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Slider: Maclaurin convergence C - arctan(x)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<h5>Example 3</h5><p>The function \(f(x)=\arctan(x)\) has Maclaurin series \(\displaystyle\sum_{n=0}^\infty \frac{(-1)^n x^{2n+1}}{2n+1}=x-\frac{x^3}{3}+\frac{x^5}{5}-\frac{x^7}{7}+\cdots\).</p><p>This applet shows what those polynomials look like, with a slider so you can vary the value of \(N\):<br></p>

<div style="margin-left:auto;margin-right:auto;width:650px;">[[jsxgraph width='640px' height='300px']]
  JXG.Options.text.fontSize = 16; 
var board1 = JXG.JSXGraph.initBoard(divid, {axis:true, grid:true, boundingbox: [-3, 2, 3.5, -2], showCopyright: false, showNavigation:false});
board1.suspendUpdate();
var func = board1.create('functiongraph', [function(t){ return Math.atan(t); },-5,5],{strokeColor: "#aaaaaa"});
var s = board1.create('slider', [[0.75,-1.5],[2.75,-1.5],[0,0,10]], {name:'N',snapWidth:1});
function isLessThan(a, b) {
  return Object.is((a-b)%1, -0);
}
var ser = board1.create('functiongraph', [
  function(t) {
    var val = 0, i, sv = s.Value()+1;
    for(i = 0; isLessThan(i,sv); i++) {
      val = val + Math.pow(-1, i) * Math.pow(t, 2 * i + 1) / (2*i+1);
    }
    return val;
  }, -10, 10], {strokeColor: "#003399",strokeWidth:2});
board1.unsuspendUpdate();
[[/jsxgraph]]</div>

<p>Here, you should see that as you increase the value of \(N\), the Maclaurin polynomials do give a better approximation of \(f(x)\) for some \(x\) values but not for others.</p><p><br></p><p><span>U</span>sing the applet, which of the following values appear to be ones for which the Maclaurin series converges?<br></p><p>[[input:ans1]][[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The series seems to be converging for \(|x|&lt;1\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2020061000</text>
    </stackversion>
    <questionvariables>
      <text>corbase:[-0.5,0.5,0];
wrongbase:[-1.5,1.5];
mcq:multiselqnalpha(corbase,3,wrongbase,2);
ta1:mcq[1];
version:mcq[2];
corr1:mcq_correct(ta1);
incorr1:mcq_incorrect(ta1);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{#version#}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>corr1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>corr1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>rest(corr1)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
