<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding unkwown angles on two transversal lines.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>In the figure below, calculate the values of the angles marked with letters.</p>
[[jsxgraph]]

    const board = JXG.JSXGraph.initBoard('jxgbox', {
        boundingbox: [-5, 5, 5, -5], 
        axis: false
    });
    // creating points.
        var P = board.create('point', [-4,3], { name: 'P', visible: false});
        var Q = board.create('point', [4,3], { name: 'Q', visible: false });
        var R = board.create('point', [-4,-1], { name: 'R', visible:false});  
        var S = board.create('point', [4,-1], { name: 'S', visible: false});

    // Define two parallel lines
    let line1 = board.create('line', [[-4, 3], [4, 3]], { strokeColor: 'blue', strokeWidth: 2 });
    let line2 = board.create('line', [[-4, -1], [4, -1]], { strokeColor: 'blue', strokeWidth: 2 });

    // Define two transversals forming a triangle
    let trans1 = board.create('line', [[-3.5, 4], [0, -4]], { strokeColor: 'red', strokeWidth: 2 });
    let trans2 = board.create('line', [[2, 4], [0, -4]], { strokeColor: 'red', strokeWidth: 2 });

    // Intersection points
    let A = board.create('point', [-3.1, 3], { name: 'A', size: 3, color: 'black',visible: false });
    let B = board.create('point', [1.7, 3], { name: 'B', size: 3, color: 'black',visible: false });
    let C = board.create('point', [-1.3, -1], { name: 'C', size: 3, color: 'black',visible: false });
    let D = board.create('point', [0.7, -1], { name: 'D', size: 3, color: 'black',visible: false });
    let E = board.create('point', [0,-4], { name: 'E', size: 3, color: 'black',visible: false });
    let F = board.create('point', [2,4], { name: 'F', size: 3, color: 'black',visible: false });
    let G = board.create('point', [-3.5,4], { name: 'G', size: 3, color: 'black',visible: false });

    // Mark corresponding angles
    let angle1 = board.create('angle', [Q,B,F], { radius: 0.8, fillColor: 'cyan', name: 'a',visible: false } );
    let angle2 = board.create('angle', [S,D,B], { radius: 0.6, fillColor: 'cyan', name: '{#val_letter_3#}' } );
    let angle3 = board.create('angle', [A, B, D], { radius: 0.8, fillColor: 'cyan',name: '{#val_letter_4#}' } );
    let angle4 = board.create('angle', [E,D,S], { radius: 0.8, fillColor: 'green', name: 'f',visible: false } );
    let angle5 = board.create('angle', [D,E,C], { radius: 0.8, fillColor: 'green', name: '{#angle_DEC#}'  });
    let angle6 = board.create('angle', [E,C,D], { radius: 0.5, fillColor: 'purple', name: '{#val_letter_1#}'  });
    let angle7 = board.create('angle', [A,C,R], { radius: 0.8, fillColor: 'yellow', name: 'w' ,visible: false }); 
    let angle8 = board.create('angle', [R,C,E], { radius: 0.6, fillColor: 'green', name: 'x',visible: false  }); 
    let angle9 = board.create('angle', [C,D,E], { radius: 0.6, fillColor: 'yean', name: '{#val_letter_2#}'  }); 
    let angle10 = board.create('angle', [B,D,C], { radius: 0.8, fillColor: 'green', name: 'h' ,visible: false }); 
    let angle11 = board.create('angle', [D,C,A], { radius: 0.6, fillColor: 'blue', name: 'z',visible: false  }); 
    let angle12 = board.create('angle', [D,B,Q], { radius: 0.6, fillColor: 'pink', name: '{#val_letter_5#}'  });
    let angle13 = board.create('angle', [F,B,A], { radius: 0.6, fillColor: 'pink', name: 'd' ,visible: false });  
    let angle14 = board.create('angle', [B,A,G], { radius: 0.8, fillColor: 'red', name: 'q',visible: false  });
    let angle15 = board.create('angle', [P,A,C], { radius: 0.8, fillColor: 'red', name: 'n',visible: false  });
    let angle16 = board.create('angle', [G,A,P], { radius: 0.6, fillColor: 'blue', name: 'p' ,visible: false });
    let angle17 = board.create('angle', [C,A,B], { radius: 0.6, fillColor: 'blue', name: '{#angle_CAD#}&deg'  });

[[/jsxgraph]]
<p>val_letter_1[[input:ans1]] \(^\circ\)[[validation:ans1]][[feedback:prt1]]</p>
<p>val_letter_2[[input:ans2]] \(^\circ\)[[validation:ans2]][[feedback:prt2]]</p>
<p>val_letter_3[[input:ans3]] \(^\circ\)[[validation:ans3]][[feedback:prt3]]</p>
<p>val_letter_4[[input:ans4]] \(^\circ\)[[validation:ans4]][[feedback:prt4]]</p>
<p>val_letter_5[[input:ans5]] \(^\circ\)[[validation:ans5]][[feedback:prt5]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>val_letter_1:rand([a,b,c,z,s]);/*generates random numbers for the unknown angle  BEG*/
val_letter_2:rand([d,f,g,t,u]);/*generates random numbers for the unknown angle  GEA*/
val_letter_3:rand([h,j,k,v]);/*generates random numbers for the unknown angle * KEB*/
val_letter_4:rand([l,m,n,w]);/*generates random numbers for the unknown angle DFK */
val_letter_5:rand([p,q,r,x,y]);/*generates random numbers for the unknown angle  HFD*/
angle_CAD:rand(53)+45;/*generates random numbers for the known angle CAD */
angle_DEC:rand_with_step(20,70,1);/*generates random numbers for the known angle DEC */
angle_ECD:angle_CAD;/*corresponding angles*/
angle_CDE:180-(angle_CDA+angle_DEC);/*interior angles of a triangle*/
angle_SDB:angle_CDE;/*opposite angles*/
angle_ABD:angle_SDB;/*alternate angles*/
angle_DBQ:180-angle_SDB;/*co-interior angles*/
ta_1:angle_ECD;/*teacher's answer*/
ta:angle_CDE;/*teacher's answer*/
ta_5:angle_DBQ;/*teacher's answer*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Angles marked {@val_letter_1@} and \({@angle_CAD@}^\circ\) are corresponding angles. Since corresponding angles are equal, Therefore;<p/>
<p>Angles marked {@val_letter_1@} =\({@angle_CAD@}^\circ\)<p/>
<p>\(\quuad\){@ta_1@}</p>

<p>Angles marked {@val_letter_1@}, {@val_letter_2@} and \({@angle_DEC@}^\circ\) are interior angles of a triangle.<p/>
<p>Interior angles of a triangle add upto \(180^\circ\). therefore:<p/>
<p>{@val_letter_1@}+{@val_letter_2@}+ \({@angle_DEC@}^\circ\) =\(180^\circ\).</p>
<p>\({@angle_CAD@}^\circ\)+{@val_letter_2@}+ \({@angle_DEC@}^\circ\) =\(180^\circ\).</p>
<p>{@val_letter_2@}=\(180^\circ\)-(\({@angle_ECD@}^\circ\)+\({@angle_DEC@}^\circ\)<p/>
<p>\(\quuad\){@ta@}</p>

<p>{@val_letter_3@} and {@val_letter_2@}  are vertically opposite angles.  <p/>
<p>But {@val_letter_2@} ={@ta@} </p>
<p>Since vertically opposite angles are equal, Therefore<p/>
<p>{@val_letter_3@} ={@ta@}</p>

<p>{@val_letter_4@} and {@val_letter_3@}  are alternate angles.  <p/>
<p>But {@val_letter_3@} ={@ta@} </p>
<p>Since alternate angles are equal, Therefore<p/>
<p>{@val_letter_4@} ={@ta@}</p>

<p>{@val_letter_5@} and {@val_letter_3@}  are co-interior angles.  <p/>
<p>But {@val_letter_3@} ={@ta@} </p>
<p>Since the summazation of two co-interior angles add upto \(180^\circ\), Therefore<p/>
<p>{@val_letter_3@}+{@val_letter_5@} =\(180^\circ\)</p>
<p>{@angle_SDB@}+{@val_letter_5@} =\(180^\circ\)</p>
<p>{@val_letter_5@} =\(180^\circ\)-{@angle_SDB@}</p>
<P>p>\(\quuad\){@ta_2@}<p/></P>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct angle</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct angle</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>3.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct angle</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>4.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct angle</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>5.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct angle</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
