<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the volume of the cylindrical objects.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the cylindrical container below. Note.\( Leave \,your \,answer \,in \,2 \,decimal \,places.\). </p>
                                          <p>Use \( \pi\) = 3.14</p>
[[jsxgraph]]

 var board = JXG.JSXGraph.initBoard('jxgbox', { // STACK: when moving to STACK change 'jxgboard' for divid
            boundingbox: [-5, 5, 5, -5],
            keepaspectratio: true,
            showCopyright: false,
            axis:false,
            showNavigation: true
        });

        // Defining the points of the elips

        var b1 = board.create('point', [-1, -1], {'visible':false, 'fixed':true});
        var b2 = board.create('point', [1, -1], {'visible':false, 'fixed':true});
        var b3 = board.create('point', [1.2, -1], {'visible':false, 'fixed':true});
        var b3o = board.create('point', [-1.2, -1], {'visible':false, 'fixed':true});
        var t1 = board.create('point', [-1, 1], {'visible':false, 'fixed':true});
        var t2 = board.create('point', [1, 1], {'visible':false, 'fixed':true});
        var t3 = board.create('point', [1.2, 1], {'visible':false, 'fixed':true});
        var t3o = board.create('point', [-1.2, 1], {'visible':false, 'fixed':true});

        // Drawing the ellipse

        var ell_bottom = board.create('ellipse', [b1, b2, b3] , {fillColor: 'pink', borders: {strokeColor: 'blue'}});
        var ell_top = board.create('ellipse', [t1, t2, t3] , {fillColor: 'pink', borders: {strokeColor: 'blue'}});

        // Drawing the height of the cylinder.
        var l1 = board.create('segment', [b3, t3]);
        var l2 = board.create('segment', [b3o, t3o]);

        // Defining the points of the the radius
        var A = board.create('point', [-1.2, 1], {name: 'A',visible:false});
        var B = board.create('point', [0, 1], {name: 'B',visible:false});

        // Drawing the radius of the cylinder.

        var line1 = board.create('segment', [A, B], { name: 'Line AB', strokeWidth: 2, 'fixed':true });

        // Labelling the height and the radius of the cylinder.

        var labelA = board.create('text', [-0.5, 1.2, '{#radius#}cm'], {fontSize: 12 ,'fixed':true});
        var labelC = board.create('text', [1.3,0, '{#height#}cm'], {fontSize: 12 ,'fixed':true});

[[/jsxgraph]]
<p></p><p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]\(cm^3\)</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p> Volume of a cylinder = Cross sectional area \(\times \) height</p>
<p> </p>
<p> \( \begin{align} \pi r^2 \times height &= 3.14 \times {@radius@}cm  \times {@radius@}cm  \times  {@height@}cm  \\ &=   {@volume@}cm^3 \end{align}\)</p>


<p>  Volume of the cylinder =   {@volume@}\(cm^3\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*Randomization for the volume of the cylinder.*/
radius:rand(99)+1;
height:rand(100)+100;
volume_1: 3.14*radius*radius*height;
volume_2: volume_1;
volume: decimalplaces(volume_2,2);
/*mistakes - Calculating the volume using radius once.*/
vol_1: volume_1: 3.14*radius*height;
vol_2:decimalplaces(vol_1,2);
/*mistake 2 - using 22/7 instead of 3.14*/
cyl_1:22/7*radius*radius*height;
cyl_2:decimal places(cyl_1,2);
/*finding volume using π×D×h*/
radii: 2(radius);
mistake_1:3.14*radii*height;
mistake_2:decimalplaces(mistake_1,2);

</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\({@radius@}\)
\({@volume_1@}\)
\({@volume_2@}\)
\({@volume@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Harriet Mogaka</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">This question requires the learner to  find the volume of a cylinder given the radius and the height.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>volume</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Volume of a cylindrical object.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>volume</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Calculating the volume using radius once.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>vol_1</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You have worked the volume correctly but it seems you forgot to use \( \pi r^2\) instead you used  \(\pi r\). Kindly refer to the worked solution and notice the difference.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Using 22/7 instead of 3.14</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>cyl_2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You seem to have used 22/7 instead of 3.14.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Finding volume using π×Dπ×h</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>You seem to have computed the sum correctly but you seem to have used \( π 2(r) h\) instead of \( π r^2 h\)</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>913422209</deployedseed>
    <deployedseed>625471714</deployedseed>
    <deployedseed>728111090</deployedseed>
    <deployedseed>623717742</deployedseed>
    <deployedseed>1601103223</deployedseed>
    <deployedseed>1109300749</deployedseed>
    <deployedseed>618178989</deployedseed>
    <deployedseed>465769445</deployedseed>
    <deployedseed>1173607849</deployedseed>
    <deployedseed>1564868074</deployedseed>
    <deployedseed>1391642787</deployedseed>
    <deployedseed>1339563844</deployedseed>
    <deployedseed>318984998</deployedseed>
    <deployedseed>276779118</deployedseed>
    <deployedseed>83936491</deployedseed>
    <deployedseed>1839012077</deployedseed>
    <deployedseed>1674444673</deployedseed>
    <deployedseed>1146606094</deployedseed>
    <deployedseed>1788698035</deployedseed>
    <deployedseed>950564556</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>volume</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Calculating the volume using radius once.</description>
      <testinput>
        <name>ans1</name>
        <value>vol_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
