<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding square roots of numbers using a calculator.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p> Find the square root of the following numbers using a calculator. </p>
<p> a) \({@number1@}\) [[input:ans1]]  [[feedback:prt1]]</p>
<p> b) \({@number2@}\) [[input:ans2]]  [[feedback:prt2]]</p>
<p> c) \({@number3@}\) [[input:ans3]] [[feedback:prt3]]</p>
<p> d) \({@number4_deci@}\) [[input:ans4]]  [[feedback:prt4]]</p>


<p hidden>[[validation:ans1]][[validation:ans2]] [[validation:ans3]] [[validation:ans4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> Worked Solution</p>
<p> To find the square root of a number using a calculator, you press the key \(\sqrt{}\)<p>
<p> Next, key in the the given number into the calculator. </p>
<p>  Press the \(=\) sign then read the result on the screen of the calculator.</p>
<p>Therefore, </p>
<p>
<ol  type = "a" >
<li>
<p>
To find the square root of \({@number1@}\) on your calculator:
</p>
<p>  Press the key \(\sqrt{}\)</p>
<p> Press the key \({@fd_number1@}\) then \({@sd_number1@}\)</p>
<p> Press the key \(=\)</p>
<p> Read the result on the screen of your calculator</p>
<p> Hence \( \sqrt{{@number1 @}}= {@sqrt_final @} \)
</li>
<li>
<p>
To find the square root of \({@number2@}\) on your calculator:
</p>
<p>  Press the key \(\sqrt{}\)</p>
<p> Press the key \({@fd_number2@},\)\({@sd_number2@}\) then \({@td_number2@}\)</p>
<p> Press the key \(=\)</p>
<p> Read the result on the screen of your calculator</p>
<p> Hence \( \sqrt{{@number2 @}} = {@sqrt_final_2@} \)
</li>
<li>
<p>
To find the square root of \({@number3@}\) on your calculator:
</p>
<p>  Press the key \(\sqrt{}\)</p>
<p> Press the key \({@fd_number3@}, {@sd_number3@} , {@td_number3@}\) then \({@fourth_number3@}\)</p>
<p> Press the key \(=\)</p>
<p> Read the result on the screen of your calculator</p>
<p> Hence \( \sqrt{{@number3 @}} = {@sqrt_final_3@} \)
</li>
<li>
<p>
  Given \({@number4_deci@}\), follow the steps below to find its square root using a calculator:
</p>
<p>Step 1:  Press the key \(\sqrt{}\)</p>
<p> Step 2:  Press the key \({@fd_number4@}, {@sd_number4@} , . , {@td_number4@}\) then \({@fourth_number4@}\)</p>
<p> Step 3: Press the key \(=\)</p>
<p> Step 4: Read the result on the screen of your calculator</p>
<p> Hence \( \sqrt{{@number4 @}} = {@root_number4@} \)
</li>
</ol>
<p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>fd_number1:rand(9)+1;
sd_number1:rand_with_prohib(1,8,[fd_number1]);
number1:fd_number1*10+sd_number1;
square_root_1:sqrt(number1);
sqrt_decimal : float(sqrt(number1)); 
is_exact_root : is(integerp(sqrt(number1)));
sqrt_final : if is_exact_root then square_root_1 else sqrt_decimal;


fd_number2:rand(8)+1;
sd_number2:rand_with_prohib(1,8,[fd_number2]);
td_number2:rand_with_prohib(1,6,[fd_number2,sd_number2]);
number2:fd_number2*100+sd_number2*10+td_number2;
square_root_2: sqrt(number2);
sqrt_decimal_2 : float(sqrt(number2)); 
is_exact_root_2 : is(integerp(sqrt(number2)));
sqrt_final_2 : if is_exact_root 
               then square_root_2 
               else sqrt_decimal_2;





fd_number3:rand(9)+1;
sd_number3:rand_with_prohib(1,9,[fd_number3]);
td_number3:rand_with_prohib(1,9,[fd_number3,sd_number3]);
fourth_number3:rand_with_prohib(2,6,[fd_number3,sd_number3,td_number3]);
number3:fd_number3*1000+sd_number3*100+td_number3*10+fourth_number3;
square_number_3:sqrt(number3);
sqrt_decimal_3 : float(sqrt(number3)); 
is_exact_root_3 : is(integerp(sqrt(number3)));
/*sqrt_final_3 : if is_exact_root then sqrt_exact else sqrt_decimal;*/
sqrt_final_3 : if is_exact_root_3 
               then sqrt_decimal_3 
               else square_number_3;



fd_number4:rand(5)+1;
sd_number4:rand_with_prohib(1,4,[fd_number4]);
td_number4:rand_with_prohib(1,6,[fd_number4,sd_number4]);
fourth_number4:rand_with_prohib(1,9,[fd_number4,sd_number4,td_number4]);
number4:fd_number4*1000+sd_number4*100+td_number4*10+fourth_number4;
decimal_div:100;
number4_deci:float(number4/decimal_div);
square_root_4:sqrt(number4_deci);
root_number4:float(square_root_4);

/*mistake 1*/

/* NUMBER 1 */
fd_number1: rand(9) + 1;
sd_number1: rand_with_prohib(1, 8, [fd_number1]);
number1: fd_number1 * 10 + sd_number1;
square_1: number1 ^ 2;


/* NUMBER 2 */
fd_number2: rand(8) + 1;
sd_number2: rand_with_prohib(1, 8, [fd_number2]);
td_number2: rand_with_prohib(1, 6, [fd_number2, sd_number2]);
number2: fd_number2 * 100 + sd_number2 * 10 + td_number2;
square_2: number2 ^ 2;


/* NUMBER 3 */
fd_number3: rand(9) + 1;
sd_number3: rand_with_prohib(1, 9, [fd_number3]);
td_number3: rand_with_prohib(1, 9, [fd_number3, sd_number3]);
fourth_number3: rand_with_prohib(2, 6, [fd_number3, sd_number3, td_number3]);
number3: fd_number3 * 1000 + sd_number3 * 100 + td_number3 * 10 + fourth_number3;
square_3: number3 ^ 2;


/* NUMBER 4 */
fd_number4: rand(5) + 1;
sd_number4: rand_with_prohib(1, 4, [fd_number4]);
td_number4: rand_with_prohib(1, 6, [fd_number4, sd_number4]);
fourth_number4: rand_with_prohib(1, 9, [fd_number4, sd_number4, td_number4]);
number4: fd_number4 * 1000 + sd_number4 * 100 + td_number4 * 10 + fourth_number4;
decimal_div: 100;
number4_deci: float(number4 / decimal_div);
square_4: number4_deci ^ 2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>sqrt_final</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>sqrt_final_2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>sqrt_final_3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>root_number4</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct square root of a 2 digit number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sqrt_final</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for a squared number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>square_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have squared the number instead of square rooting. When finding the square root of a number using a calculator, use the key \(\sqrt{}\). Look at the step by step worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct square root of the given 3 digit number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>sqrt_final_2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for a squared number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>square_2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have squared the number instead of square rooting. When finding the square root of a number using a calculator, use the key \(\sqrt{}\). Look at the step by step worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct square root of a 4 digit number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>sqrt_final_3</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for a squared number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>square_3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have squared the number instead of square rooting. When finding the square root of a number using a calculator, use the key \(\sqrt{}\). Look at the step by step worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct square root of a  decimal number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>root_number4</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for a squared number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>square_4</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have squared the number instead of square rooting. When finding the square root of a number using a calculator, use the key \(\sqrt{}\). Look at the step by step worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1775598203</deployedseed>
  </question>
</quiz>
