<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a given rectangle.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the area of the following rectangle.</p>
<p>area = [[input:ans1]] \(m^2\)[[validation:ans1]][[feedback:prt1]]</p>
[[jsxgraph]]
 var board = JXG.JSXGraph.initBoard('jxgbox', {boundingbox: [-5, 30, 30, -5],
                keepAspectRatio:true, showCopyright:false, showNavigation:false ,axis:false});
        const A = board.create('point', [0,0], { name: 'A', visible: false});
        const B = board.create('point', [0,10], { name: 'B', visible: false });
        const C = board.create('point', [20,10], { name: 'C', visible: false});  
        const D = board.create('point', [20,0], { name: 'D', visible: false}); 
        board.create('polygon', [A, B, C,D], {borders: {
                                               strokeWidth: 1,
                                               strokeColor: 'black',},
                                              fillColor: 'none',});

        board.create('text', [-4,5, '{#width#} m'], { fontSize: 10 });
        board.create('text', [10,11, '{#length#} m'], { fontSize: 10 });
        board.create('text', [21,5, '9 m'], { fontSize: 0 });
        board.create('text', [10,5, '24 m'], { fontSize: 0 });
[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>The area of a rectangle is calculated by multiplying its length and width. </p>
<p>Given a rectangle with a length of \({@length@}\) meters and a width of \({@width@}\) meters, the area is found by multiplying \({@length@}\)  meters by \({@width@}\)  meters, resulting in an area of \({@area@}\)  square meters.
</p>Therefore area = \({@length@}\) \times  \({@width@}\) = \({@area@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>units:rand([m]);/*units*/
square_unit:(units)^2;/*square units*/
length:rand_with_step(50,100,2);/*length of the rectangle*/
width:rand_with_step(6,49,2);/*width of the rectangle*/
area:length*width;/*area of the rectangle*/
ta:area;/*teacher's answer*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(@ta@\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Eric Morara</p>
<p id="concept">https://ecampus.idems.international/question/bank/history/history.php?entryid=126329&returnurl=%2Fquestion%2Fedit.php%3Fcourseid%3D430%26cat%3D18807%252C121817&courseid=430</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The area of a rectangle is calculated by multiplying its length and width </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
