<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Expressing numbers in standard form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Express the following number in standard from:</p>
<p>\({@number_selected@} = \) [[input:ans1]] \( \times 10 \) ^ [[input:ans2]] [[validation:ans1]] [[validation:ans2]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To write \({@number_selected@}\) in standard form, first express the number (\({@number_selected@}\)) as a number between \( 1 \) and \( 10 \) multiplied by a power of \( 10 \). </p>
<p> \({@number_selected@} = {@ta_number_selected@} \times {@divisor_selected@} \)</p>
<p>Next, express \({@divisor_selected@}\) as a power of \( 10 \): </p>
<p>\({@divisor_selected@} = 10^ {@decimal_places_selected@} \)</p>
<p>So, in standard form:</p>
<p> \({@number_selected@} =  {@ta_number_selected@} \times 10^ {@decimal_places_selected@} \)</p>
<p></p>
<p></p>
<p></p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*Generate random numbers*/
number_1:rand_with_step(10,99,1);
number_2:rand_with_step(100,999,1);
number_3:rand_with_step(1000,9999,1);
number_4:rand_with_step(10000,99999,1);

random_index:rand(4)+1;

number_list:[number_1,number_2,number_3,number_4];
decimal_places_list:[1,2,3,4];
divisor_list:[10,100,1000,10000];

number_selected:number_list[random_index]
decimal_places_selected:decimal_places_list[random_index];
divisor_selected:divisor_list[random_index];

ta_number_selected:float(number_selected/divisor_selected);



</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Number: \({@number_selected@}\) </p>
<p>Decimal places: \({@decimal_places_selected@}\)</p>
<p>Divisor: \({@divisor_selected@}\)</p>
<p>Teacher's answer: \({@ta_number_selected@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The student is required to convert the given number into standard form, expressed as \( p \times 10^n \), where \( p \) is a number greater than or equal to 1 but less than 10, and \( n \) is an integer exponent of 10.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_number_selected</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>decimal_places_selected</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct placement of the decimal point</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_number_selected</tans>
        <testoptions>decimal_places_selected</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correct exponent</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>decimal_places_selected</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
