<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Expressing numbers to the required significant figures</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Express each of the following numbers to the given number of significant figures in brackets:</p>
<p>a) \({@first_number@}\)  \( ({@random_significant_number@})\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>b) \({@second_number_decimal@}\)  \( ({@random_significant_number2@})\)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>c) \({@third_number@}\)  \( ({@random_significant_number3@})\)[[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Part a: </p>
<p>The number {@first_number@} has \(5\) significant figures.</p>
<p>To round up \({@first_number@}\) to \({@random_significant_number@}\) significant figures, we consider the {@place_list[random_significant_number+1]@} significant digit.</p>
<p>Since {@first_number_digits[random_significant_number+1]@} {@message@} to the  {@place_list[random_significant_number]@} significant digit and we remove the remaining significant figures</p>
<p>Thus, {@first_number@} to \({@random_significant_number@}\) significant figures is {@ta1@}</p>
<p>Part b: </p>
<p>The number {@second_number_decimal@} has \(4\) significant figures.</p>
<p>To round up \({@second_number_decimal@}\) to \({@random_significant_number2@}\) significant figures, we consider the {@place_list[random_significant_number2+1]@} significant digit. </p>
<p>Since {@second_number_digits[random_significant_number2+1]@} {@second_decision@} to the  {@place_list[random_significant_number2]@} significant digit and we replace the remaining significant digits with zeros.
</p>
<p>Thus, {@second_number_decimal@} to \({@random_significant_number2@}\) significant figures is {@ta2@}</p>
<p></p>
<p>Part c: </p>
<p>The number {@third_number@} has \(4\) significant figures.</p>
<p>To round up \({@third_number@}\) to \({@random_significant_number3@}\) significant figures, we consider the {@place_list[random_significant_number3+1]@} significant digit. </p>
<p>Since {@third_number_digits[random_significant_number3+1]@} {@third_decision@} to the  {@place_list[random_significant_number3]@} significant digit and we replace the remaining significant digits with zeros.
</p>
<p>Thus, {@third_number@} to \({@random_significant_number3@}\) significant figures is {@ta3@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[first_digit:rand(9)+1;
second_digit:rand_with_prohib(1,9,[first_digit]);
third_digit:rand_with_prohib(1,9,[first_digit,second_digit]);
fourth_digit:rand_with_prohib(1,9,[first_digit,second_digit,third_digit]);
fifth_digit:rand_with_prohib(1,9,[first_digit,second_digit,third_digit,fourth_digit]);
stackfltfmt:"~f";
first_number:float(((first_digit*10000)+(second_digit*1000)+(third_digit*100)+(fourth_digit*10)+(fifth_digit))/100);
first_number_digits:[first_digit,second_digit,third_digit,fourth_digit,fifth_digit];
place_list:["first","second","third","fourth","fifth"];
random_significant_number:rand(4)+1;
if first_number_digits[random_significant_number+1] = 5 then
   message:"is equal to 5. Therefore, we add 1 "
elseif first_number_digits[random_significant_number+1] > 5 then
   message:"is greater than 5. Therefore, we add 1 "
else 
   message:"is less than 5. Therefore, we don't add 1 ";

rounding_message(first_number_digits, random_significant_number) := block(
    [next_digit, message],
    next_digit : first_number_digits[random_significant_number + 1],
    if next_digit = 5 then
        message : "is equal to 5. Therefore, we add 1"
    elseif next_digit > 5 then
        message : "is greater than 5. Therefore, we add 1"
    else
        message : "is less than 5. Therefore, we don't add 1",
    return(message)
)$

random_significant_number2:rand(3)+1;
second_number:(first_digit*1000)+(second_digit*100)+(third_digit*10)+fourth_digit;
second_number_decimal:float(second_number/1000000);
second_number_digits:[first_digit,second_digit,third_digit,fourth_digit];
second_decision:rounding_message(second_number_digits,random_significant_number2)




third_first_digit:rand(9)+1;
third_second_digit:rand_with_prohib(1,9,[third_first_digit]);
third_third_digit:rand_with_prohib(1,9,[third_first_digit,third_second_digit]);
third_fourth_digit:rand_with_prohib(1,9,[third_first_digit,third_second_digit,third_third_digit]);
third_number:(third_fourth_digit*1000)+(third_third_digit*100)+(third_second_digit*10)+third_first_digit;
third_number_digits:[third_fourth_digit,third_second_digit,third_third_digit,third_first_digit];
random_significant_number3:rand(3)+1;
third_decision:rounding_message(third_number_digits,random_significant_number3);

ta1:significantfigures(first_number,random_significant_number);
ta2:significantfigures(second_number_decimal,random_significant_number2);
ta3:significantfigures(third_number,random_significant_number3);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\( {@first_number@}\)</p>
<p>\( {@second_number_decimal@}\)</p>
<p>\( {@third_number@}\)</p>
<p>\( {@ta1@}\)</p>
<p>\( {@ta2@}\)</p>
<p>\( {@ta3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The student is required to round the given number to the specified number of significant figures.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Chechs for the correct signficant figures</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>random_significant_number</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct siginificant figures</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>random_significant_number2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for significant figures in the given number</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>random_significant_number3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
