<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Rounding off decimal numbers to the required decimal places</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Round off \({@decimal_number@}\) to the number of decimal places indicated below:</p>
<p>(a) Correct to \(1\) decimal place.  [[input:ans1]] [[validation:ans1]]</p>
<p>(b) Correct to \(2\) decimal places. [[input:ans2]] [[validation:ans2]]</p>
<p>(c)Correct to  \(3\) decimal places.  [[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>We first place each digit in its correct place value position.</p>
<table border="1">
<thead>
<tr>
<th scope="col">Number</th>
<th scope="col">Ones</th>
<th scope="col">Decimal Point</th>
<th scope="col">Tenths</th>
<th scope="col">Hundredths</th>
<th scope="col">Thousandths</th>
<th scope="col">Ten Thousandths</th>
</tr>
</thead>
  <tbody>
    <tr>
      <td><p align="center">\({@decimal_number@}\)</p></td>
      <td><p align="center">\({@first_digit@}\)</p></td>
      <td><p align="center">.</p></td>
      <td><p align="center">\({@second_digit@}\)</p></td>
      <td><p align="center">\({@third_digit@}\)</p></td>
      <td><p align="center">\({@fourth_digit@}\)</p></td>
      <td><p align="center">\({@fifth_digit@}\)</p></td>
    </tr>
  </tbody>
</table>
<p>To round {@decimal_number@} to one decimal place, we examine the digit in the hundredths place.</p>
<p>The  hundredths digit is {@third_digit@}, which is {@message_1@} to the Tenths digit.</p>
<p>Thus {@decimal_number@} rounded to one decimal place is {@ta_1@}</p>

<p>To round {@decimal_number@} to two decimal places, we examine the digit in the Thousandths place.</p>
<p>The thousandths digit is {@fourth_digit@}, which is {@message_2@} to the Tenths digit.</p>
<p>Thus {@decimal_number@} rounded to two decimal place is {@ta_2@}</p>

<p>To round {@decimal_number@} to three decimal places, we examine the digit in the Ten Thousandths place.</p>
<p>The Ten Thousandths digit is {@fifth_digit@}, which is {@message_3@} to the Tenths digit.</p>
<p>Thus {@decimal_number@} rounded to three decimal places is {@ta_3@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[first_digit:rand(9)+1;
second_digit:rand_with_prohib(1,9,[first_digit]);
third_digit:rand_with_prohib(1,9,[first_digit,second_digit]);
fourth_digit:rand_with_prohib(1,9,[first_digit,second_digit,third_digit]);
fifth_digit:rand_with_prohib(1,9,[first_digit,second_digit,third_digit,fourth_digit]);

number:(first_digit*10000)+(second_digit*1000)+(third_digit*100)+(fourth_digit*10)+fifth_digit;
decimal_number:float(number/10000);

if third_digit >=5 then 
   (second_digit_ta:second_digit+1,message_1:"greater than or equal to 5. Therefore, we add 1 " )
else 
   (second_digit_ta:second_digit,message_1:"less than 5. Therefore, we don't add 1 " );
ta_1:float(((first_digit*10)+second_digit_ta)/10);

if fourth_digit >= 5 then 
   (third_digit_ta:third_digit+1 ,message_2:"greater than or equal to 5. Therefore, we add 1 " )
else 
   (third_digit_ta:third_digit,message_2:"less than 5. Therefore, we don't add 1 " );
ta_2:float(((first_digit*100)+(second_digit*10)+third_digit_ta)/100);

if fifth_digit >= 5 then 
   (fourth_digit_ta:fourth_digit+1 ,message_3:"greater than or equal to 5. Therefore, we add 1 " )
else 
   (fourth_digit_ta:fourth_digit,message_3:"less than 5. Therefore, we don't add 1 " );
ta_3:float(((first_digit*1000)+(second_digit*100)+(third_digit*10)+fourth_digit_ta)/1000);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@decimal_number@}\)</p>
<p>\({@ta_1@}\)</p>
<p>\({@ta_2@}\)</p>
<p>\({@ta_3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The student is required to round the given number to the specified number of significant figures.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct decimal place</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.333333</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.333333</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.333333</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
