<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the difference in distance covered by two people in Kilometers.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Steve lives \({@ fraction_1@}\) km from school. John lives  \({@ fraction_2@}\) km from school. By how many kilometers is John's home closer to the school compared to Steve's?</p>
<p> Distance \(=\)[[input:ans1]] [[validation:ans1]] km [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>The distance between Steve’s home and the school \(= \)\({@ fraction_1@}\) km </p>
<p>The distance between John’s home and the school \(=\) \({@ fraction_2@}\) km </p>
<p>To find how closer John is from the school than Steve,</p>
<p>We subtract the distance between John’s home and the school from the distance between Steve’s home and the school as shown below.</p>
<p>\(\quad = \)\({@ fraction_1@}\) km - \({@ fraction_2@}\) km<p>
<p>\(\quad = {@ta@}\) km </p>
<p>Therefore, John is \({@ta@}\) km closer to school than Steve.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>num_1:rand(4)+4;/*numerator of the first fraction*/
den_1:rand(4)+6;/*denominator of the second fraction*/
fraction_1:(num_1/den_1);/*represents distance between Steve's home and the school*/
num_2:rand(3)+1;/*numerator of the second fraction*/
den_2:rand(4)+7;/*denominator of the second fraction*/
fraction_2:(num_2/den_2);/*represents the distance between John's home and the school*/
ta:fraction_1 -fraction_2;/*subtracts the first fraction from the second fraction*/
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \({@ta@}\)</p>
<p>\({@ fraction_1@}\)</p>
<p>\({@ fraction_2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Sheila Cherotich</p>
<p id="concept">https://docs.google.com/document/d/11_e0Fu8E8rdFcpjK1lYP-ZHJzzIO104IFqShVNvCA9s/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the student to give the difference in distance between where Steve lives from the school and where John lives from the school in order to identify how John is closer from the school than Steve. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct distance to determine how closer John is to the school than Steve.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
