<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying numbers which are divisible by 10 from a given list</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Which of the following numbers are divisible by \(10\) ?</p>

<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To identify numbers divisible by 10, we need to apply the divisibility test. </p>
<p>According to this test, a number is divisible by 10 if its last digit is 0.</p>
<p>Given the numbers,</p>
<p>\({@num_div1@}\)</p>
<p>\({@num_div2@}\)</p>
<p>\({@num_div3@}\)</p>
<p>\({@num_div4@}\)</p>
<p>\({@num_div5@}\)</p>
<p>\({@num_div6@}\)</p>
<p>\({@num_div7@}\)</p>
<p>\({@num_div8@}\)</p>
<p>\({@num_div9@}\)</p>
<p>\({@num_div10@}\)</p>
<p>we need to check if the last digit of each number is zero.<p>
<p>If the last digit of each number is zero then the number is  divisible by 10 similarly if the last digit of each number is not zero then the number is not divisible by 10.</p>
<p> Therefore the numbers which are divisible by 10 are</p>
<p>\({@num_div2@}\)</p>
<p>\({@num_div4@}\)</p>
<p>\({@num_div5@}\)</p>
<p>\({@num_div7@}\)</p>
<p>\({@num_div10@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>num_div1:rand_with_step(111,10000,10);/*randomazation with step of 10 from 111 to 10000*/
num_div2:rand_with_step(8000,10000,1000);/*randomazation with step of 1000 from 8000 to 10000*/
num_div3:rand_with_step(999,9999,10);/*randomazation with step of 10 from 999 to 99999*/
num_div4:rand_with_step(100000,599999,10);/*randomazation with step of 10 from 100000 to 599999*/
num_div5:rand_with_step(600000,999999,10);/*randomazation with step of 10 from 600000 to 999999*/
num_div6:rand_with_step(55001,90000,10);/*randomazation with step of 10 from 55001 to 90000*/
num_div7:rand_with_step(20000,70000,10);/*randomazation with step of 10 from 20000 to 70000*/
num_div8:rand_with_step(9999,99999,10);/*randomazation with step of 10 from 9999 to 99999*/
num_div9:rand_with_step(200001,499999,10);/*randomazation with step of 10 from 200001 to 499999*/
num_div10:rand_with_step(10000,20000,10);/*randomazation with step of 10 from 10000 to 20000*/
correct_answer:rand_selection([

[1, true, num_div2],

[2, true, num_div4],[3, true, num_div5],[4, true, num_div7],[5, true, num_div10]

], 3);

incorrect_answers:rand_selection([

[6, false, num_div1],

[7, false, num_div3],

[8, false, num_div6],

[9, false, num_div8],

[10, false, num_div9]

], 2);

ta: random_permutation(append(correct_answer,incorrect_answers));


ta_index: sort(makelist(x[1], x, ta));

ta_largest: last(ta_index);

/*ta:random_permutation([[num_div1,false],[num_div4,true],[num_div3,false], [num_div8,false],[num_div6,false]]);*/
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>({@not_div@}\)</p>
<p>({@ans@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Eric Morara </p>
<p id="concept">https://ecampus.idems.international/question/bank/history/history.php?entryid=126134&returnurl=%2Fquestion%2Fedit.php%3Fcourseid%3D430&courseid=430</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">A divisibility test is a quick way to determine if one number can be divided by another without performing the actual division. These tests rely on certain rules and patterns of numbers. By applying these rules, you can easily check if a number is divisible by a given divisor.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks the numbers divisible by 10</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ans)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>0</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
