<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Numberline</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Illustrate the compound inequality in the line as provided.</p>
<p>{@lower_value@} &le; {@letters@} &le; {@upper_value@}</p>

<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>


[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox', {
            boundingbox: [1, 1, 15, -1],
            axis: true,
            showNavigation: true
        });

        var pointA = board.create('point', [4, 0], {fixed: false, size: 10, face: 'o', strokeColor: 'black', fillColor: 'black'});
        var pointB = board.create('point', [8, 0], {fixed: false, size: 10, face: 'o', strokeColor: 'black', fillColor: 'black'});
        var line = board.create('line', [pointA, pointB], {straightFirst: false, straightLast: false, strokeWidth: 8});

        function updateEquation() {
            var a = Math.round(pointA.X());
            var b = Math.round(pointB.X());
            document.getElementById('equation').innerHTML = `Equation: \\(${#lower_value#} \\leq x \\leq ${#upper_value#}\\)`;
            MathJax.typeset();
        }

        function updateCircleType(point) {
            var circleType = document.getElementById('circleType' + point).value;
            var fillColor = circleType === 'shaded' ? 'black' : 'white';
            if (point === 'A') {
                pointA.setAttribute({fillColor: fillColor});
            } else {
                pointB.setAttribute({fillColor: fillColor});
            }
            board.update();
        }

        pointA.on('drag', updateEquation);
        pointB.on('drag', updateEquation);
[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>letters:rand([x,y,z,m,n,k,q,r,s,t,p,h,b,t,w]);
lower_value:rand(15)-10;
upper_value:rand_with_prohib(10,20,[lower_value]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>ta</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
