<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying the area of a circle in the given figure.</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the area of the circle given below.  \( \text{Use}\,\pi = 3.142\) leaving your answer correct to \(\textbf{One decimal place}\). </p>


<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Document</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/jsxgraph/1.4.0/jsxgraphcore.js"></script>
</head>
<body>
    <div id="circleBox" class="jxgbox" style="width: 400px; height: 400px;"></div>
    <script type="text/javascript">
        var board = JXG.JSXGraph.initBoard('circleBox', {
            boundingbox: [-30, 30, 30, -30],
            axis: false ,showCopyright: false
        });

        var r = {#radii#}; // <<< Change this to any desired radius

        var center = board.create('point', [0, 0], { fixed: true, name: '' });

        // Hidden edge point (used only to define circle radius)
        var edge = board.create('point', [r, 0], {
            visible: false,
            fixed: true ,fixed:true
        });

        var circle = board.create('circle', [center, edge], {
            strokeColor: 'blue'  ,fixed:true
        });

        // Radius line
        var radiusLine = board.create('line', [center, edge], {
            straightFirst: false,
            straightLast: false,
            strokeColor: 'red' ,fixed:true
        });

        // Radius label
        board.create('text', [r / 3, 2, r + " cm"], {
            fontSize: 16 ,fixed:true
        });

        
    </script>



</body>
</html>



<p>Area \(=\)[[input:ans1]]\(cm^2\) [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p> Write the formula for area of a circle</p>
<p>\(\text{Area}=\pi r^2\)</p>
<p>Substitute the given radius and value of π</p>

<p>\(\begin{align}  \text{Area} &={@Pi@} \times  {@radii@}^2\\
&={@Pi@} \times  {@radii@} \times {@radii@} \\
&={@areas@}\,cm^2


\end{align}\)</p>

<p>Therefore;</p>
<p>\(\text{Area}= {@Area@}\,cm^2\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*Area of a circle figure*/
/* Generate random radii between 2 and 20 for 30 students */
radii: rand_with_step(7,27,1)+2;


/* Use π = 3.142 and compute area A = πr² */
Pi: 3.142;
areas:  Pi*radii^2;
/*Computing the area of a circle*/
Area:decimalplaces(areas,1);

/*Mistake 1  Using the wrong formula*/
Areas_1:2*Pi*radii;
Area1:decimalplaces(Areas_1,1);

/*Mistake Rounding incorrectly*/
Areas_2: pi*radii^2;
Areas2:decimalplaces(Areas_2,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{Radius of the circle is}\) {@radii@}</p>
<p> \(\text{Area}=\){@Area@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Henry Onyango</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">The student is required to find the area of the circle drawn and give their answer correct to one decimal place.</p]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Area</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the corect area of the circle drawn.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>Area</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect using of the formular</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>Area1</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You might have used the formula for circumference: \(\text{C}=2 \pi r\) instead of \(\text{A}=\pi r^2\)<p>
<p>\(\textbf{What You Should Do}\):<p>
<p>Carefully read the question and identify whether you're being asked to find area or circumference. For area, the correct formula is:<p>
<p>\(\text{A}=\pi r^2\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for the Rounding incorrectly</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>Areas2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You might have rounded up the answer to two decimal places, which made it incorrect.<p>
<p>\(\textbf{What to do:}\)</p>
<p> Your final answer must be correct to one decimal place. This means you should look at the second decimal digit and decide whether to round up or leave it the same.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.5</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1321420302</deployedseed>
    <deployedseed>360935701</deployedseed>
    <deployedseed>1544311462</deployedseed>
    <deployedseed>1140898296</deployedseed>
    <deployedseed>784721238</deployedseed>
    <deployedseed>1580002427</deployedseed>
    <deployedseed>1256776366</deployedseed>
    <deployedseed>1940156403</deployedseed>
    <deployedseed>1283133040</deployedseed>
    <deployedseed>357745478</deployedseed>
    <deployedseed>1818480997</deployedseed>
    <deployedseed>1302424360</deployedseed>
    <deployedseed>1409188274</deployedseed>
    <deployedseed>370804684</deployedseed>
    <deployedseed>142120895</deployedseed>
    <deployedseed>2034918027</deployedseed>
    <deployedseed>2122260462</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>Area</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Checks for the incorrect using of the formular</description>
      <testinput>
        <name>ans1</name>
        <value>Area1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>checks for the Rounding incorrectly</description>
      <testinput>
        <name>ans1</name>
        <value>Areas2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
