<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the interior angles of a quadrilateral in degrees</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The interior angles of a quadrilateral are \({@angle_1@}\), \({@angle_2@}\), \({@angle_3@}\) \(\text{and}\) \({@angle_4@}.\) Find the angles in degrees.(Write your answer in 2 decimal places)</p>
<p>\({@angle_1@}= \)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>\({@angle_2@}= \)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>\({@angle_3@}= \)[[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>
<p>\({@angle_4@}= \)[[input:ans4]] [[validation:ans4]][[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p>The sum of interior angles of a quadrilateral is \(360^\circ \)</p>
<p>\({@angle_1@}+{@angle_2@}+ {@angle_3@} + {@angle_4@} = \) \(360^\circ \)   </p>
<p>\({@sum_angles@}\) = \(360^\circ \)</p>
<p>\({@letter_1@}\) = \({@angle_value@}^\circ\)</p>


<p>The angles are \({@angle_value@} \times {@number_1@}\),\({@angle_value@} \times {@number_2@}\),\({@angle_value@}\times{@number_3@}\),\({@angle_value@}\times{@number_4@}\)</p>
<p>\(= {@ta_angle_1@},{@ta_angle_2@},{@ta_angle_3@},{@ta_angle_4@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>letter_1:rand([a,q,w,r,t,y,x,g,z,m,n,p,l,z,s,k,h,j,b,c]);/*generates random letters for the angle variable*/
number_1:rand(9)+1;/*generates random number for the coefficient of the first angle */
number_2:rand_with_step(1,8,2);/*generates random number for the coefficient of the second angle */
number_3:rand_with_prohib(1,5,[number_1,number_2]);/*generates random number for the coefficient of the third angle */
number_4:rand_with_prohib(1,7,[number_1,number_2,number_3]);/*generates random number for the coefficient of the fourth angle */
angle_1:letter_1*number_1;/*the first angle*/
angle_2:letter_1*number_2;/*the second angle*/
angle_3:letter_1*number_3;/*the third angle*/
angle_4:letter_1*number_4;/*the fourth angle*/
sum_numbers:number_1+number_2+number_3+number_4;/*sum of the coefficient of the angles*/
sum_angles:angle_1+angle_2+angle_3+angle_4;/*sum of the angles*/
angle_value:decimalplaces(360/sum_numbers,2);/*value for the angle letter*/
ta_angle_1:decimalplaces(angle_value*number_1,2);/*correct answer for the first angle*/
ta_angle_2:decimalplaces(angle_value*number_2,2);/*correct answer for the second angle*/
ta_angle_3:decimalplaces(angle_value*number_3,2);/*correct answer for the third angle*/
ta_angle_4:decimalplaces(angle_value*number_4,2);/*correct answer for the fourth angle*/

/*Mistake_1*/
angle_value_incorrect:decimalplaces(180/sum_numbers,2);/*value for the angle letter when sum of interior angles is 180*/
ta_angle_1_incorrect:decimalplaces(angle_value_incorrect*number_1,2);/*incorrect answer for the first angle*/
ta_angle_2_incorrect:decimalplaces(angle_value_incorrect*number_2,2);/*incorrect answer for the second angle*/
ta_angle_3_incorrect:decimalplaces(angle_value_incorrect*number_3,2);/*incorrect answer for the third angle*/
ta_angle_4_incorrect:decimalplaces(angle_value_incorrect*number_4,2);/*incorrect answer for the fourth angle*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@letter_1@}\)</p>
<p>\({@angle_1@}\)</p>
<p>\({@angle_2@}\)</p>
<p>\({@angle_3@}\)</p>
<p>\({@angle_4@}\)</p>
<p>\({@ta_angle_1@}\)</p>
<p>\({@ta_angle_2@}\)</p>
<p>\({@ta_angle_3@}\)</p>
<p>\({@ta_angle_4@}\)</p>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_angle_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_angle_2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_angle_3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta_angle_4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the first angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta_angle_1</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect answer for the first angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta_angle_1_incorrect</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have equated the sum of interior angles to \(180^\circ\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the second angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta_angle_2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect answer for the second angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta_angle_2_incorrect</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have equated the sum of interior angles to \(180^\circ\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the third angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta_angle_3</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect answer for the third angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta_angle_3_incorrect</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have equated the sum of interior angles to \(180^\circ\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the fourth angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>ta_angle_4</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect answer for the fourth angle in the quadrilateral.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>ta_angle_4_incorrect</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have equated the sum of interior angles to \(180^\circ\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>693016910</deployedseed>
    <deployedseed>1684947985</deployedseed>
    <deployedseed>503468405</deployedseed>
    <deployedseed>1576518766</deployedseed>
    <deployedseed>83756167</deployedseed>
    <deployedseed>1564142719</deployedseed>
    <deployedseed>912990674</deployedseed>
    <deployedseed>1268448380</deployedseed>
    <deployedseed>1046094387</deployedseed>
    <deployedseed>1138403157</deployedseed>
    <deployedseed>1221735825</deployedseed>
    <deployedseed>776892026</deployedseed>
    <deployedseed>275338976</deployedseed>
    <deployedseed>471252168</deployedseed>
    <deployedseed>1973567959</deployedseed>
  </question>
</quiz>
