<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Formulating simple algebraic expressions from the given information.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>\(\text{@persons@}\) buys \(\text{@fruits@}\) from a kiosk. He then buys \({@second_kiosk@}\) more \(\text{@fruits@}\)  from another kiosk.</p>
<p>(i) Use a letter to represent the number of \(\text{@fruits@}\) \(\text{@persons@}\) bought from the first kiosk.
</p>
<p>Number of \(\text{@fruits@} = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>(ii) Write down an algebraic expression for the total number of \(\text{@fruits@}\) \(\text{@persons@}\) bought.</p>
<p>Expression \(=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>(i)  Let the {@fruits@} {@persons@} bought from the first kiosk be {@first_kiosk@}.</p>
<p>(ii) The total number of {@fruits@}  {@persons@} bought is the sum of the {@fruits@}from the first kiosk which is = {@first_kiosk@} and the \({@second_kiosk@}\) additional {@fruits@} he bought from the second kiosk.</P>
<p>The algebraic expression for the total number of {@fruits@} is:</p>
<p> \(\quad \quad =\) {@expression_fruit@} </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[persons:rand(["Karisa","Mellisa","Esther","Jack","Jepkoech","Murunga"]);/*different names of fruit buyers*/
fruits:rand(["mangoes","apples","oranges","pineapples"]);/*different kind of fruits bought*/
first_kiosk:rand([a,b,c,d,e,f,g,h,i,j,k,l,m,n,o,p,q,r,s,t,u,v,w,x,y,z]);/*representation of  fruits bought from the first kiosk*/
second_kiosk:rand(20)+1;/*fruits bought from the second kiosk*/
expression_fruit:first_kiosk+second_kiosk;/*expression of the total fruits bought*/
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>{@first_kiosk@}<p>
<P>{@second_kiosk@}<p>
<p>{@expression_fruit@}<p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author"> Sheila Cherotich</p>
<p id="concept">https://docs.google.com/document/d/16-PneValzoJk_RNdcPGMx6SU24IPdD6mLWoM70lSKz4/edit?tab=t.0#heading=h.4n0o8kgcl43p</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">The question requires the learner to form an algebraic expression to represent the total number of a particular kind of fruit a person buys from two different kiosks.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>first_kiosk</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>expression_fruit</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct representation of fruits bought from the first kiosk.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>first_kiosk</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Mistake 1*/
expression_fruit_mistake1:first_kiosk+second_kiosk = 0;/*equation of the total fruits bought*/
</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct expression for the total number of fruits bought from the first and the second kiosk.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>expression_fruit</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for an equation.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>expression_fruit_mistake1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have written an equation instead of an expression. An expression does not contain an equal sign.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
