<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Forming algebraic expressions involving division</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A water tank has a capacity of \({@tank_capacity@}\) litres. A water drum of capacity \(\text{@drum_capacity@}\)  litres was used to empty the tank when full of water. How many drums were drawn from the tank ?</p>
<p>Number of drums \(=\) [[input:ans1]]  drums[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Total capacity of the tank \(= {@tank_capacity@}\) litres</p>
<p>Capacity of one drum \(= {@drum_capacity@}\) litres</p>
<p>To find the number of drums, divide the total capacity of the tank by the capacity of one drum.</p>
<p>\(\text{ Number of drums} = \frac{\text{Total capacity of tank}}{\text{Capacity of one drum}}\)</p>
<p>\(\qquad \qquad = {@number_of_drums@}\) drums</p>

<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>tank_capacity : rand(3000)+3000; /* capacity of the tank in litres */
drum_capacity : rand([c,d,h,k,b,p,q,g,y]); /* capacity of the drum in litres */
number_of_drums : tank_capacity / drum_capacity; /* Calculate the number of drums */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Tank capacity {@tank_capacity@}</p>
<p>Drum capacity {@drum_capacity@}</p>
<p>Number of drums {@number_of_drums@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">https://docs.google.com/document/d/1RmqGT2I1b183pRZeEUm-N3k-SWWJztCiqxJv9sWTll8/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">To determine how many drums were drawn from the tank, we start by noting that the total capacity of the tank in litres, and the capacity of one drum in litres. To find the number of drums required to empty the tank, we divide the total capacity of the tank by the capacity of one drum.which represents the number of drums needed to empty the tank. ​</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>number_of_drums</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct algebraic expression representing the number of drums.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>number_of_drums</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
