<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the perimeter of  a three quarter-circle.</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[p>A three-quarter circle has a radius OA of \({@radius@}\) cm.calculate the total perimeter (circumference) of the shape.</p>

<!DOCTYPE html>
<html>
<head>
  <title>Sector OAB</title>
  <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraphcore.js"></script>
  <link rel="stylesheet" href="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraph.css" />
</head>
<body>

<div id="box" class="jxgbox" style="width: 500px; height: 500px;"></div>

<script type="text/javascript">
  var board = JXG.JSXGraph.initBoard('box', {
    boundingbox: [-10, 10, 10, -10],
    axis: false,
    showCopyright: false,
    showNavigation: false
  });

  var radius = 8;
  var angleDeg = 270;
  var angleRad = angleDeg * Math.PI / 180;

  // Invisible point O with label
  var O = board.create('point', [0,0], { name: 'O', fixed: true, visible: false, showInfobox: false });
  board.create('text', [O.X(), O.Y() - 0.5, 'O'], { fontSize: 14 });

  // Invisible point A with label
  var A = board.create('point', [radius, 0], { name: 'A', fixed: true, visible: false, showInfobox: false });
  board.create('text', [A.X() + 0.3, A.Y(), 'A'], { fontSize: 14 });

  // Invisible point B with label
  var B = board.create('point', [radius * Math.cos(angleRad), radius * Math.sin(angleRad)], { name: 'B', fixed: true, visible: false, showInfobox: false });
  board.create('text', [B.X(), B.Y() -0.4, 'B'], { fontSize: 14 });

  var OA = board.create('line', [O, A], { straightFirst: false, straightLast: false });
  var OB = board.create('line', [O, B], { straightFirst: false, straightLast: false });

  var arc = board.create('arc', [O, A, B], {
    strokeWidth: 2,
    strokeColor: 'black'
  });

//   // Label the angle AOB = 60°
//   var angle = board.create('nonreflexangle', [A, O, B], {
//     radius: 1.9,
//     name: '60°',
//     label: {
//       fontSize: 12
//     },
//     visible: true
//   });

  // Label line OB = 6 cm
  var OBLabel = board.create('midpoint', [O, B], { visible: false });
  var textOB = board.create('text', [
    () => OBLabel.X(),
    () => OBLabel.Y() + -1.6,
    () => OBLabel.Y() -0.7,
    ' @quarter-circle@m'
  ], { fontSize: 14 });
</script>

</body>
</html>
<p></p><p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]m</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Arc Length = \( \frac{3}{4} \times 2 \pi r\)</p>
<p>Perimeter = Arc Length + 2r</p>
<p>Radius r =\({@radius@}\) cm

Arc length + 2r:
</p>
<p>
\( \frac{3}{4} \times \frac{110}{360} \times 2 \times \({@radius@} + 2  times {@radius@} \)
Total perimeter:
=\({@peri @}\)cm
</p>
<p>To 2dp = \({@ta@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>quarter-circle:rand(210)+7;

 theta_deg: (180) + 20; /* Randomize angle θ in degrees in the range [20°, 200°) */
radius:(1001)+7;/* Randomize radius in degrees in the range [7, 1001] with a step of 7 in each */
/* Convert θ from degrees to fraction of a full circle */
theta_frac: theta_deg / 360;/*Taking theta value selected over 360*/

/* Use pi as 22/7 instead of 3.1416 */
pi_val : 3/4 * 22/7;
two_radii:radius*2;

solution:pi_val*theta_frac*two_radii;/*Arc length*/
/*Mistakes 1*/
/* Compute area of the sector: A = (θ/360) × π × r²  instead pof peimeter formula*/
peri : decimalplaces(theta_frac * pi_val * radius*radius,2);
ta_1:peri;
/*Mistake 2  using the Arc Length Formula Instead of perimeter of a Sector*/

ta_2:theta_frac*2*pi_val*radius;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>peri1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>perimeter of a quarter-</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>peri1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
