<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Representing of integers on a number line.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Represent \(({@int_number@})\) on the number line below.</p>
<p>Drag the green point along the number line to the correct integer.</p>
[[jsxgraph input-ref-ans1='ans1Ref']]
    var board = JXG.JSXGraph.initBoard(divid, {
         boundingbox: [-10, 10, 12, -10],  
        keepaspectratio: true, 
        showCopyright: true, 
        axis: false,
        showNavigation:true

    });
    const axis = board.create('line', [[-8.5, 0], [10.5, 0]], { 
        straightFirst: false,
        straightLast: false,
        strokeWidth: 2,
        strokeColor: 'black',
        firstArrow:true,
        lastArrow:true,
        fixed: true
    });
    for (let x = -8; x <= 10; x++) {
        board.create('line', [[x, 0.25], [x, -0.25]], {
            strokeWidth: 1,
            strokeColor: 'black',
            straightFirst: false,
            straightLast: false,
            fixed: true
        });
        board.create('text', [x, -0.35, x.toString()], {
            anchorX: 'middle',
            anchorY: 'top',
            fontSize: 12,
            fixed: true
        });
    }
 var point1 = board.create('point', [{#circle_false_pos#}, 0], {
        name: '',
        size: 5,
        fixed: false, 
        strokeColor: 'green',
        fillColor: 'green', 
        attractors: [axis],
        attractorDistance: 1000.0,
        snapToGrid: true,
        snapSizeX: 1,
        snapSizeY: 1
    });
stack_jxg.bind_point(ans1Ref, point1);
board.update();

[[/jsxgraph]]
<p style='display:none'>[[input:ans1]] [[validation:ans1]]</p> [[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
[[jsxgraph input-ref-ans1='ans1Ref']]
    var board = JXG.JSXGraph.initBoard(divid, {
        boundingbox: [-10, 10, 12, -10],  
        keepaspectratio: true, 
        showCopyright: true, 
        axis: false,
        showNavigation:true
    });
    const axis = board.create('line', [[-8.5, 0], [10.5, 0]], { 
        straightFirst: false,
        straightLast: false,
        strokeWidth: 2,
        strokeColor: 'black',
        firstArrow:true,
        lastArrow:true,
        fixed: true
    });
    for (let x = -8; x <= 10; x++) {
        board.create('line', [[x, 0.25], [x, -0.25]], {
            strokeWidth: 1,
            strokeColor: 'black',
            straightFirst: false,
            straightLast: false,
            fixed: true
        });
        board.create('text', [x, -0.35, x.toString()], {
            anchorX: 'middle',
            anchorY: 'top',
            fontSize: 12,
            fixed: true
        });
    }
 var point1 = board.create('point', [{#int_number#}, 0], {
        name: '',
        size: 5,
        fixed: false, 
        strokeColor: 'green',
        fillColor: 'green', 
        attractors: [axis],
        attractorDistance: 1000.0,
        snapToGrid: true,
        snapSizeX: 1,
        snapSizeY: 1
    });
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>int_number:rand_with_step(-10,10,1);
circle_false_pos:rand_with_prohib(-10,10,[int_number]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Integer: \({@int_number@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id= "author">Daniel Murunga</p>
<p id= "concept"></p>
<p id= "reviewer"></p>
<p id= "reviewer"></p>
<p id="description">This question requires the student to plot given integers on the number line.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>int_number</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks if the correct integer is plotted on the number line.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1[1]</sans>
        <tans>int_number</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
