<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the surface area of the cube from its net.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The figure below shows a cube. Sketch the net of the cube and work out the surface area of the cube from its net.</p>
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Wireframe Cube</title>
</head>
<body>
  <svg width="375" height="375" viewBox="0 0 375 375">
    <!-- Front face -->
    <polygon points="90,90 202.5,90 202.5,202.5 90,202.5" stroke="blue" fill="none" />
    <!-- Back face -->
    <polygon points="135,45 247.5,45 247.5,157.5 135,157.5" stroke="blue" fill="none" />
    <!-- Connect corners -->
    <line x1="90" y1="90" x2="135" y2="45" stroke="blue" />
    <line x1="202.5" y1="90" x2="247.5" y2="45" stroke="blue" />
    <line x1="202.5" y1="202.5" x2="247.5" y2="157.5" stroke="blue" />
    <line x1="90" y1="202.5" x2="135" y2="157.5" stroke="blue" />

    <!-- Vertex labels (invisible) -->
    <text x="90" y="90" font-family="Arial" font-size="18" fill="none">A</text>
    <text x="202.5" y="90" font-family="Arial" font-size="18" fill="none">B</text>
    <text x="202.5" y="202.5" font-family="Arial" font-size="18" fill="none">C</text>
    <text x="90" y="202.5" font-family="Arial" font-size="18" fill="none">D</text>
    <text x="135" y="45" font-family="Arial" font-size="18" fill="none">E</text>
    <text x="247.5" y="45" font-family="Arial" font-size="18" fill="none">F</text>
    <text x="247.5" y="157.5" font-family="Arial" font-size="18" fill="none">G</text>
    <text x="135" y="157.5" font-family="Arial" font-size="18" fill="none">H</text>

    <!-- Midpoints (circles hidden), labels visible -->

    <!-- "3cm" for midpoint of DC -->
    <circle cx="146.25" cy="202.5" r="4.5" fill="none" />
    <text x="150" y="222" font-family="Arial" font-size="18" fill="black">{#side#} cm</text>

    <!-- "1cm" for midpoint of CG -->
    <circle cx="225" cy="180" r="4.5" fill="none" />
    <text x="229" y="198" font-family="Arial" font-size="18" fill="black">{#side#} cm</text>

    <!-- "2cm" for midpoint of GF -->
    <circle cx="247.5" cy="101.25" r="4.5" fill="none" />
    <text x="252" y="120" font-family="Arial" font-size="18" fill="black">{#side#} cm</text>
  </svg>
</body>
</html>
<p>Surface area=[[input:ans1]] cm\(^2\) [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>(a).The figure below shows the net of the cube</p>
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Cube Net with Selected Dimensions</title>
  <style>
    .net {
      display: grid;
      grid-template-columns: repeat(3, 80px);
      grid-template-rows: repeat(5, 80px);
      gap: 2px;
      justify-content: center;
      align-items: center;
      margin-top: 40px;
    }

    .cell {
      position: relative;
      width: 80px;
      height: 80px;
      display: flex;
      align-items: center;
      justify-content: center;
      font-weight: bold;
      border: 1px solid #333;
      box-sizing: border-box;
    }

    .A { background-color: yellow; grid-column: 2; grid-row: 1; }
    .B { background-color: blue; grid-column: 2; grid-row: 2; }
    .C { background-color: orange; grid-column: 2; grid-row: 3; }
    .D { background-color: yellow; grid-column: 2; grid-row: 4; }
    .E { background-color: cyan; grid-column: 1; grid-row: 3; }
    .F { background-color: purple; grid-column: 3; grid-row: 3; }

    .dimension {
      position: absolute;
      font-size: 10px;
      background: white;
      padding: 1px 4px;
      border: 1px solid #999;
      border-radius: 4px;
    }

    .top { top: -24px; left: 50%; transform: translateX(-50%); }
    .bottom { bottom: -24px; left: 50%; transform: translateX(-50%); }
    .left { left: -40px; top: 50%; transform: translateY(-50%); }
    .right { right: -40px; top: 50%; transform: translateY(-50%); }
  </style>
</head>
<body>

  <div class="net">
    <div class="cell A">
      A
      <span class="dimension top">{#side#} cm</span>
      <span class="dimension left">{#side#} cm</span>
      <span class="dimension right">{#side#} cm</span>
    </div>
    <div class="cell B">
      B
      
      <span class="dimension bottom">{#side#} cm</span>
      <span class="dimension left">{#side#} cm</span>
      <span class="dimension right">{#side#} cm</span>
    </div>
    <div class="cell C">
      C
      <!-- Dimensions fully removed from square C -->
    </div>
    <div class="cell D">
      D
      <span class="dimension bottom">{#side#} cm</span>
      <span class="dimension right">{#side#} cm</span>
    </div>
    <div class="cell E">
      E
    
      <span class="dimension bottom">{#side#} cm</span>
      <span class="dimension left">{#side#} cm</span>
     
    </div>
    <div class="cell F">
      F
     
      <span class="dimension bottom">{#side#} cm</span>
      
      <span class="dimension right">{#side#} cm</span>
    </div>
  </div>

</body>
</html>

<p>(b).Calculating  the surface area.</p>
<p>\(\begin{align}\text{The area of one face}&=length \times width\\&={@side@} cm\times {@side@} cm\\&={@area_of_squareA@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{There are 6 faces. Therefore, the surface area of the cube}&=6 \times \text{area of one face}\\&=6 \times {@area_of_squareA@} cm^2\\&={@ta@} cm^2\end{align}\)</p>
<p><b>Alternative</b></p>
<p>\(\begin{align}\text{Area of square A}&=length \times width\\&={@side@} cm\times {@side@} cm\\&={@area_of_squareA@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Area of square B}&=length \times width\\&={@side@} cm\times {@side@} cm\\&={@area_of_squareA@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Area of square C}&=length \times width\\&={@side@} cm\times {@side@} cm\\&={@area_of_squareA@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Area of square D}&=length \times width\\&={@side@} cm\times {@side@} cm\\&={@area_of_squareA@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Area of square E}&=length \times width\\&={@side@} cm\times {@side@} cm\\&={@area_of_squareA@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Area of square F}&=length \times width\\&={@side@} cm\times {@side@} cm\\&={@area_of_squareA@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Surface area of the cube}&={@area_of_squareA@} cm^2+{@area_of_squareB@} cm^2+{@area_of_squareC@} cm^2+{@area_of_squareD@} cm^2+{@area_of_squareE@} cm^2+{@area_of_squareF@} cm^2\\&={@ta@} cm^2\end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/*surface area of a cube*/
side:rand_with_step(5,99,1);/*dimension of the side of the cube*/
area_of_squareA:side*side;/*area of  square A*/
area_of_squareB:side*side;/*area of  square B*/
area_of_squareC:side*side;/*area of  square C*/
area_of_squareD:side*side;/*area of  square D*/
area_of_squareE:side*side;/*area of  square E*/
area_of_squareF:side*side;/*area of  square F*/
surface_area:area_of_squareA*6;/*total surface area of the cube from its net*/
ta:surface_area;/*teacher's answer*/

/*mistakes*/
ta_1:area_of_squareA*5;/*not including all faces*/
perimeter:(side+side+side+side)*6;
ta_2:perimeter;
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Surface area of the cube =6 \times area of one face={@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct surface area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description><![CDATA[<p> Adding Instead of Multiplying</p>]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you found the perimeter instead of the area. Remember, perimeter measures the distance around the edge, but surface area is all about the space inside each face. Try squaring the side length instead — you’re close, just a small switch in the formula will get you there</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>0</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description><![CDATA[<p>Not including all faces</p>]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You’re almost there, it looks like you missed one of the cube’s faces. A cube has 6 equal square faces, so make sure all of them are counted. Double-check your net and give it another try</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1787018075</deployedseed>
    <deployedseed>2062995598</deployedseed>
    <deployedseed>1601273657</deployedseed>
    <deployedseed>350542555</deployedseed>
    <deployedseed>777690842</deployedseed>
    <deployedseed>1803501000</deployedseed>
    <deployedseed>228221649</deployedseed>
    <deployedseed>2041729038</deployedseed>
    <deployedseed>1271633358</deployedseed>
    <deployedseed>1396585628</deployedseed>
    <deployedseed>1994389680</deployedseed>
    <deployedseed>287356643</deployedseed>
    <deployedseed>363660712</deployedseed>
    <deployedseed>292131946</deployedseed>
    <deployedseed>1290839965</deployedseed>
    <deployedseed>889988027</deployedseed>
    <deployedseed>1546515111</deployedseed>
    <deployedseed>709814206</deployedseed>
    <deployedseed>252827283</deployedseed>
    <deployedseed>2126261550</deployedseed>
    <deployedseed>1426889628</deployedseed>
    <deployedseed>1440502928</deployedseed>
    <deployedseed>1047909913</deployedseed>
    <deployedseed>425776401</deployedseed>
    <deployedseed>340656734</deployedseed>
    <deployedseed>1464225579</deployedseed>
    <deployedseed>505480315</deployedseed>
    <deployedseed>2107764411</deployedseed>
    <deployedseed>262196927</deployedseed>
    <deployedseed>1331015599</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>correct surface area</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description><![CDATA[<p> Adding Instead of Multiplying</p>]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description><![CDATA[<p>Not including all faces</p>]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
