<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the surface area of a cylinder from its net.</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the surface area of the cylinder below, with diameter {@diameter@} cm and perpendicular height {@height@} cm from its net.</p>

<!DOCTYPE html>
<html>
<head>
  <title>Proper Cylinder - JSXGraph</title>
  <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraphcore.js"></script>
  <link rel="stylesheet" type="text/css" href="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraph.css" />
  <style>
    #jxgbox {
      width: 600px;
      height: 500px;
    }
  </style>
</head>
<body>

    <div id="jxgbox" class="jxgbox" style="width: 500px; height: 400px;"></div>
    <script type="text/javascript">
        var board = JXG.JSXGraph.initBoard('jxgbox', {
            boundingbox: [-10, 100, 30, -10],
            axis: false,showCopyright: false
        });

        var radius = 10.5; // Half of 21 cm diameter
        var height = 60;

        // Bottom ellipse
        var bottom = board.create('ellipse', [
            [0, 0],
            [radius, 0],
            [0, 3]
        ], {strokeColor: '#000', fillColor: '#ccc', fillOpacity: 0.3 ,fixed:true});

        // Top ellipse
        var top = board.create('ellipse', [
            [0, height],
            [radius, height],
            [0, height + 3]
        ], {strokeColor: '#000', fillColor: '#ccc', fillOpacity: 0.3,fixed:true});

        // Side lines
        var leftLine = board.create('line', [[-1.7, -0.1], [-1.7, 60]], {straightFirst: false, straightLast: false,fixed:true});
        var rightLine = board.create('line', [[12.2, 0], [12.2, 60]], {straightFirst: false, straightLast: false,fixed:true});

        // Labels
        var heightLabel = board.create('text', [radius + 2, height / 2, '{#height#} cm'],{fixed:true});
        var diameterLabel = board.create('text', [5, 62, ' {#diameter#} cm'],{fixed:true});


        var pointA = board.create('point', [-1.7, 60], {
        name: 'A',
        visible: false,fixed:false
    });

    var pointB = board.create('point', [12.2, 60], {
        name: 'B',
        visible: false ,fixed:true
    });

    var line = board.create('segment', [pointA, pointB], {
        strokeColor: 'blue',fixed:true
    });
    </script>
</body>
</html>

<p>Surface area=[[input:ans1]]cm\(^2\) [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>The figure below shows the net of the cylinder</p>
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Net of a Cylinder</title>
</head>
<body>
  <svg width="300" height="400" xmlns="http://www.w3.org/2000/svg">
    <!-- Top Circle -->
    <circle cx="150" cy="60" r="49" stroke="black" fill="cyan" />

    <!-- Bottom Circle -->
    <circle cx="150" cy="300" r="49" stroke="black" fill="cyan" />

    <!-- Rectangle -->
    <rect x="28" y="110" width="244" height="140" stroke="black" fill="lightgreen" />

    <!-- Top Circle Arrow and Label -->
    <line x1="150" y1="11" x2="150" y2="60" stroke="black" marker-end="url(#arrow)" />
    <text x="155" y="40" font-size="12">{#radius#} cm</text>

    <!-- Bottom Circle Radius from Center -->
    <line x1="150" y1="350" x2="150" y2="300" stroke="black" marker-end="url(#arrow)" />
    <text x="155" y="325" font-size="12">{#radius#} cm</text>

    <!-- Rectangle Dimension Labels -->
    <text x="140" y="125" font-size="12">{#length#} cm</text>
    <text x="275" y="180" font-size="12" transform="rotate(90,275,180)">{#height#} cm</text>

    <!-- Points: A, B, C, D -->
    <text x="270" y="105" font-size="12">A</text>
    <text x="270" y="255" font-size="12">B</text>
    <text x="10" y="255" font-size="12">C</text>
    <text x="10" y="105" font-size="12">D</text>

    <!-- Arrow Marker Definition -->
    <defs>
      <marker id="arrow" markerWidth="10" markerHeight="10" refX="5" refY="5" orient="auto">
        <path d="M 0 0 L 10 5 L 0 10 z" fill="black"/>
      </marker>
    </defs>
  </svg>
</body>
</html>

<p><b>Surface area</b></p>
<p>\(\begin{align}\text{Circumference of a circle}=\text{Length of the rectangle}\end{align}\)</p>
<p>\(\begin{align}\text{Radius}&=\frac{Diameter}{2}\\&=\frac{{@diameter@}}{2}\\&={@radius@} cm\end{align}\)</p>
<p>\(\begin{align}\text{Circumference of a circle}&=2 \pi  r\\&=2 \times  \frac{22}{7}\times {@radius@} cm\\&={@length@} cm\end{align}\)</p>
<p>\(\begin{align}\text{Length}&={@length@} cm\end{align}\)</p>
<p>\(\begin{align}\text{Width}&={@height@} cm\end{align}\)</p>
<p>\(\begin{align}\text{Area of the rectangle}&=Length \times width\\&={@length@} cm \times {@height@} cm\\&={@area_of_rectangle@} cm^2\end{align}\)
<p>\(\begin{align}\text{Area of the top circle}&=\pi r^2\\&=\frac{22}{7} \times {@radius@} cm\times{@radius@} cm\\&={@area_of_circle@} cm^2\end{align}\)</p>
<p>The area of the top and the area of the bottom circle of a cylinder are the same.</p>
<p>\(\begin{align}\text{Area of the bottom circle}&=\pi r^2\\&=\frac{22}{7} \times {@radius@} cm\times{@radius@} cm\\&={@area_of_circle@} cm^2\end{align}\)</p>
<p>\(\begin{align}\text{Total surface area}&={@area_of_rectangle@} cm^2+{@area_of_circle@} cm^2+{@area_of_circle@} cm^2\\&={@ta@} cm^2\end{align}\)<p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>radius:rand_with_step(7,280,7);/*radius of the circular part*/
diameter:2*radius;/*diameter of the circle*/
height:rand(36)+5;/*height of the cylinder*/
length:2*22/7*radius;/*circumference which gives the length of the rectangle when cylinder is opened to form its net*/
area_of_circle:22/7*radius*radius;/*the area of a circle*/
area_of_rectangle:length*height;/*the area of a rectangle*/
surface_area:area_of_circle+area_of_circle+area_of_rectangle;/*total surface area*/
ta:surface_area;/*teacher's answer*/

/*mistakes*/
ta_1:area_of_circle+area_of_rectangle;
area_rectangle_1:2*radius*height;
ta_2:area_rectangle_1+area_of_circle+area_of_circle;
area_of_circle1:22/7*diameter*diameter;
ta_3:area_of_circle1+area_of_circle1+area_of_rectangle;
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Circumference of the circle (\(2\pi\times radius\))is equal to the length of the rectangle={@length@}</p>
<p>Width of the rectangle is equal to the height of the cylinder ={@height@}</p>
<p>Surface area={@ta@}<p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Eric Morara</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The student is required to find the surface area of the cylinder from its net.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for the correct surface area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check for incorrect working of the surface area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you only included one of the circles when calculating the surface area. Remember, a cylinder has two circular faces — one at the top and one at the bottom — and both need to be counted. You're on the right track, just make sure to include both circles next time. Give it another try — you’re almost there!</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>0</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>check for incorrect working of the surface area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you mixed up the dimensions of the rectangle when working out the surface area. Remember, the length of the rectangle should match the circumference of the circle, not the diameter or radius. You're making good progress — double-check the dimensions and give it another try!</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>check for incorrect working of the surface area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you used the wrong formula for the area of a circle. Remember, the correct formula is \(\pi r^2\)
 , where r is the radius of the circle. You're definitely on the right path — just be careful with the formula and try again!<p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1272167702</deployedseed>
    <deployedseed>488460472</deployedseed>
    <deployedseed>332085674</deployedseed>
    <deployedseed>1002698102</deployedseed>
    <deployedseed>1902264411</deployedseed>
    <deployedseed>1277483144</deployedseed>
    <deployedseed>1586201424</deployedseed>
    <deployedseed>590663201</deployedseed>
    <deployedseed>361380415</deployedseed>
    <deployedseed>1202926783</deployedseed>
    <deployedseed>581128501</deployedseed>
    <deployedseed>1392748729</deployedseed>
    <deployedseed>1226912170</deployedseed>
    <deployedseed>971526969</deployedseed>
    <deployedseed>1201072641</deployedseed>
    <deployedseed>322246920</deployedseed>
    <deployedseed>1487161194</deployedseed>
    <deployedseed>964090773</deployedseed>
    <deployedseed>694250684</deployedseed>
    <deployedseed>268081238</deployedseed>
    <deployedseed>226027682</deployedseed>
    <deployedseed>371989812</deployedseed>
    <deployedseed>311714335</deployedseed>
    <deployedseed>1203884129</deployedseed>
    <deployedseed>83565417</deployedseed>
    <deployedseed>1487579460</deployedseed>
    <deployedseed>5610670</deployedseed>
    <deployedseed>730612516</deployedseed>
    <deployedseed>1927091979</deployedseed>
    <deployedseed>994011604</deployedseed>
    <deployedseed>2026699276</deployedseed>
    <deployedseed>1437724924</deployedseed>
    <deployedseed>1029596074</deployedseed>
    <deployedseed>413317107</deployedseed>
    <deployedseed>795063920</deployedseed>
    <deployedseed>913649519</deployedseed>
    <deployedseed>423570798</deployedseed>
    <deployedseed>552210986</deployedseed>
    <deployedseed>527561005</deployedseed>
    <deployedseed>555057075</deployedseed>
    <deployedseed>2145044038</deployedseed>
    <deployedseed>1143229759</deployedseed>
    <deployedseed>1538952453</deployedseed>
    <deployedseed>668152436</deployedseed>
    <deployedseed>1546152757</deployedseed>
    <deployedseed>770298381</deployedseed>
    <deployedseed>1827769277</deployedseed>
    <deployedseed>476338429</deployedseed>
    <deployedseed>1489417179</deployedseed>
    <deployedseed>302294800</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description><![CDATA[<p>It looks like you only included one of the circles when calculating the surface area. Remember, a cylinder has two circular faces — one at the top and one at the bottom — and both need to be counted. You're on the right track, just make sure to inc]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>check for the correct surface area</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description><![CDATA[<p>It looks like you mixed up the dimensions of the rectangle when working out the surface area. Remember, the length of the rectangle should match the circumference of the circle, not the diameter or radius. You're making good progress — double-check t]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description><![CDATA[<p>It looks like you used the wrong formula for the area of a circle. Remember, the correct formula is \(\pi r^2\)  , where r is the radius of the circle. You're definitely on the right path — just be careful with the formula and try again!<p>]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta_3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
