<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Applying linear inequality symbols to inequality statements.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Compare the following numbers using the greater than or less than symbols.</p>
<p>(a) \({@digit_1@}\) and \({@digit_2@} = \)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>(b) \({@digit_3@}\) and \({@digit_4@} = \)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Note that:</p>
<p>The inequality sign for greater than is >.</p>
<p>The inequality sign for less than is <.</p>
<p>Therefore;</p>
<p>(a) \({@digit_1@}\) is less than \({@digit_2@}\) is written as \({@digit_1@}\) < \({@digit_2@} \) </p>
<p>(b) \({@digit_3@}\) is greater than \({@digit_4@}\) is written as \({@digit_3@}\) > \({@digit_4@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[digit_1:rand(20)+1;/*generates the first digit for part (a)*/
digit_2:rand_with_step(21,50,2);/*generates the second digit for part (a)*/
digit_3:rand_with_prohib(15,30,[digit_1]);/*generates the first digit for part (b)*/
digit_4:rand_with_prohib(1,14,[digit_1]);/*generates the second digit for part (b)*/
ta_a:digit_1<digit_2;/*correct answer for part (a)*/
ta_b:digit_3>digit_4;/*correct answer for part (b)*/


]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@digit_1@}\)</p> 
<p>\({@digit_2@}\)</p> 
<p>\({@digit_3@}\)</p> 
<p> \({@digit_4@}\)</p>
<p>\({@ta_a@}\)</p> 
<p>\({@ta_b@}\)</p> 

]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author"> Isdora Ochieng</p>
<p id="concept">https://docs.google.com/document/d/1xx5JVAEoItVnWWsOer9xJCbKgVo49ATJ5zFDtqdiCNg/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">The question requires the learner to compare pairs of numbers using less than or greater than symbols to inequality statements</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_a</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_b</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/*Mistake 1*/
ta_a_incorrect:digit_1>digit_2;/*incorrect answer for part (a) when greater than sign is used in place of less than sign*/]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for part (a)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_a</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for part (b)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>553419200</deployedseed>
    <deployedseed>421137894</deployedseed>
    <deployedseed>733705004</deployedseed>
    <deployedseed>1957620850</deployedseed>
    <deployedseed>704158869</deployedseed>
    <deployedseed>1913294306</deployedseed>
    <deployedseed>304565355</deployedseed>
    <deployedseed>443918656</deployedseed>
    <deployedseed>1865837396</deployedseed>
    <deployedseed>1958968788</deployedseed>
    <deployedseed>465810401</deployedseed>
    <deployedseed>1241109021</deployedseed>
    <deployedseed>1257437681</deployedseed>
    <deployedseed>1699446742</deployedseed>
    <deployedseed>1432517628</deployedseed>
  </question>
</quiz>
