<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Representing the number line using an inequality symbol.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Represent the number line using an inequality symbol.  [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
[[jsxgraph]]
   	const board = JXG.JSXGraph.initBoard('jxgbox', {
        	boundingbox: [-6, 0.5, 11, -0.5],  
        	axis: false,
        	showNavigation: false,
        	keepAspectRatio:true,
        	showCopyright:false
    	});
    	const axis = board.create('line', [[-5.5, 0], [10.5, 0]], {
        	straightFirst: false,
        	straightLast: false,
        	strokeWidth: 2,
        	strokeColor: 'black'
    	});

    	board.create('line', [[-5.5, 0], [-5.2, 0.1]], { strokeColor: 'black', strokeWidth: 1 ,straightFirst: false,straightLast: false});
    	board.create('line', [[-5.5, 0], [-5.2, -0.1]], { strokeColor: 'black', strokeWidth: 1 ,straightFirst: false,straightLast: false});
    	board.create('line', [[10.5, 0], [10.2, 0.1]], { strokeColor: 'black', strokeWidth: 1,straightFirst: false,straightLast: false });
    	board.create('line', [[10.5, 0], [10.2, -0.1]], { strokeColor: 'black', strokeWidth: 1,straightFirst: false,straightLast: false });

    	for (let x = -5; x <= 10; x++) {
        	board.create('line', [[x, 0.25], [x, -0.25]], {
            	strokeWidth: 1,
            	strokeColor: 'black',
            	straightFirst: false,
            	straightLast: false
        	});
        	board.create('text', [x, -0.35, x.toString()], {
            	anchorX: 'middle',
            	anchorY: 'top',
            	fontSize: 12
        	});
    	}
    	const rangeLine = board.create('line', [[{#n2#}, 0.25], [{#n1-0.23#}, 0.25]], {
        	straightFirst: false,
        	straightLast: false,
        	strokeWidth: 1,
        	strokeColor: 'green',
        	firstArrow: true
    	});

    	const openCircle = board.create('circle', [[{#n1#}, 0.25], 0.2], {
        	strokeWidth: 1,
        	strokeColor: 'green',
        	fillColor: 'none'
    	});


[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To create the inequality, we first need to determine the type of point used on the number line: open or closed. </p>
<p>The given number line shows an empty circle, indicating an open point. Open points signify that the number itself is not part of the solution set.</p>
<p>Therefore, since the open circle is at \({@n1@}\) and the line extends to the left (representing values less than \({@n1@}\)), the inequality is \({@ta@}\).
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Generate random numbers */
n1: rand_with_step(-4, 9, 1);
n2: rand_with_step(-4, 9, 1);

/* Ensure number1 is always greater than number2 */
if n2 > n1 then (
	k: n1,
	n1: n2,
	n2: k
)elseif n1 = n2 then (
	n1: n1 + 1 );

/* Define the inequality */
ta: x < n1;
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">https://docs.google.com/document/d/1HpwUhk2L2qg6bmcmiCFVbKT1zCxyvWYx5xJ_V3-7lBk/edit?tab=t.0</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">To represent the number line using an inequality symbol, first identify the point where the arrow starts or stops. Check whether this point has an open circle, which means the value is excluded, or a closed circle, which means the value is included. Then, observe the direction of the arrow to determine if the values are less than or greater than the point. Use <, >, \(\leq\), or \(\geq \)based on whether the boundary point is included or excluded and the direction of the arrow. Finally, write the inequality to describe the values represented on the number line​.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
